Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f1540,plain,(
  $false),
  inference(subsumption_resolution,[],[f1102,f1524])).
fof(f1524,plain,(
  cont(sK2,sK0)),
  inference(backtracking_split_refutation,[],[f1016,f1024_D,f1517])).
fof(f1517,plain,(
  $false | $spl11),
  inference(subsumption_resolution,[],[f1516,f300])).
fof(f300,plain,(
  cont(sK2,sum(sK0,sK1))),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  cont(sK2,sum(sK0,sK1)) & ! [X3] : ((~p(X3,sK2) | ~cont(X3,sK0)) & (~p(X3,sK2) | ~cont(X3,sK1)))),
  inference(skolemisation,[status(esa)],[f166])).
fof(f166,plain,(
  ? [X0,X1,X2] : (cont(X2,sum(X0,X1)) & ! [X3] : ((~p(X3,X2) | ~cont(X3,X0)) & (~p(X3,X2) | ~cont(X3,X1))))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ~! [X0,X1,X2] : (cont(X2,sum(X0,X1)) => ? [X3] : ((p(X3,X2) & cont(X3,X0)) | (p(X3,X2) & cont(X3,X1))))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : ! [X2] : (cont(X2,sum(X0,X1)) => ? [X3] : ((p(X3,X2) & cont(X3,X0)) | (p(X3,X2) & cont(X3,X1))))),
  file(unknown,unknown)).
fof(f1516,plain,(
  ~cont(sK2,sum(sK0,sK1)) | $spl11),
  inference(subsumption_resolution,[],[f1496,f1024])).
fof(f1496,plain,(
  cont(sK2,sK0) | ~cont(sK2,sum(sK0,sK1)) | $spl11),
  inference(resolution,[],[f406,f1048])).
fof(f1048,plain,(
  ~cont(difference(sK2,sK0),sK1) | $spl11),
  inference(resolution,[],[f1032,f302])).
fof(f302,plain,(
  ( ! [X3] : (~p(X3,sK2) | ~cont(X3,sK1)) )),
  inference(cnf_transformation,[],[f237])).
fof(f1032,plain,(
  p(difference(sK2,sK0),sK2) | $spl11),
  inference(resolution,[],[f1025,f347])).
fof(f347,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X117] : ! [X118] : (~zex(difference(X117,X118)) => p(difference(X117,X118),X117))),
  file(unknown,unknown)).
fof(f1025,plain,(
  ~zex(difference(sK2,sK0)) | $spl11),
  inference(resolution,[],[f1024,f711])).
fof(f711,plain,(
  ( ! [X1] : (cont(sK2,X1) | ~zex(difference(sK2,X1))) )),
  inference(resolution,[],[f391,f433])).
fof(f433,plain,(
  ~zex(sK2)),
  inference(resolution,[],[f320,f300])).
fof(f320,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X142] : ! [X143] : (zex(X142) => (~cont(X143,X142) & ~cont(X142,X143)))),
  file(unknown,unknown)).
fof(f391,plain,(
  ( ! [X0,X1] : (zex(X0) | cont(X0,X1) | ~zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X27] : ! [X28] : (zex(difference(X27,X28)) <=> (zex(X27) | cont(X27,X28)))),
  file(unknown,unknown)).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : (~cont(X2,sum(X0,X1)) | cont(X2,X0) | cont(difference(X2,X0),X1))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((~cont(X2,sum(X0,X1)) | cont(X2,X0)) | cont(difference(X2,X0),X1))),
  inference(ennf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X11] : ! [X12] : ! [X13] : ((cont(X13,sum(X11,X12)) & ~cont(X13,X11)) => cont(difference(X13,X11),X12))),
  file(unknown,unknown)).
fof(f1024,plain,(
  ~cont(sK2,sK0) | $spl11),
  inference(cnf_transformation,[],[f1024_D])).
fof(f1024_D,plain,(
  ~cont(sK2,sK0) <=> ~$spl11),
  introduced(backtracking_splitting_component,[])).
fof(f1016,plain,(
  ~po(sK2,sK13(sK2,sK2)) | ~cont(sK2,sK0)),
  inference(resolution,[],[f902,f301])).
fof(f301,plain,(
  ( ! [X3] : (~p(X3,sK2) | ~cont(X3,sK0)) )),
  inference(cnf_transformation,[],[f237])).
fof(f902,plain,(
  ( ! [X4] : (p(sK2,X4) | ~po(X4,sK13(X4,sK2))) )),
  inference(resolution,[],[f344,f433])).
fof(f344,plain,(
  ( ! [X0,X1] : (zex(X0) | ~po(X1,sK13(X1,X0)) | p(X0,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1] : (zex(X0) | (po(X0,sK13(X1,X0)) & ~po(X1,sK13(X1,X0))) | p(X0,X1))),
  inference(skolemisation,[status(esa)],[f200])).
fof(f200,plain,(
  ! [X0,X1] : (zex(X0) | ? [X2] : (po(X0,X2) & ~po(X1,X2)) | p(X0,X1))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ! [X0,X1] : ((zex(X0) | ? [X2] : (po(X0,X2) & ~po(X1,X2))) | p(X0,X1))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : ((~zex(X0) & ! [X2] : (po(X0,X2) => po(X1,X2))) => p(X0,X1))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1] : ((~zex(X0) & ! [X2] : (po(X0,X2) => po(X1,X2))) => p(X0,X1))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X29] : ! [X30] : ((~zex(X29) & ! [X31] : (po(X29,X31) => po(X30,X31))) => p(X29,X30))),
  file(unknown,unknown)).
fof(f1102,plain,(
  ~cont(sK2,sK0)),
  inference(resolution,[],[f1096,f301])).
fof(f1096,plain,(
  p(sK2,sK2)),
  inference(backtracking_split_refutation,[],[f1003,f1020_D,f1094])).
fof(f1094,plain,(
  $false | $spl7),
  inference(subsumption_resolution,[],[f1092,f1020])).
fof(f1092,plain,(
  p(sK2,sK2) | $spl7),
  inference(resolution,[],[f1027,f884])).
fof(f884,plain,(
  ( ! [X4] : (po(sK2,sK13(X4,sK2)) | p(sK2,X4)) )),
  inference(resolution,[],[f343,f433])).
fof(f343,plain,(
  ( ! [X0,X1] : (zex(X0) | po(X0,sK13(X1,X0)) | p(X0,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f1027,plain,(
  ~po(sK2,sK13(sK2,sK2)) | $spl7),
  inference(resolution,[],[f1020,f902])).
fof(f1020,plain,(
  ~p(sK2,sK2) | $spl7),
  inference(cnf_transformation,[],[f1020_D])).
fof(f1020_D,plain,(
  ~p(sK2,sK2) <=> ~$spl7),
  introduced(backtracking_splitting_component,[])).
fof(f1003,plain,(
  ( ! [X1] : (~po(X1,sK13(X1,sK2)) | ~p(sK2,sK2) | ~zex(sK22(X1,sK2))) )),
  inference(resolution,[],[f902,f661])).
fof(f661,plain,(
  ( ! [X26,X25] : (~p(X25,sK2) | ~p(X25,X26) | ~zex(sK22(X26,sK2))) )),
  inference(resolution,[],[f387,f552])).
fof(f552,plain,(
  ( ! [X3] : (~po(sK2,X3) | ~zex(sK22(X3,sK2))) )),
  inference(resolution,[],[f385,f458])).
fof(f458,plain,(
  ( ! [X0] : (~p(X0,sK2) | ~zex(X0)) )),
  inference(resolution,[],[f456,f445])).
fof(f445,plain,(
  ( ! [X4,X5] : (~less(X4,X5) | ~p(X4,X5)) )),
  inference(resolution,[],[f389,f373])).
fof(f373,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X78] : ! [X79] : (eqdim(X78,X79) <=> (~less(X78,X79) & ~less(X79,X78)))),
  file(unknown,unknown)).
fof(f389,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X89] : ! [X90] : (p(X89,X90) <=> (cont(X89,X90) & eqdim(X89,X90)))),
  file(unknown,unknown)).
fof(f456,plain,(
  ( ! [X2] : (less(X2,sK2) | ~zex(X2)) )),
  inference(resolution,[],[f329,f433])).
fof(f329,plain,(
  ( ! [X0,X1] : (zex(X1) | less(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : (~zex(X0) | zex(X1) | less(X0,X1))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : ((~zex(X0) | zex(X1)) | less(X0,X1))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X75] : ! [X76] : ((zex(X75) & ~zex(X76)) => less(X75,X76))),
  file(unknown,unknown)).
fof(f385,plain,(
  ( ! [X0,X1] : (p(sK22(X1,X0),X0) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1] : ((~po(X0,X1) | (p(sK22(X1,X0),X0) & p(sK22(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(skolemisation,[status(esa)],[f290])).
fof(f290,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(rectify,[],[f289])).
fof(f289,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X2] : (~p(X2,X0) | ~p(X2,X1)) | po(X0,X1)))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0] : ! [X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X93] : ! [X94] : (po(X93,X94) <=> ? [X95] : (p(X95,X93) & p(X95,X94)))),
  file(unknown,unknown)).
fof(f387,plain,(
  ( ! [X0,X3,X1] : (po(X0,X1) | ~p(X3,X1) | ~p(X3,X0)) )),
  inference(cnf_transformation,[],[f291])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 323
Passive clauses: 763
Generated clauses: 1235
Final active clauses: 287
Final passive clauses: 328
Input formulas: 58
Initial clauses: 116

Selected by SInE selection: 58
SInE iterations: 1

Duplicate literals: 12
Fw subsumption resolutions: 40

Simple tautologies: 12
Forward subsumptions: 412

Binary resolution: 1029
Factoring: 3

Backtracking splits: 21
Backtracking splits refuted: 5
Backtracking splits refuted at zero level: 5

Memory used: 511KB
Time elapsed: 0.042 s
------------------------------
% Success in time 0.062 s
