Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 1932
Passive clauses: 19412
Generated clauses: 36324
Input formulas: 38
Initial clauses: 87

Selected by SInE selection: 38
SInE iterations: 1

Duplicate literals: 261
Fw subsumption resolutions: 1566

Simple tautologies: 94
Forward subsumptions: 14319

Binary resolution: 33287
Factoring: 90

Backtracking splits: 614
Backtracking splits refuted: 265
Backtracking splits refuted at zero level: 103

Memory used: 6396KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2402
Passive clauses: 28230
Generated clauses: 75964
Input formulas: 38
Initial clauses: 86

Selected by SInE selection: 38
SInE iterations: 1

Duplicate literals: 4716
Fw subsumption resolutions: 11618

Simple tautologies: 1795
Equational tautologies: 10
Forward subsumptions: 30491

Binary resolution: 46992
Factoring: 189
Forward superposition: 7110
Backward superposition: 5581

Unique components: 1

Memory used: 11513KB
Time elapsed: 1.400 s
------------------------------
% remaining time: 1181 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fsr=off:fde=none:gsp=input_only:lcm=reverse:nwc=10:ptb=off:ssec=off:stl=60:sos=on:sgo=on:sio=off:spl=backtracking:sfv=off:sp=reverse_arity_23 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f19169,plain,(
  $false),
  inference(subsumption_resolution,[],[f19168,f3203])).
fof(f3203,plain,(
  ~cont(sK16,sK16)),
  inference(superposition,[],[f333,f3194])).
fof(f3194,plain,(
  difference(sK0,sK0) = sK16),
  inference(resolution,[],[f334,f285])).
fof(f285,plain,(
  zex(sK16)),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  zex(sK16)),
  inference(skolemisation,[status(esa)],[f105])).
fof(f105,plain,(
  ? [X0] : zex(X0)),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ? [X29] : zex(X29)),
  file(unknown,unknown)).
fof(f334,plain,(
  ( ! [X0] : (~zex(X0) | difference(sK0,sK0) = X0) )),
  inference(resolution,[],[f308,f223])).
fof(f223,plain,(
  ( ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X36] : ! [X37] : ((zex(X36) & zex(X37)) => X36 = X37)),
  file(unknown,unknown)).
fof(f308,plain,(
  zex(difference(sK0,sK0))),
  inference(resolution,[],[f304,f273])).
fof(f273,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X15] : ! [X16] : (zex(difference(X15,X16)) <=> (zex(X15) | cont(X15,X16)))),
  file(unknown,unknown)).
fof(f304,plain,(
  cont(sK0,sK0)),
  inference(resolution,[],[f299,f215])).
fof(f215,plain,(
  ( ! [X0] : (zex(X0) | cont(X0,X0)) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X81] : (~zex(X81) <=> cont(X81,X81))),
  file(unknown,unknown)).
fof(f299,plain,(
  ~zex(sK0)),
  inference(resolution,[],[f288,f218])).
fof(f218,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X87] : ! [X88] : (zex(X87) => (~cont(X88,X87) & ~cont(X87,X88)))),
  file(unknown,unknown)).
fof(f288,plain,(
  cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK0)),
  inference(resolution,[],[f286,f281])).
fof(f281,plain,(
  ( ! [X2,X0,X1] : (~p(X2,intersection(X0,X1)) | cont(X2,X0)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(flattening,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(nnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X24] : ! [X25] : ! [X26] : ((cont(X26,X24) & cont(X26,X25) & eqdim(X26,intersection(X24,X25))) <=> p(X26,intersection(X24,X25)))),
  file(unknown,unknown)).
fof(f286,plain,(
  p(sK13(difference(sK0,sK1),intersection(sK0,sK1)),intersection(sK0,sK1))),
  inference(resolution,[],[f202,f259])).
fof(f259,plain,(
  ( ! [X0,X1] : (~po(X0,X1) | p(sK13(X1,X0),X0)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : ((~po(X0,X1) | (p(sK13(X1,X0),X0) & p(sK13(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(skolemisation,[status(esa)],[f187])).
fof(f187,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(rectify,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X2] : (~p(X2,X0) | ~p(X2,X1)) | po(X0,X1)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X56] : ! [X57] : (po(X56,X57) <=> ? [X58] : (p(X58,X56) & p(X58,X57)))),
  file(unknown,unknown)).
fof(f202,plain,(
  po(intersection(sK0,sK1),difference(sK0,sK1))),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  po(intersection(sK0,sK1),difference(sK0,sK1))),
  inference(skolemisation,[status(esa)],[f106])).
fof(f106,plain,(
  ? [X0,X1] : po(intersection(X0,X1),difference(X0,X1))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ~po(intersection(X0,X1),difference(X0,X1))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : ~po(intersection(X0,X1),difference(X0,X1))),
  file(unknown,unknown)).
fof(f333,plain,(
  ~cont(difference(sK0,sK0),difference(sK0,sK0))),
  inference(resolution,[],[f308,f216])).
fof(f216,plain,(
  ( ! [X0] : (~zex(X0) | ~cont(X0,X0)) )),
  inference(cnf_transformation,[],[f163])).
fof(f19168,plain,(
  cont(sK16,sK16)),
  inference(forward_demodulation,[],[f19166,f19139])).
fof(f19139,plain,(
  sK13(difference(sK0,sK1),intersection(sK0,sK1)) = sK16),
  inference(backtracking_split_refutation,[],[f7488,f7503_D,f19134])).
fof(f19134,plain,(
  $false | $spl328),
  inference(subsumption_resolution,[],[f300,f19133])).
fof(f19133,plain,(
  zex(sK13(difference(sK0,sK1),intersection(sK0,sK1))) | $spl328),
  inference(forward_demodulation,[],[f19128,f16627])).
fof(f16627,plain,(
  difference(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK16) = sK13(difference(sK0,sK1),intersection(sK0,sK1))),
  inference(resolution,[],[f16595,f300])).
fof(f16595,plain,(
  ( ! [X17] : (zex(X17) | difference(X17,sK16) = X17) )),
  inference(backtracking_split_refutation,[],[f16522,f16593_D,f16594])).
fof(f16594,plain,(
  $false | $spl561),
  inference(subsumption_resolution,[],[f284,f16593])).
fof(f284,plain,(
  mindim(sK15)),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  mindim(sK15)),
  inference(skolemisation,[status(esa)],[f104])).
fof(f104,plain,(
  ? [X0] : mindim(X0)),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ? [X40] : mindim(X40)),
  file(unknown,unknown)).
fof(f16593,plain,(
  ~mindim(sK15) | $spl561),
  inference(cnf_transformation,[],[f16593_D])).
fof(f16593_D,plain,(
  ~mindim(sK15) <=> ~$spl561),
  introduced(backtracking_splitting_component,[])).
fof(f16522,plain,(
  ( ! [X17] : (zex(X17) | ~mindim(sK15) | difference(X17,sK16) = X17) )),
  inference(resolution,[],[f212,f6813])).
fof(f6813,plain,(
  ( ! [X5] : (less(X5,sK15) | difference(X5,sK16) = X5) )),
  inference(resolution,[],[f6707,f220])).
fof(f220,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | difference(X0,X1) = X0) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : (~less(X1,X0) | difference(X0,X1) = X0)),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X4] : ! [X5] : (less(X5,X4) => difference(X4,X5) = X4)),
  file(unknown,unknown)).
fof(f6707,plain,(
  ( ! [X51] : (less(sK16,X51) | less(X51,sK15)) )),
  inference(resolution,[],[f3293,f3418])).
fof(f3418,plain,(
  ~zex(sK15)),
  inference(backtracking_split_refutation,[],[f3405,f3406_D,f3415])).
fof(f3415,plain,(
  $false | $spl223),
  inference(subsumption_resolution,[],[f284,f3409])).
fof(f3409,plain,(
  ~mindim(sK15) | $spl223),
  inference(resolution,[],[f3406,f211])).
fof(f211,plain,(
  ( ! [X0] : (~zex(X0) | ~mindim(X0)) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1)))) & (zex(X0) | (less(sK5(X0),X0) & ~zex(sK5(X0))) | mindim(X0)))),
  inference(skolemisation,[status(esa)],[f161])).
fof(f161,plain,(
  ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1)))) & (zex(X0) | ? [X2] : (less(X2,X0) & ~zex(X2)) | mindim(X0)))),
  inference(rectify,[],[f160])).
fof(f160,plain,(
  ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1)))) & (zex(X0) | ? [X1] : (less(X1,X0) & ~zex(X1)) | mindim(X0)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1)))) & ((zex(X0) | ? [X1] : (less(X1,X0) & ~zex(X1))) | mindim(X0)))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : (mindim(X0) <=> (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1))))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (mindim(X0) <=> (~zex(X0) & ! [X1] : (less(X1,X0) => zex(X1))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X47] : (mindim(X47) <=> (~zex(X47) & ! [X48] : (less(X48,X47) => zex(X48))))),
  file(unknown,unknown)).
fof(f3406,plain,(
  zex(sK15) | $spl223),
  inference(cnf_transformation,[],[f3406_D])).
fof(f3406_D,plain,(
  zex(sK15) <=> ~$spl223),
  introduced(backtracking_splitting_component,[])).
fof(f3405,plain,(
  eqdim(sK15,sK15) | zex(sK15)),
  inference(forward_demodulation,[],[f3403,f3347])).
fof(f3347,plain,(
  difference(sK15,sK16) = sK15),
  inference(resolution,[],[f3343,f220])).
fof(f3343,plain,(
  less(sK16,sK15)),
  inference(resolution,[],[f3290,f284])).
fof(f3290,plain,(
  ( ! [X2] : (~mindim(X2) | less(sK16,X2)) )),
  inference(forward_demodulation,[],[f3269,f3194])).
fof(f3269,plain,(
  ( ! [X2] : (less(difference(sK0,sK0),X2) | ~mindim(X2)) )),
  inference(resolution,[],[f335,f211])).
fof(f335,plain,(
  ( ! [X1] : (zex(X1) | less(difference(sK0,sK0),X1)) )),
  inference(resolution,[],[f308,f224])).
fof(f224,plain,(
  ( ! [X0,X1] : (~zex(X0) | zex(X1) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : (~zex(X0) | zex(X1) | less(X0,X1))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ((~zex(X0) | zex(X1)) | less(X0,X1))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X38] : ! [X39] : ((zex(X38) & ~zex(X39)) => less(X38,X39))),
  file(unknown,unknown)).
fof(f3403,plain,(
  zex(sK15) | eqdim(sK15,difference(sK15,sK16))),
  inference(superposition,[],[f229,f3347])).
fof(f229,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | eqdim(X0,difference(X0,X1))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) | eqdim(X0,difference(X0,X1)))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : (~zex(difference(X0,X1)) => eqdim(X0,difference(X0,X1)))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1] : (~zex(difference(X0,X1)) => eqdim(X0,difference(X0,X1)))),
  inference(rectify,[],[f2])).
fof(f2,axiom,(
  ! [X2] : ! [X3] : (~zex(difference(X2,X3)) => eqdim(X2,difference(X2,X3)))),
  file(unknown,unknown)).
fof(f3293,plain,(
  ( ! [X6,X7] : (zex(X7) | less(sK16,X6) | less(X6,X7)) )),
  inference(forward_demodulation,[],[f3272,f3194])).
fof(f3272,plain,(
  ( ! [X6,X7] : (less(difference(sK0,sK0),X6) | zex(X7) | less(X6,X7)) )),
  inference(resolution,[],[f335,f224])).
fof(f212,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | zex(X1) | ~mindim(X0)) )),
  inference(cnf_transformation,[],[f162])).
fof(f19128,plain,(
  zex(difference(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK16)) | $spl328),
  inference(resolution,[],[f7503,f273])).
fof(f300,plain,(
  ~zex(sK13(difference(sK0,sK1),intersection(sK0,sK1)))),
  inference(resolution,[],[f288,f219])).
fof(f219,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f112])).
fof(f7503,plain,(
  cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK16) | $spl328),
  inference(cnf_transformation,[],[f7503_D])).
fof(f7503_D,plain,(
  cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK16) <=> ~$spl328),
  introduced(backtracking_splitting_component,[])).
fof(f7488,plain,(
  sK13(difference(sK0,sK1),intersection(sK0,sK1)) = sK16 | cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK16)),
  inference(backtracking_split_refutation,[],[f7477,f7478_D,f7479])).
fof(f7479,plain,(
  $false | $spl324),
  inference(subsumption_resolution,[],[f474,f7478])).
fof(f474,plain,(
  leq(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1)),
  inference(resolution,[],[f289,f222])).
fof(f222,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1)) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X27] : ! [X28] : (cont(X27,X28) => leq(X27,X28))),
  file(unknown,unknown)).
fof(f289,plain,(
  cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1)),
  inference(resolution,[],[f286,f282])).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (~p(X2,intersection(X0,X1)) | cont(X2,X1)) )),
  inference(cnf_transformation,[],[f199])).
fof(f7478,plain,(
  ~leq(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1) | $spl324),
  inference(cnf_transformation,[],[f7478_D])).
fof(f7478_D,plain,(
  ~leq(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1) <=> ~$spl324),
  introduced(backtracking_splitting_component,[])).
fof(f7477,plain,(
  sK13(difference(sK0,sK1),intersection(sK0,sK1)) = sK16 | cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK16) | ~leq(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1)),
  inference(forward_demodulation,[],[f7476,f3194])).
fof(f7476,plain,(
  cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK16) | ~leq(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1) | difference(sK0,sK0) = sK13(difference(sK0,sK1),intersection(sK0,sK1))),
  inference(forward_demodulation,[],[f7475,f3194])).
fof(f7475,plain,(
  cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),difference(sK0,sK0)) | ~leq(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1) | difference(sK0,sK0) = sK13(difference(sK0,sK1),intersection(sK0,sK1))),
  inference(forward_demodulation,[],[f7138,f3190])).
fof(f3190,plain,(
  difference(sK0,sK0) = difference(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1)),
  inference(resolution,[],[f334,f472])).
fof(f472,plain,(
  zex(difference(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1))),
  inference(resolution,[],[f289,f273])).
fof(f7138,plain,(
  cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),difference(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1)) | ~leq(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1) | difference(sK0,sK0) = sK13(difference(sK0,sK1),intersection(sK0,sK1))),
  inference(resolution,[],[f2752,f3172])).
fof(f3172,plain,(
  ( ! [X0] : (cont(X0,X0) | difference(sK0,sK0) = X0) )),
  inference(resolution,[],[f334,f215])).
fof(f2752,plain,(
  ( ! [X0] : (~cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),X0) | cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),difference(X0,sK1)) | ~leq(X0,sK1)) )),
  inference(resolution,[],[f2153,f274])).
fof(f274,plain,(
  ( ! [X2,X0,X1] : (~less(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)) | ~cont(X2,X0) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X0) | ~less(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X0) | ~less(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & cont(X2,X0) & less(intersection(X2,X1),X2)) => cont(X2,difference(X0,X1)))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & cont(X2,X0) & less(intersection(X2,X1),X2)) => cont(X2,difference(X0,X1)))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X6] : ! [X7] : ! [X8] : ((leq(X6,X7) & cont(X8,X6) & less(intersection(X8,X7),X8)) => cont(X8,difference(X6,X7)))),
  file(unknown,unknown)).
fof(f2153,plain,(
  less(intersection(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1),sK13(difference(sK0,sK1),intersection(sK0,sK1)))),
  inference(backtracking_split_refutation,[],[f350,f514_D,f2150])).
fof(f2150,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f802,f2145])).
fof(f2145,plain,(
  less(sK13(sK0,intersection(sK0,sK1)),sK0) | $spl20),
  inference(backtracking_split_refutation,[],[f1097,f1386_D,f2143])).
fof(f2143,plain,(
  $false | ($spl20 | $spl79)),
  inference(subsumption_resolution,[],[f2126,f2140])).
fof(f2140,plain,(
  ~less(sK1,sK13(sK0,intersection(sK0,sK1))) | ($spl20 | $spl79)),
  inference(resolution,[],[f2136,f221])).
fof(f221,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~less(X1,X0)) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : (~less(X0,X1) | ~less(X1,X0))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X31] : ! [X32] : (less(X31,X32) => ~less(X32,X31))),
  file(unknown,unknown)).
fof(f2136,plain,(
  less(sK13(sK0,intersection(sK0,sK1)),sK1) | ($spl20 | $spl79)),
  inference(backtracking_split_refutation,[],[f1145,f1387_D,f2133])).
fof(f2133,plain,(
  $false | ($spl20 | $spl79 | $spl80)),
  inference(subsumption_resolution,[],[f2126,f2131])).
fof(f2131,plain,(
  ~less(sK1,sK13(sK0,intersection(sK0,sK1))) | $spl80),
  inference(resolution,[],[f1387,f263])).
fof(f263,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | ~less(X1,X0)) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X41] : ! [X42] : (eqdim(X41,X42) <=> (~less(X41,X42) & ~less(X42,X41)))),
  file(unknown,unknown)).
fof(f1387,plain,(
  eqdim(sK13(sK0,intersection(sK0,sK1)),sK1) | $spl80),
  inference(cnf_transformation,[],[f1387_D])).
fof(f1387_D,plain,(
  eqdim(sK13(sK0,intersection(sK0,sK1)),sK1) <=> ~$spl80),
  introduced(backtracking_splitting_component,[])).
fof(f1145,plain,(
  less(sK13(sK0,intersection(sK0,sK1)),sK1) | eqdim(sK13(sK0,intersection(sK0,sK1)),sK1) | $spl20),
  inference(resolution,[],[f787,f265])).
fof(f265,plain,(
  ( ! [X0,X1] : (~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X43] : ! [X44] : (leq(X43,X44) <=> (less(X43,X44) | eqdim(X43,X44)))),
  file(unknown,unknown)).
fof(f787,plain,(
  leq(sK13(sK0,intersection(sK0,sK1)),sK1) | $spl20),
  inference(forward_demodulation,[],[f474,f564])).
fof(f564,plain,(
  difference(sK0,sK1) = sK0 | $spl20),
  inference(resolution,[],[f533,f220])).
fof(f533,plain,(
  less(sK1,sK0) | $spl20),
  inference(backtracking_split_refutation,[],[f528,f530_D,f531])).
fof(f531,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f523,f530])).
fof(f523,plain,(
  ~eqdim(sK0,sK1) | $spl20),
  inference(resolution,[],[f514,f267])).
fof(f267,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f192])).
fof(f530,plain,(
  eqdim(sK0,sK1) | $spl24),
  inference(cnf_transformation,[],[f530_D])).
fof(f530_D,plain,(
  eqdim(sK0,sK1) <=> ~$spl24),
  introduced(backtracking_splitting_component,[])).
fof(f528,plain,(
  less(sK1,sK0) | eqdim(sK0,sK1) | $spl20),
  inference(resolution,[],[f522,f264])).
fof(f264,plain,(
  ( ! [X0,X1] : (less(X1,X0) | less(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f522,plain,(
  ~less(sK0,sK1) | $spl20),
  inference(resolution,[],[f514,f266])).
fof(f266,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f192])).
fof(f2126,plain,(
  less(sK1,sK13(sK0,intersection(sK0,sK1))) | ($spl20 | $spl79)),
  inference(resolution,[],[f1386,f1280])).
fof(f1280,plain,(
  ( ! [X1] : (~eqdim(sK0,X1) | less(sK1,X1)) ) | $spl20),
  inference(resolution,[],[f561,f267])).
fof(f561,plain,(
  ( ! [X0] : (~leq(sK0,X0) | less(sK1,X0)) ) | $spl20),
  inference(resolution,[],[f533,f277])).
fof(f277,plain,(
  ( ! [X2,X0,X1] : (~less(X0,X1) | ~leq(X1,X2) | less(X0,X2)) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : (~less(X0,X1) | ~leq(X1,X2) | less(X0,X2))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : ((~less(X0,X1) | ~leq(X1,X2)) | less(X0,X2))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X33] : ! [X34] : ! [X35] : ((less(X33,X34) & leq(X34,X35)) => less(X33,X35))),
  file(unknown,unknown)).
fof(f1386,plain,(
  eqdim(sK0,sK13(sK0,intersection(sK0,sK1))) | $spl79),
  inference(cnf_transformation,[],[f1386_D])).
fof(f1386_D,plain,(
  eqdim(sK0,sK13(sK0,intersection(sK0,sK1))) <=> ~$spl79),
  introduced(backtracking_splitting_component,[])).
fof(f1097,plain,(
  less(sK13(sK0,intersection(sK0,sK1)),sK0) | eqdim(sK0,sK13(sK0,intersection(sK0,sK1))) | $spl20),
  inference(resolution,[],[f784,f264])).
fof(f784,plain,(
  ~less(sK0,sK13(sK0,intersection(sK0,sK1))) | $spl20),
  inference(forward_demodulation,[],[f779,f564])).
fof(f779,plain,(
  ~less(difference(sK0,sK1),sK13(difference(sK0,sK1),intersection(sK0,sK1)))),
  inference(resolution,[],[f352,f263])).
fof(f352,plain,(
  eqdim(sK13(difference(sK0,sK1),intersection(sK0,sK1)),difference(sK0,sK1))),
  inference(resolution,[],[f287,f257])).
fof(f257,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X52] : ! [X53] : (p(X52,X53) <=> (cont(X52,X53) & eqdim(X52,X53)))),
  file(unknown,unknown)).
fof(f287,plain,(
  p(sK13(difference(sK0,sK1),intersection(sK0,sK1)),difference(sK0,sK1))),
  inference(resolution,[],[f202,f260])).
fof(f260,plain,(
  ( ! [X0,X1] : (~po(X0,X1) | p(sK13(X1,X0),X1)) )),
  inference(cnf_transformation,[],[f188])).
fof(f802,plain,(
  ~less(sK13(sK0,intersection(sK0,sK1)),sK0) | $spl20),
  inference(forward_demodulation,[],[f778,f564])).
fof(f778,plain,(
  ~less(sK13(difference(sK0,sK1),intersection(sK0,sK1)),difference(sK0,sK1))),
  inference(resolution,[],[f352,f262])).
fof(f262,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f514,plain,(
  ~leq(sK0,sK1) | $spl20),
  inference(cnf_transformation,[],[f514_D])).
fof(f514_D,plain,(
  ~leq(sK0,sK1) <=> ~$spl20),
  introduced(backtracking_splitting_component,[])).
fof(f350,plain,(
  less(intersection(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK1),sK13(difference(sK0,sK1),intersection(sK0,sK1))) | ~leq(sK0,sK1)),
  inference(resolution,[],[f287,f275])).
fof(f275,plain,(
  ( ! [X2,X0,X1] : (~p(X2,difference(X0,X1)) | less(intersection(X2,X1),X2) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~p(X2,difference(X0,X1)) | less(intersection(X2,X1),X2))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~p(X2,difference(X0,X1))) | less(intersection(X2,X1),X2))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & p(X2,difference(X0,X1))) => less(intersection(X2,X1),X2))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & p(X2,difference(X0,X1))) => less(intersection(X2,X1),X2))),
  inference(rectify,[],[f6])).
fof(f6,axiom,(
  ! [X12] : ! [X13] : ! [X14] : ((leq(X12,X13) & p(X14,difference(X12,X13))) => less(intersection(X14,X13),X14))),
  file(unknown,unknown)).
fof(f19166,plain,(
  cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK16)),
  inference(backtracking_split_refutation,[],[f9088,f19155_D,f19163])).
fof(f19163,plain,(
  $false | $spl605),
  inference(subsumption_resolution,[],[f19155,f538])).
fof(f538,plain,(
  leq(sK1,sK1)),
  inference(resolution,[],[f521,f222])).
fof(f521,plain,(
  cont(sK1,sK1)),
  inference(resolution,[],[f475,f215])).
fof(f475,plain,(
  ~zex(sK1)),
  inference(resolution,[],[f289,f218])).
fof(f19155,plain,(
  ~leq(sK1,sK1) | $spl605),
  inference(cnf_transformation,[],[f19155_D])).
fof(f19155_D,plain,(
  ~leq(sK1,sK1) <=> ~$spl605),
  introduced(backtracking_splitting_component,[])).
fof(f9088,plain,(
  cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),sK16) | ~leq(sK1,sK1)),
  inference(forward_demodulation,[],[f7137,f8805])).
fof(f8805,plain,(
  difference(sK1,sK1) = sK16),
  inference(forward_demodulation,[],[f3184,f3194])).
fof(f3184,plain,(
  difference(sK0,sK0) = difference(sK1,sK1)),
  inference(resolution,[],[f334,f536])).
fof(f536,plain,(
  zex(difference(sK1,sK1))),
  inference(resolution,[],[f521,f273])).
fof(f7137,plain,(
  cont(sK13(difference(sK0,sK1),intersection(sK0,sK1)),difference(sK1,sK1)) | ~leq(sK1,sK1)),
  inference(resolution,[],[f2752,f289])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 2674
Passive clauses: 5535
Generated clauses: 21264
Final active clauses: 990
Final passive clauses: 806
Input formulas: 38
Initial clauses: 84

Duplicate literals: 64
Trivial inequalities: 14
Bw subsumption resolutions: 1329
Fw demodulations: 3812
Condensations: 23

Simple tautologies: 162
Equational tautologies: 280
Forward subsumptions: 10007

Binary resolution: 10794
Factoring: 2
Forward superposition: 390
Backward superposition: 1151

Backtracking splits: 606
Backtracking splits refuted: 344
Backtracking splits refuted at zero level: 26

Memory used: 2686KB
Time elapsed: 0.164 s
------------------------------
% Success in time 2.03 s
