Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 1983
Passive clauses: 20560
Generated clauses: 33294
Input formulas: 58
Initial clauses: 115

Selected by SInE selection: 58
SInE iterations: 1

Duplicate literals: 208
Fw subsumption resolutions: 1876

Simple tautologies: 86
Forward subsumptions: 9608

Binary resolution: 29354
Factoring: 55

Backtracking splits: 1331
Backtracking splits refuted: 195
Backtracking splits refuted at zero level: 20

Memory used: 6780KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f2588,plain,(
  $false),
  inference(subsumption_resolution,[],[f2587,f1320])).
fof(f1320,plain,(
  ~less(sK0,sK1)),
  inference(resolution,[],[f1312,f316])).
fof(f316,plain,(
  ( ! [X0] : (~less(X0,X0)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X65] : ~less(X65,X65)),
  file(unknown,unknown)).
fof(f1312,plain,(
  ( ! [X1] : (less(sK1,X1) | ~less(sK0,X1)) )),
  inference(subsumption_resolution,[],[f901,f875])).
fof(f875,plain,(
  ( ! [X1] : (~less(sK0,X1) | ~less(X1,sK1)) )),
  inference(resolution,[],[f873,f376])).
fof(f376,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X78] : ! [X79] : (leq(X78,X79) <=> (less(X78,X79) | eqdim(X78,X79)))),
  file(unknown,unknown)).
fof(f873,plain,(
  ( ! [X20] : (~leq(X20,sK1) | ~less(sK0,X20)) )),
  inference(subsumption_resolution,[],[f872,f630])).
fof(f630,plain,(
  ~zex(sK1)),
  inference(factoring,[],[f629])).
fof(f629,plain,(
  ( ! [X0] : (~zex(sK1) | ~zex(X0)) )),
  inference(subsumption_resolution,[],[f628,f301])).
fof(f301,plain,(
  ~zex(sK0)),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  zex(sum(sK0,sK1)) & ~zex(sK0)),
  inference(skolemisation,[status(esa)],[f166])).
fof(f166,plain,(
  ? [X0,X1] : (zex(sum(X0,X1)) & ~zex(X0))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ~! [X0,X1] : (zex(sum(X0,X1)) => zex(X0))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : (zex(sum(X0,X1)) => zex(X0))),
  file(unknown,unknown)).
fof(f628,plain,(
  ( ! [X0] : (zex(sK0) | ~zex(sK1) | ~zex(X0)) )),
  inference(duplicate_literal_removal,[],[f623])).
fof(f623,plain,(
  ( ! [X0] : (~zex(sK1) | ~zex(X0) | zex(sK0) | ~zex(X0)) )),
  inference(resolution,[],[f604,f328])).
fof(f328,plain,(
  ( ! [X0,X1] : (less(X0,X1) | zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : (~zex(X0) | zex(X1) | less(X0,X1))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : ((~zex(X0) | zex(X1)) | less(X0,X1))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X73] : ! [X74] : ((zex(X73) & ~zex(X74)) => less(X73,X74))),
  file(unknown,unknown)).
fof(f604,plain,(
  ( ! [X5] : (~less(X5,sK0) | ~zex(sK1) | ~zex(X5)) )),
  inference(subsumption_resolution,[],[f603,f301])).
fof(f603,plain,(
  ( ! [X5] : (zex(sK0) | ~zex(X5) | ~zex(sK1) | ~less(X5,sK0)) )),
  inference(superposition,[],[f534,f324])).
fof(f324,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = X1 | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : (~less(X0,X1) | sum(X0,X1) = X1)),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : (less(X0,X1) => sum(X0,X1) = X1)),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => sum(X0,X1) = X1)),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X4] : ! [X5] : (less(X4,X5) => sum(X4,X5) = X5)),
  file(unknown,unknown)).
fof(f534,plain,(
  ( ! [X0] : (zex(sum(X0,sK0)) | ~zex(X0) | ~zex(sK1)) )),
  inference(superposition,[],[f474,f317])).
fof(f317,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = sum(X1,X0)) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : sum(X0,X1) = sum(X1,X0)),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)),
  inference(rectify,[],[f2])).
fof(f2,axiom,(
  ! [X2] : ! [X3] : sum(X2,X3) = sum(X3,X2)),
  file(unknown,unknown)).
fof(f474,plain,(
  ( ! [X0] : (zex(sum(sK0,X0)) | ~zex(X0) | ~zex(sK1)) )),
  inference(superposition,[],[f300,f454])).
fof(f454,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0)) )),
  inference(resolution,[],[f412,f421])).
fof(f421,plain,(
  ( ! [X0,X1] : (~sQ1_eqProxy(X0,X1) | X0 = X1) )),
  introduced(equality_proxy_axiom,[])).
fof(f412,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X0,X1) | ~zex(X1) | ~zex(X0)) )),
  inference(equality_proxy_replacement,[],[f327])).
fof(f327,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X71] : ! [X72] : ((zex(X71) & zex(X72)) => X71 = X72)),
  file(unknown,unknown)).
fof(f300,plain,(
  zex(sum(sK0,sK1))),
  inference(cnf_transformation,[],[f237])).
fof(f872,plain,(
  ( ! [X20] : (~leq(X20,sK1) | ~less(sK0,X20) | zex(sK1)) )),
  inference(resolution,[],[f400,f599])).
fof(f599,plain,(
  ~less(sK0,sK1) | zex(sK1)),
  inference(superposition,[],[f300,f324])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (less(X0,X2) | ~leq(X1,X2) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : (~less(X0,X1) | ~leq(X1,X2) | less(X0,X2))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((~less(X0,X1) | ~leq(X1,X2)) | less(X0,X2))),
  inference(ennf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X68] : ! [X69] : ! [X70] : ((less(X68,X69) & leq(X69,X70)) => less(X68,X70))),
  file(unknown,unknown)).
fof(f901,plain,(
  ( ! [X1] : (~less(sK0,X1) | less(sK1,X1) | less(X1,sK1)) )),
  inference(resolution,[],[f874,f374])).
fof(f374,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | less(X1,X0) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X76] : ! [X77] : (eqdim(X76,X77) <=> (~less(X76,X77) & ~less(X77,X76)))),
  file(unknown,unknown)).
fof(f874,plain,(
  ( ! [X0] : (~eqdim(X0,sK1) | ~less(sK0,X0)) )),
  inference(resolution,[],[f873,f377])).
fof(f377,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f283])).
fof(f2587,plain,(
  less(sK0,sK1)),
  inference(subsumption_resolution,[],[f2581,f1969])).
fof(f1969,plain,(
  ~less(sK1,sK0)),
  inference(subsumption_resolution,[],[f1962,f301])).
fof(f1962,plain,(
  zex(sK0) | ~less(sK1,sK0)),
  inference(superposition,[],[f300,f597])).
fof(f597,plain,(
  ( ! [X2,X3] : (sum(X3,X2) = X3 | ~less(X2,X3)) )),
  inference(superposition,[],[f324,f317])).
fof(f2581,plain,(
  less(sK1,sK0) | less(sK0,sK1)),
  inference(resolution,[],[f2528,f374])).
fof(f2528,plain,(
  ~eqdim(sK0,sK1)),
  inference(resolution,[],[f2527,f377])).
fof(f2527,plain,(
  ~leq(sK0,sK1)),
  inference(subsumption_resolution,[],[f2511,f630])).
fof(f2511,plain,(
  ~leq(sK0,sK1) | zex(sK1)),
  inference(resolution,[],[f2499,f314])).
fof(f314,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X134] : (~zex(X134) <=> cont(X134,X134))),
  file(unknown,unknown)).
fof(f2499,plain,(
  ( ! [X0] : (~cont(X0,sK1) | ~leq(sK0,sK1)) )),
  inference(resolution,[],[f1130,f300])).
fof(f1130,plain,(
  ( ! [X10,X8,X9] : (~zex(sum(X10,X9)) | ~leq(X10,X9) | ~cont(X8,X9)) )),
  inference(resolution,[],[f397,f318])).
fof(f318,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X140] : ! [X141] : (zex(X140) => (~cont(X141,X140) & ~cont(X140,X141)))),
  file(unknown,unknown)).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (cont(X2,sum(X0,X1)) | ~cont(X2,X1) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X1) | cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X1)) | cont(X2,sum(X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & cont(X2,X1)) => cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & cont(X2,X1)) => cont(X2,sum(X0,X1)))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X6] : ! [X7] : ! [X8] : ((leq(X6,X7) & cont(X8,X7)) => cont(X8,sum(X6,X7)))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 350
Passive clauses: 1062
Generated clauses: 2281
Final active clauses: 350
Final passive clauses: 684
Input formulas: 58
Initial clauses: 114

Selected by SInE selection: 58
SInE iterations: 1

Duplicate literals: 55
Fw subsumption resolutions: 145

Simple tautologies: 89
Equational tautologies: 5
Forward subsumptions: 949

Binary resolution: 1066
Factoring: 10
Forward superposition: 830
Backward superposition: 63

Unique components: 2

Memory used: 639KB
Time elapsed: 0.052 s
------------------------------
% Success in time 0.498 s
