Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f20957,plain,(
  $false),
  inference(subsumption_resolution,[],[f20956,f20803])).
fof(f20803,plain,(
  cont(intersection(sK2,sK0),sK2)),
  inference(resolution,[],[f20788,f394])).
fof(f394,plain,(
  ( ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X55] : ! [X56] : (~zex(intersection(X55,X56)) => cont(intersection(X55,X56),X55))),
  file(unknown,unknown)).
fof(f20788,plain,(
  ~zex(intersection(sK2,sK0))),
  inference(resolution,[],[f20784,f544])).
fof(f544,plain,(
  ( ! [X8] : (less(X8,sK2) | ~zex(X8)) )),
  inference(resolution,[],[f374,f495])).
fof(f495,plain,(
  ~zex(sK2)),
  inference(resolution,[],[f363,f338])).
fof(f338,plain,(
  cont(sK2,sum(sK0,sK1))),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  eqdim(sK0,sK1) & cont(sK2,sum(sK0,sK1)) & ~cont(sK2,sK0) & ~cont(sK2,sK1) & ~zex(sK0) & ~zex(sK1) & ~p(intersection(sK2,sK0),sK2)),
  inference(skolemisation,[status(esa)],[f190])).
fof(f190,plain,(
  ? [X0,X1,X2] : (eqdim(X0,X1) & cont(X2,sum(X0,X1)) & ~cont(X2,X0) & ~cont(X2,X1) & ~zex(X0) & ~zex(X1) & ~p(intersection(X2,X0),X2))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ? [X0,X1,X2] : ((eqdim(X0,X1) & cont(X2,sum(X0,X1)) & ~cont(X2,X0) & ~cont(X2,X1) & ~zex(X0) & ~zex(X1)) & ~p(intersection(X2,X0),X2))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ~! [X0,X1,X2] : ((eqdim(X0,X1) & cont(X2,sum(X0,X1)) & ~cont(X2,X0) & ~cont(X2,X1) & ~zex(X0) & ~zex(X1)) => p(intersection(X2,X0),X2))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ~! [X0] : ! [X1] : ! [X2] : ((eqdim(X0,X1) & cont(X2,sum(X0,X1)) & ~cont(X2,X0) & ~cont(X2,X1) & ~zex(X0) & ~zex(X1)) => p(intersection(X2,X0),X2))),
  inference(rectify,[],[f9])).
fof(f9,negated_conjecture,(
  ~! [X19] : ! [X20] : ! [X21] : ((eqdim(X19,X20) & cont(X21,sum(X19,X20)) & ~cont(X21,X19) & ~cont(X21,X20) & ~zex(X19) & ~zex(X20)) => p(intersection(X21,X19),X21))),
  file(unknown,unknown)).
fof(f363,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X160] : ! [X161] : (zex(X160) => (~cont(X161,X160) & ~cont(X160,X161)))),
  file(unknown,unknown)).
fof(f374,plain,(
  ( ! [X0,X1] : (zex(X1) | less(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : (~zex(X0) | zex(X1) | less(X0,X1))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((~zex(X0) | zex(X1)) | less(X0,X1))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X93] : ! [X94] : ((zex(X93) & ~zex(X94)) => less(X93,X94))),
  file(unknown,unknown)).
fof(f20784,plain,(
  ~less(intersection(sK2,sK0),sK2)),
  inference(backtracking_split_refutation,[],[f20750,f20751_D,f20783])).
fof(f20783,plain,(
  $false | $spl497),
  inference(subsumption_resolution,[],[f20780,f337])).
fof(f337,plain,(
  eqdim(sK0,sK1)),
  inference(cnf_transformation,[],[f273])).
fof(f20780,plain,(
  ~eqdim(sK0,sK1) | $spl497),
  inference(resolution,[],[f20758,f2162])).
fof(f2162,plain,(
  ( ! [X26] : (eqdim(X26,sum(X26,sK1)) | ~eqdim(X26,sK1)) )),
  inference(subsumption_resolution,[],[f2140,f560])).
fof(f560,plain,(
  ( ! [X1] : (~eqdim(X1,sK1) | ~zex(X1)) )),
  inference(resolution,[],[f543,f431])).
fof(f431,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X96] : ! [X97] : (eqdim(X96,X97) <=> (~less(X96,X97) & ~less(X97,X96)))),
  file(unknown,unknown)).
fof(f543,plain,(
  ( ! [X7] : (less(X7,sK1) | ~zex(X7)) )),
  inference(resolution,[],[f374,f342])).
fof(f342,plain,(
  ~zex(sK1)),
  inference(cnf_transformation,[],[f273])).
fof(f2140,plain,(
  ( ! [X26] : (eqdim(X26,sum(X26,sK1)) | zex(X26) | ~eqdim(X26,sK1)) )),
  inference(resolution,[],[f376,f342])).
fof(f376,plain,(
  ( ! [X0,X1] : (zex(X1) | eqdim(X0,sum(X0,X1)) | zex(X0) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1] : (~eqdim(X0,X1) | zex(X0) | zex(X1) | eqdim(X0,sum(X0,X1)))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | zex(X0) | zex(X1)) | eqdim(X0,sum(X0,X1)))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((eqdim(X0,X1) & ~zex(X0) & ~zex(X1)) => eqdim(X0,sum(X0,X1)))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1] : ((eqdim(X0,X1) & ~zex(X0) & ~zex(X1)) => eqdim(X0,sum(X0,X1)))),
  inference(rectify,[],[f6])).
fof(f6,axiom,(
  ! [X11] : ! [X12] : ((eqdim(X11,X12) & ~zex(X11) & ~zex(X12)) => eqdim(X11,sum(X11,X12)))),
  file(unknown,unknown)).
fof(f20758,plain,(
  ~eqdim(sK0,sum(sK0,sK1)) | $spl497),
  inference(resolution,[],[f20752,f1272])).
fof(f1272,plain,(
  ( ! [X4,X3] : (eqdim(X3,X4) | ~eqdim(X4,X3)) )),
  inference(subsumption_resolution,[],[f1243,f432])).
fof(f432,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f327])).
fof(f1243,plain,(
  ( ! [X4,X3] : (eqdim(X3,X4) | less(X3,X4) | ~eqdim(X4,X3)) )),
  inference(resolution,[],[f433,f431])).
fof(f433,plain,(
  ( ! [X0,X1] : (less(X1,X0) | eqdim(X0,X1) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f327])).
fof(f20752,plain,(
  ~eqdim(sum(sK0,sK1),sK0) | $spl497),
  inference(resolution,[],[f20751,f421])).
fof(f421,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f316])).
fof(f316,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X98] : ! [X99] : (leq(X98,X99) <=> (less(X98,X99) | eqdim(X98,X99)))),
  file(unknown,unknown)).
fof(f20751,plain,(
  ~leq(sum(sK0,sK1),sK0) | $spl497),
  inference(cnf_transformation,[],[f20751_D])).
fof(f20751_D,plain,(
  ~leq(sum(sK0,sK1),sK0) <=> ~$spl497),
  introduced(backtracking_splitting_component,[])).
fof(f20750,plain,(
  ~less(intersection(sK2,sK0),sK2) | ~leq(sum(sK0,sK1),sK0)),
  inference(subsumption_resolution,[],[f20748,f338])).
fof(f20748,plain,(
  ~less(intersection(sK2,sK0),sK2) | ~cont(sK2,sum(sK0,sK1)) | ~leq(sum(sK0,sK1),sK0)),
  inference(resolution,[],[f20736,f447])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (cont(X2,difference(X0,X1)) | ~less(intersection(X2,X1),X2) | ~cont(X2,X0) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X0) | ~less(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X0) | ~less(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))),
  inference(ennf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & cont(X2,X0) & less(intersection(X2,X1),X2)) => cont(X2,difference(X0,X1)))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & cont(X2,X0) & less(intersection(X2,X1),X2)) => cont(X2,difference(X0,X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X36] : ! [X37] : ! [X38] : ((leq(X36,X37) & cont(X38,X36) & less(intersection(X38,X37),X38)) => cont(X38,difference(X36,X37)))),
  file(unknown,unknown)).
fof(f20736,plain,(
  ~cont(sK2,difference(sum(sK0,sK1),sK0))),
  inference(subsumption_resolution,[],[f20690,f1403])).
fof(f1403,plain,(
  ~cont(sum(sK0,sK1),sK0)),
  inference(resolution,[],[f1389,f338])).
fof(f1389,plain,(
  ( ! [X0] : (~cont(X0,sK0) | ~cont(sK2,X0)) )),
  inference(resolution,[],[f456,f339])).
fof(f339,plain,(
  ~cont(sK2,sK0)),
  inference(cnf_transformation,[],[f273])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))),
  inference(ennf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X157] : ! [X158] : ! [X159] : ((cont(X157,X158) & cont(X158,X159)) => cont(X157,X159))),
  file(unknown,unknown)).
fof(f20690,plain,(
  ~cont(sK2,difference(sum(sK0,sK1),sK0)) | cont(sum(sK0,sK1),sK0)),
  inference(resolution,[],[f2570,f492])).
fof(f492,plain,(
  cont(sum(sK0,sK1),sum(sK0,sK1))),
  inference(resolution,[],[f490,f357])).
fof(f357,plain,(
  ( ! [X0] : (zex(X0) | cont(X0,X0)) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X154] : (~zex(X154) <=> cont(X154,X154))),
  file(unknown,unknown)).
fof(f490,plain,(
  ~zex(sum(sK0,sK1))),
  inference(resolution,[],[f362,f338])).
fof(f362,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f198])).
fof(f2570,plain,(
  ( ! [X17,X16] : (~cont(sK2,difference(X16,X17)) | ~cont(X16,sum(X17,sK1)) | cont(X16,X17)) )),
  inference(resolution,[],[f458,f1390])).
fof(f1390,plain,(
  ( ! [X1] : (~cont(X1,sK1) | ~cont(sK2,X1)) )),
  inference(resolution,[],[f456,f340])).
fof(f340,plain,(
  ~cont(sK2,sK1)),
  inference(cnf_transformation,[],[f273])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1))) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (~cont(X2,sum(X0,X1)) | cont(X2,X0) | cont(difference(X2,X0),X1))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : ((~cont(X2,sum(X0,X1)) | cont(X2,X0)) | cont(difference(X2,X0),X1))),
  inference(ennf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X29] : ! [X30] : ! [X31] : ((cont(X31,sum(X29,X30)) & ~cont(X31,X29)) => cont(difference(X31,X29),X30))),
  file(unknown,unknown)).
fof(f20956,plain,(
  ~cont(intersection(sK2,sK0),sK2)),
  inference(subsumption_resolution,[],[f20955,f343])).
fof(f343,plain,(
  ~p(intersection(sK2,sK0),sK2)),
  inference(cnf_transformation,[],[f273])).
fof(f20955,plain,(
  p(intersection(sK2,sK0),sK2) | ~cont(intersection(sK2,sK0),sK2)),
  inference(resolution,[],[f20948,f436])).
fof(f436,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f328])).
fof(f328,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X107] : ! [X108] : (p(X107,X108) <=> (cont(X107,X108) & eqdim(X107,X108)))),
  file(unknown,unknown)).
fof(f20948,plain,(
  eqdim(intersection(sK2,sK0),sK2)),
  inference(subsumption_resolution,[],[f20939,f20784])).
fof(f20939,plain,(
  eqdim(intersection(sK2,sK0),sK2) | less(intersection(sK2,sK0),sK2)),
  inference(resolution,[],[f20924,f433])).
fof(f20924,plain,(
  ~less(sK2,intersection(sK2,sK0))),
  inference(backtracking_split_refutation,[],[f20789,f20799_D,f20922])).
fof(f20922,plain,(
  $false | $spl500),
  inference(subsumption_resolution,[],[f20919,f20803])).
fof(f20919,plain,(
  ~cont(intersection(sK2,sK0),sK2) | $spl500),
  inference(resolution,[],[f20799,f5849])).
fof(f5849,plain,(
  ( ! [X6,X5] : (~less(X6,X5) | ~cont(X5,X6)) )),
  inference(subsumption_resolution,[],[f5793,f432])).
fof(f5793,plain,(
  ( ! [X6,X5] : (eqdim(X5,X6) | ~cont(X5,X6) | ~less(X6,X5)) )),
  inference(resolution,[],[f1134,f369])).
fof(f369,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1] : (~less(X0,X1) | ~less(X1,X0))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X86] : ! [X87] : (less(X86,X87) => ~less(X87,X86))),
  file(unknown,unknown)).
fof(f1134,plain,(
  ( ! [X4,X5] : (less(X4,X5) | eqdim(X4,X5) | ~cont(X4,X5)) )),
  inference(resolution,[],[f419,f370])).
fof(f370,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X69] : ! [X70] : (cont(X69,X70) => leq(X69,X70))),
  file(unknown,unknown)).
fof(f419,plain,(
  ( ! [X0,X1] : (~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f317])).
fof(f20799,plain,(
  less(sK2,intersection(sK2,sK0)) | $spl500),
  inference(cnf_transformation,[],[f20799_D])).
fof(f20799_D,plain,(
  less(sK2,intersection(sK2,sK0)) <=> ~$spl500),
  introduced(backtracking_splitting_component,[])).
fof(f20789,plain,(
  eqdim(sK2,intersection(sK2,sK0)) | less(sK2,intersection(sK2,sK0))),
  inference(resolution,[],[f20784,f433])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1576
Passive clauses: 12673
Generated clauses: 20616
Final active clauses: 1514
Final passive clauses: 10545
Input formulas: 66
Initial clauses: 131

Selected by SInE selection: 66
SInE iterations: 1

Duplicate literals: 44
Fw subsumption resolutions: 1340

Simple tautologies: 59
Forward subsumptions: 6242

Binary resolution: 18505
Factoring: 13

Backtracking splits: 508
Backtracking splits refuted: 39
Backtracking splits refuted at zero level: 4

Memory used: 4861KB
Time elapsed: 0.284 s
------------------------------
% Success in time 0.311 s
