Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2044
Passive clauses: 21859
Generated clauses: 33419
Input formulas: 58
Initial clauses: 116

Selected by SInE selection: 58
SInE iterations: 1

Duplicate literals: 214
Fw subsumption resolutions: 2423

Simple tautologies: 97
Forward subsumptions: 7293

Binary resolution: 28694
Factoring: 68

Backtracking splits: 1556
Backtracking splits refuted: 229
Backtracking splits refuted at zero level: 63

Memory used: 6908KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f19866,plain,(
  $false),
  inference(subsumption_resolution,[],[f19861,f302])).
fof(f302,plain,(
  ~zex(sK0)),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  less(sK1,sK0) & ~zex(sK0) & ~p(sK0,intersection(sK0,sum(sK0,sK1)))),
  inference(skolemisation,[status(esa)],[f167])).
fof(f167,plain,(
  ? [X0,X1] : (less(X1,X0) & ~zex(X0) & ~p(X0,intersection(X0,sum(X0,X1))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ? [X0,X1] : ((less(X1,X0) & ~zex(X0)) & ~p(X0,intersection(X0,sum(X0,X1))))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ~! [X0,X1] : ((less(X1,X0) & ~zex(X0)) => p(X0,intersection(X0,sum(X0,X1))))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : ((less(X1,X0) & ~zex(X0)) => p(X0,intersection(X0,sum(X0,X1))))),
  file(unknown,unknown)).
fof(f19861,plain,(
  zex(sK0)),
  inference(resolution,[],[f19812,f316])).
fof(f316,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X134] : (~zex(X134) <=> cont(X134,X134))),
  file(unknown,unknown)).
fof(f19812,plain,(
  ~cont(sK0,sK0)),
  inference(subsumption_resolution,[],[f19807,f301])).
fof(f301,plain,(
  less(sK1,sK0)),
  inference(cnf_transformation,[],[f238])).
fof(f19807,plain,(
  ~cont(sK0,sK0) | ~less(sK1,sK0)),
  inference(superposition,[],[f19769,f630])).
fof(f630,plain,(
  ( ! [X2,X3] : (sum(X3,X2) = X3 | ~less(X2,X3)) )),
  inference(superposition,[],[f326,f319])).
fof(f319,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = sum(X1,X0)) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : sum(X0,X1) = sum(X1,X0)),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)),
  inference(rectify,[],[f2])).
fof(f2,axiom,(
  ! [X2] : ! [X3] : sum(X2,X3) = sum(X3,X2)),
  file(unknown,unknown)).
fof(f326,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = X1 | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : (~less(X0,X1) | sum(X0,X1) = X1)),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : (less(X0,X1) => sum(X0,X1) = X1)),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => sum(X0,X1) = X1)),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X4] : ! [X5] : (less(X4,X5) => sum(X4,X5) = X5)),
  file(unknown,unknown)).
fof(f19769,plain,(
  ~cont(sK0,sum(sK0,sK1))),
  inference(subsumption_resolution,[],[f19754,f16491])).
fof(f16491,plain,(
  ( ! [X12,X13] : (cont(X12,X12) | ~cont(X12,X13)) )),
  inference(subsumption_resolution,[],[f16476,f328])).
fof(f328,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X49] : ! [X50] : (cont(X49,X50) => leq(X49,X50))),
  file(unknown,unknown)).
fof(f16476,plain,(
  ( ! [X12,X13] : (cont(X12,X12) | ~cont(X12,X13) | ~leq(X12,X13)) )),
  inference(resolution,[],[f3697,f2251])).
fof(f2251,plain,(
  ( ! [X0] : (zex(difference(X0,X0))) )),
  inference(resolution,[],[f2249,f349])).
fof(f349,plain,(
  ( ! [X0,X1] : (p(difference(X0,X1),X0) | zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X115] : ! [X116] : (~zex(difference(X115,X116)) => p(difference(X115,X116),X115))),
  file(unknown,unknown)).
fof(f2249,plain,(
  ( ! [X6,X7] : (~p(difference(X6,X7),X7)) )),
  inference(subsumption_resolution,[],[f2235,f439])).
fof(f439,plain,(
  ( ! [X2,X1] : (~p(X1,X2) | ~zex(X1)) )),
  inference(resolution,[],[f389,f321])).
fof(f321,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X140] : ! [X141] : (zex(X140) => (~cont(X141,X140) & ~cont(X140,X141)))),
  file(unknown,unknown)).
fof(f389,plain,(
  ( ! [X0,X1] : (cont(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X87] : ! [X88] : (p(X87,X88) <=> (cont(X87,X88) & eqdim(X87,X88)))),
  file(unknown,unknown)).
fof(f2235,plain,(
  ( ! [X6,X7] : (~p(difference(X6,X7),X7) | zex(difference(X6,X7))) )),
  inference(resolution,[],[f1046,f989])).
fof(f989,plain,(
  ( ! [X0] : (p(X0,X0) | zex(X0)) )),
  inference(duplicate_literal_removal,[],[f985])).
fof(f985,plain,(
  ( ! [X0] : (p(X0,X0) | zex(X0) | p(X0,X0) | zex(X0)) )),
  inference(resolution,[],[f345,f344])).
fof(f344,plain,(
  ( ! [X0,X1] : (po(X0,sK12(X1,X0)) | p(X0,X1) | zex(X0)) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1] : (zex(X0) | (po(X0,sK12(X1,X0)) & ~po(X1,sK12(X1,X0))) | p(X0,X1))),
  inference(skolemisation,[status(esa)],[f201])).
fof(f201,plain,(
  ! [X0,X1] : (zex(X0) | ? [X2] : (po(X0,X2) & ~po(X1,X2)) | p(X0,X1))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1] : ((zex(X0) | ? [X2] : (po(X0,X2) & ~po(X1,X2))) | p(X0,X1))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : ((~zex(X0) & ! [X2] : (po(X0,X2) => po(X1,X2))) => p(X0,X1))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1] : ((~zex(X0) & ! [X2] : (po(X0,X2) => po(X1,X2))) => p(X0,X1))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X27] : ! [X28] : ((~zex(X27) & ! [X29] : (po(X27,X29) => po(X28,X29))) => p(X27,X28))),
  file(unknown,unknown)).
fof(f345,plain,(
  ( ! [X0,X1] : (~po(X1,sK12(X1,X0)) | p(X0,X1) | zex(X0)) )),
  inference(cnf_transformation,[],[f261])).
fof(f1046,plain,(
  ( ! [X2,X0,X1] : (~p(X0,difference(X1,X2)) | ~p(X0,X2)) )),
  inference(resolution,[],[f1045,f388])).
fof(f388,plain,(
  ( ! [X0,X3,X1] : (po(X0,X1) | ~p(X3,X1) | ~p(X3,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1] : ((~po(X0,X1) | (p(sK21(X1,X0),X0) & p(sK21(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(skolemisation,[status(esa)],[f291])).
fof(f291,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(rectify,[],[f290])).
fof(f290,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X2] : (~p(X2,X0) | ~p(X2,X1)) | po(X0,X1)))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0] : ! [X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X91] : ! [X92] : (po(X91,X92) <=> ? [X93] : (p(X93,X91) & p(X93,X92)))),
  file(unknown,unknown)).
fof(f1045,plain,(
  ( ! [X0,X1] : (~po(X1,difference(X0,X1))) )),
  inference(subsumption_resolution,[],[f1044,f734])).
fof(f734,plain,(
  ( ! [X4,X5] : (~po(X4,X5) | ~zex(X5)) )),
  inference(resolution,[],[f387,f440])).
fof(f440,plain,(
  ( ! [X4,X3] : (~p(X3,X4) | ~zex(X4)) )),
  inference(resolution,[],[f389,f320])).
fof(f320,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f387,plain,(
  ( ! [X0,X1] : (p(sK21(X1,X0),X1) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f292])).
fof(f1044,plain,(
  ( ! [X0,X1] : (~po(X1,difference(X0,X1)) | zex(difference(X0,X1))) )),
  inference(subsumption_resolution,[],[f341,f723])).
fof(f723,plain,(
  ( ! [X4,X5] : (~po(X4,X5) | ~zex(X4)) )),
  inference(resolution,[],[f386,f440])).
fof(f386,plain,(
  ( ! [X0,X1] : (p(sK21(X1,X0),X0) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f292])).
fof(f341,plain,(
  ( ! [X0,X1] : (~po(X1,difference(X0,X1)) | zex(difference(X0,X1)) | zex(X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : (zex(X1) | zex(difference(X0,X1)) | ~po(X1,difference(X0,X1)))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : ((zex(X1) | zex(difference(X0,X1))) | ~po(X1,difference(X0,X1)))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : ((~zex(X1) & ~zex(difference(X0,X1))) => ~po(X1,difference(X0,X1)))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : ! [X1] : ((~zex(X1) & ~zex(difference(X0,X1))) => ~po(X1,difference(X0,X1)))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X121] : ! [X122] : ((~zex(X122) & ~zex(difference(X121,X122))) => ~po(X122,difference(X121,X122)))),
  file(unknown,unknown)).
fof(f3697,plain,(
  ( ! [X4,X5,X3] : (~zex(difference(X3,X4)) | cont(X3,X4) | ~cont(X3,X5) | ~leq(X4,X5)) )),
  inference(resolution,[],[f1387,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (cont(X2,sum(X0,X1)) | ~cont(X2,X1) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X1) | cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X1)) | cont(X2,sum(X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & cont(X2,X1)) => cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & cont(X2,X1)) => cont(X2,sum(X0,X1)))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X6] : ! [X7] : ! [X8] : ((leq(X6,X7) & cont(X8,X7)) => cont(X8,sum(X6,X7)))),
  file(unknown,unknown)).
fof(f1387,plain,(
  ( ! [X6,X7,X5] : (~cont(X5,sum(X6,X7)) | cont(X5,X6) | ~zex(difference(X5,X6))) )),
  inference(resolution,[],[f407,f321])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : (~cont(X2,sum(X0,X1)) | cont(X2,X0) | cont(difference(X2,X0),X1))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((~cont(X2,sum(X0,X1)) | cont(X2,X0)) | cont(difference(X2,X0),X1))),
  inference(ennf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X9] : ! [X10] : ! [X11] : ((cont(X11,sum(X9,X10)) & ~cont(X11,X9)) => cont(difference(X11,X9),X10))),
  file(unknown,unknown)).
fof(f19754,plain,(
  ~cont(sK0,sK0) | ~cont(sK0,sum(sK0,sK1))),
  inference(resolution,[],[f19752,f2439])).
fof(f2439,plain,(
  less(sK0,intersection(sK0,sum(sK0,sK1))) | ~cont(sK0,sK0) | ~cont(sK0,sum(sK0,sK1))),
  inference(subsumption_resolution,[],[f2426,f404])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))),
  inference(ennf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X37] : ! [X38] : ! [X39] : ((cont(X39,X37) & cont(X39,X38)) => leq(X39,intersection(X37,X38)))),
  file(unknown,unknown)).
fof(f2426,plain,(
  ~cont(sK0,sum(sK0,sK1)) | ~cont(sK0,sK0) | less(sK0,intersection(sK0,sum(sK0,sK1))) | ~leq(sK0,intersection(sK0,sum(sK0,sK1)))),
  inference(resolution,[],[f1805,f377])).
fof(f377,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | less(X0,X1) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X78] : ! [X79] : (leq(X78,X79) <=> (less(X78,X79) | eqdim(X78,X79)))),
  file(unknown,unknown)).
fof(f1805,plain,(
  ~eqdim(sK0,intersection(sK0,sum(sK0,sK1))) | ~cont(sK0,sum(sK0,sK1)) | ~cont(sK0,sK0)),
  inference(resolution,[],[f408,f303])).
fof(f303,plain,(
  ~p(sK0,intersection(sK0,sum(sK0,sK1)))),
  inference(cnf_transformation,[],[f238])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (p(X2,intersection(X0,X1)) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0)) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X40] : ! [X41] : ! [X42] : ((cont(X42,X40) & cont(X42,X41) & eqdim(X42,intersection(X40,X41))) <=> p(X42,intersection(X40,X41)))),
  file(unknown,unknown)).
fof(f19752,plain,(
  ( ! [X2] : (~less(sK0,intersection(sK0,X2))) )),
  inference(subsumption_resolution,[],[f19744,f301])).
fof(f19744,plain,(
  ( ! [X2] : (~less(sK0,intersection(sK0,X2)) | ~less(sK1,sK0)) )),
  inference(superposition,[],[f19708,f630])).
fof(f19708,plain,(
  ( ! [X8] : (~less(sK0,intersection(sum(sK0,sK1),X8))) )),
  inference(subsumption_resolution,[],[f19641,f1329])).
fof(f1329,plain,(
  ( ! [X1] : (~less(sK0,X1) | ~zex(X1)) )),
  inference(subsumption_resolution,[],[f1324,f492])).
fof(f492,plain,(
  ( ! [X0] : (~less(sK0,X0) | ~zex(X0) | ~zex(sK1)) )),
  inference(superposition,[],[f429,f461])).
fof(f461,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0)) )),
  inference(resolution,[],[f414,f423])).
fof(f423,plain,(
  ( ! [X0,X1] : (~sQ1_eqProxy(X0,X1) | X0 = X1) )),
  introduced(equality_proxy_axiom,[])).
fof(f414,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X0,X1) | ~zex(X1) | ~zex(X0)) )),
  inference(equality_proxy_replacement,[],[f329])).
fof(f329,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X71] : ! [X72] : ((zex(X71) & zex(X72)) => X71 = X72)),
  file(unknown,unknown)).
fof(f429,plain,(
  ~less(sK0,sK1)),
  inference(resolution,[],[f327,f301])).
fof(f327,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : (~less(X0,X1) | ~less(X1,X0))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X66] : ! [X67] : (less(X66,X67) => ~less(X67,X66))),
  file(unknown,unknown)).
fof(f1324,plain,(
  ( ! [X1] : (~less(sK0,X1) | ~zex(X1) | zex(sK1)) )),
  inference(resolution,[],[f1321,f458])).
fof(f458,plain,(
  ( ! [X4,X3] : (~less(X3,X4) | ~zex(X4) | zex(X3)) )),
  inference(resolution,[],[f330,f327])).
fof(f330,plain,(
  ( ! [X0,X1] : (less(X0,X1) | zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : (~zex(X0) | zex(X1) | less(X0,X1))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : ((~zex(X0) | zex(X1)) | less(X0,X1))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X73] : ! [X74] : ((zex(X73) & ~zex(X74)) => less(X73,X74))),
  file(unknown,unknown)).
fof(f1321,plain,(
  ( ! [X1] : (less(sK1,X1) | ~less(sK0,X1)) )),
  inference(subsumption_resolution,[],[f913,f907])).
fof(f907,plain,(
  ( ! [X1] : (~less(sK0,X1) | ~less(X1,sK1)) )),
  inference(resolution,[],[f903,f378])).
fof(f378,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f284])).
fof(f903,plain,(
  ( ! [X14] : (~leq(X14,sK1) | ~less(sK0,X14)) )),
  inference(resolution,[],[f402,f429])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (less(X0,X2) | ~leq(X1,X2) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : (~less(X0,X1) | ~leq(X1,X2) | less(X0,X2))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((~less(X0,X1) | ~leq(X1,X2)) | less(X0,X2))),
  inference(ennf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X68] : ! [X69] : ! [X70] : ((less(X68,X69) & leq(X69,X70)) => less(X68,X70))),
  file(unknown,unknown)).
fof(f913,plain,(
  ( ! [X1] : (~less(sK0,X1) | less(sK1,X1) | less(X1,sK1)) )),
  inference(resolution,[],[f906,f376])).
fof(f376,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | less(X1,X0) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X76] : ! [X77] : (eqdim(X76,X77) <=> (~less(X76,X77) & ~less(X77,X76)))),
  file(unknown,unknown)).
fof(f906,plain,(
  ( ! [X0] : (~eqdim(X0,sK1) | ~less(sK0,X0)) )),
  inference(resolution,[],[f903,f379])).
fof(f379,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f284])).
fof(f19641,plain,(
  ( ! [X8] : (~less(sK0,intersection(sum(sK0,sK1),X8)) | zex(intersection(sum(sK0,sK1),X8))) )),
  inference(resolution,[],[f19450,f348])).
fof(f348,plain,(
  ( ! [X0,X1] : (cont(intersection(X0,X1),X0) | zex(intersection(X0,X1))) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X35] : ! [X36] : (~zex(intersection(X35,X36)) => cont(intersection(X35,X36),X35))),
  file(unknown,unknown)).
fof(f19450,plain,(
  ( ! [X2] : (~cont(X2,sum(sK0,sK1)) | ~less(sK0,X2)) )),
  inference(resolution,[],[f19377,f328])).
fof(f19377,plain,(
  ( ! [X2] : (~leq(X2,sum(sK0,sK1)) | ~less(sK0,X2)) )),
  inference(resolution,[],[f19360,f402])).
fof(f19360,plain,(
  ~less(sK0,sum(sK0,sK1))),
  inference(factoring,[],[f19238])).
fof(f19238,plain,(
  ( ! [X142] : (~less(sK0,sum(X142,sK1)) | ~less(X142,sum(X142,sK1))) )),
  inference(subsumption_resolution,[],[f19191,f1329])).
fof(f19191,plain,(
  ( ! [X142] : (~less(X142,sum(X142,sK1)) | zex(sum(X142,sK1)) | ~less(sK0,sum(X142,sK1))) )),
  inference(resolution,[],[f2902,f908])).
fof(f908,plain,(
  ( ! [X2] : (~cont(X2,sK1) | ~less(sK0,X2)) )),
  inference(resolution,[],[f903,f328])).
fof(f2902,plain,(
  ( ! [X6,X7] : (cont(sum(X6,X7),X7) | ~less(X6,sum(X6,X7)) | zex(sum(X6,X7))) )),
  inference(resolution,[],[f1397,f316])).
fof(f1397,plain,(
  ( ! [X2,X0,X1] : (~cont(X0,sum(X1,X2)) | cont(X0,X2) | ~less(X1,X0)) )),
  inference(subsumption_resolution,[],[f1394,f621])).
fof(f621,plain,(
  ( ! [X6,X7] : (~cont(X6,X7) | ~less(X7,X6)) )),
  inference(subsumption_resolution,[],[f619,f321])).
fof(f619,plain,(
  ( ! [X6,X7] : (zex(X6) | ~cont(X6,X7) | ~less(X7,X6)) )),
  inference(superposition,[],[f394,f325])).
fof(f325,plain,(
  ( ! [X0,X1] : (difference(X0,X1) = X0 | ~less(X1,X0)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : (~less(X1,X0) | difference(X0,X1) = X0)),
  inference(ennf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X14] : ! [X15] : (less(X15,X14) => difference(X14,X15) = X14)),
  file(unknown,unknown)).
fof(f394,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X25] : ! [X26] : (zex(difference(X25,X26)) <=> (zex(X25) | cont(X25,X26)))),
  file(unknown,unknown)).
fof(f1394,plain,(
  ( ! [X2,X0,X1] : (cont(X0,X2) | cont(X0,X1) | ~cont(X0,sum(X1,X2)) | ~less(X1,X0)) )),
  inference(superposition,[],[f407,f325])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1276
Passive clauses: 11260
Generated clauses: 19558
Final active clauses: 1276
Final passive clauses: 9696
Input formulas: 58
Initial clauses: 115

Selected by SInE selection: 58
SInE iterations: 1

Duplicate literals: 245
Fw subsumption resolutions: 1119
Fw demodulations: 23

Simple tautologies: 345
Equational tautologies: 12
Forward subsumptions: 6857

Binary resolution: 10186
Factoring: 47
Forward superposition: 6739
Backward superposition: 1102
Self superposition: 2

Unique components: 3

Memory used: 4221KB
Time elapsed: 0.290 s
------------------------------
% Success in time 0.749 s
