Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2051
Passive clauses: 18509
Generated clauses: 35201
Input formulas: 50
Initial clauses: 106

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 170
Fw subsumption resolutions: 1330

Simple tautologies: 80
Forward subsumptions: 14005

Binary resolution: 32221
Factoring: 41

Backtracking splits: 1067
Backtracking splits refuted: 158
Backtracking splits refuted at zero level: 44

Memory used: 6652KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 3418
Passive clauses: 43812
Generated clauses: 92349
Input formulas: 50
Initial clauses: 105

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 3328
Fw subsumption resolutions: 6887

Simple tautologies: 746
Equational tautologies: 30
Forward subsumptions: 34091

Binary resolution: 48449
Factoring: 180
Forward superposition: 17776
Backward superposition: 16054
Self superposition: 5

Unique components: 4

Memory used: 21236KB
Time elapsed: 1.400 s
------------------------------
% remaining time: 1181 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fsr=off:fde=none:gsp=input_only:lcm=reverse:nwc=10:ptb=off:ssec=off:stl=60:sos=on:sgo=on:sio=off:spl=backtracking:sfv=off:sp=reverse_arity_23 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 10271
Passive clauses: 46949
Generated clauses: 179566
Input formulas: 50
Initial clauses: 103

Duplicate literals: 400
Trivial inequalities: 226
Bw subsumption resolutions: 11736
Fw demodulations: 36324
Condensations: 50

Simple tautologies: 441
Equational tautologies: 895
Forward subsumptions: 84808

Binary resolution: 92053
Factoring: 18
Forward superposition: 2688
Backward superposition: 9717

Backtracking splits: 4082
Backtracking splits refuted: 2173
Backtracking splits refuted at zero level: 32

Memory used: 36459KB
Time elapsed: 3.300 s
------------------------------
% remaining time: 1148 next slice time: 95
dis+2_4_bs=off:ep=on:nwc=1.5:nicw=on:ptb=off:ssec=off:sac=on:sio=off:spl=backtracking_70 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f12417,plain,(
  $false),
  inference(subsumption_resolution,[],[f12416,f1465])).
fof(f1465,plain,(
  p(sK1,sK1)),
  inference(subsumption_resolution,[],[f1464,f262])).
fof(f262,plain,(
  ~zex(sK1)),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ~zex(sK0) & ~zex(sK1) & ~cont(sK1,sK0) & ! [X2] : (~p(X2,sK1) | ~less(intersection(X2,sK0),X2))),
  inference(skolemisation,[status(esa)],[f143])).
fof(f143,plain,(
  ? [X0,X1] : (~zex(X0) & ~zex(X1) & ~cont(X1,X0) & ! [X2] : (~p(X2,X1) | ~less(intersection(X2,X0),X2)))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ? [X0,X1] : ((~zex(X0) & ~zex(X1) & ~cont(X1,X0)) & ! [X2] : (~p(X2,X1) | ~less(intersection(X2,X0),X2)))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ~! [X0,X1] : ((~zex(X0) & ~zex(X1) & ~cont(X1,X0)) => ? [X2] : (p(X2,X1) & less(intersection(X2,X0),X2)))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : ((~zex(X0) & ~zex(X1) & ~cont(X1,X0)) => ? [X2] : (p(X2,X1) & less(intersection(X2,X0),X2)))),
  file(unknown,unknown)).
fof(f1464,plain,(
  zex(sK1) | p(sK1,sK1)),
  inference(forward_demodulation,[],[f1445,f1366])).
fof(f1366,plain,(
  difference(sK1,sK20) = sK1),
  inference(resolution,[],[f1353,f262])).
fof(f1353,plain,(
  ( ! [X18] : (zex(X18) | difference(X18,sK20) = X18) )),
  inference(resolution,[],[f440,f363])).
fof(f363,plain,(
  zex(sK20)),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  zex(sK20)),
  inference(skolemisation,[status(esa)],[f141])).
fof(f141,plain,(
  ? [X0] : zex(X0)),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ? [X61] : zex(X61)),
  file(unknown,unknown)).
fof(f440,plain,(
  ( ! [X0,X1] : (~zex(X1) | zex(X0) | difference(X0,X1) = X0) )),
  inference(resolution,[],[f285,f289])).
fof(f289,plain,(
  ( ! [X0,X1] : (less(X0,X1) | zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : (~zex(X0) | zex(X1) | less(X0,X1))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : ((~zex(X0) | zex(X1)) | less(X0,X1))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X70] : ! [X71] : ((zex(X70) & ~zex(X71)) => less(X70,X71))),
  file(unknown,unknown)).
fof(f285,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | difference(X0,X1) = X0) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1] : (~less(X1,X0) | difference(X0,X1) = X0)),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X17] : ! [X18] : (less(X18,X17) => difference(X17,X18) = X17)),
  file(unknown,unknown)).
fof(f1445,plain,(
  p(sK1,sK1) | zex(difference(sK1,sK20))),
  inference(superposition,[],[f304,f1366])).
fof(f304,plain,(
  ( ! [X0,X1] : (p(difference(X0,X1),X0) | zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ! [X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X30] : ! [X31] : (~zex(difference(X30,X31)) => p(difference(X30,X31),X30))),
  file(unknown,unknown)).
fof(f12416,plain,(
  ~p(sK1,sK1)),
  inference(subsumption_resolution,[],[f12390,f6310])).
fof(f6310,plain,(
  cont(intersection(sK1,sK0),sK1)),
  inference(resolution,[],[f6264,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~p(X2,intersection(X0,X1)) | cont(X2,X0)) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X56] : ! [X57] : ! [X58] : ((cont(X58,X56) & cont(X58,X57) & eqdim(X58,intersection(X56,X57))) <=> p(X58,intersection(X56,X57)))),
  file(unknown,unknown)).
fof(f6264,plain,(
  p(intersection(sK1,sK0),intersection(sK1,sK0))),
  inference(subsumption_resolution,[],[f6263,f5758])).
fof(f5758,plain,(
  ~zex(intersection(sK1,sK0))),
  inference(subsumption_resolution,[],[f5739,f1465])).
fof(f5739,plain,(
  ~zex(intersection(sK1,sK0)) | ~p(sK1,sK1)),
  inference(resolution,[],[f5660,f264])).
fof(f264,plain,(
  ( ! [X2] : (~less(intersection(X2,sK0),X2) | ~p(X2,sK1)) )),
  inference(cnf_transformation,[],[f202])).
fof(f5660,plain,(
  ( ! [X36] : (less(X36,sK1) | ~zex(X36)) )),
  inference(resolution,[],[f1937,f1493])).
fof(f1493,plain,(
  cont(sK1,sK1)),
  inference(resolution,[],[f1465,f342])).
fof(f342,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X84] : ! [X85] : (p(X84,X85) <=> (cont(X84,X85) & eqdim(X84,X85)))),
  file(unknown,unknown)).
fof(f1937,plain,(
  ( ! [X4,X5,X3] : (~cont(X5,X4) | less(X3,X4) | ~zex(X3)) )),
  inference(subsumption_resolution,[],[f1927,f281])).
fof(f281,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X119] : ! [X120] : (zex(X119) => (~cont(X120,X119) & ~cont(X119,X120)))),
  file(unknown,unknown)).
fof(f1927,plain,(
  ( ! [X4,X5,X3] : (less(X3,X4) | ~cont(X5,X4) | zex(X5) | ~zex(X3)) )),
  inference(resolution,[],[f654,f289])).
fof(f654,plain,(
  ( ! [X6,X8,X7] : (~less(X6,X8) | less(X6,X7) | ~cont(X8,X7)) )),
  inference(resolution,[],[f354,f287])).
fof(f287,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X59] : ! [X60] : (cont(X59,X60) => leq(X59,X60))),
  file(unknown,unknown)).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (~leq(X1,X2) | less(X0,X2) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : (~less(X0,X1) | ~leq(X1,X2) | less(X0,X2))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((~less(X0,X1) | ~leq(X1,X2)) | less(X0,X2))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X65] : ! [X66] : ! [X67] : ((less(X65,X66) & leq(X66,X67)) => less(X65,X67))),
  file(unknown,unknown)).
fof(f6263,plain,(
  zex(intersection(sK1,sK0)) | p(intersection(sK1,sK0),intersection(sK1,sK0))),
  inference(forward_demodulation,[],[f6199,f5765])).
fof(f5765,plain,(
  intersection(sK1,sK0) = difference(intersection(sK1,sK0),sK20)),
  inference(resolution,[],[f5758,f1353])).
fof(f6199,plain,(
  p(intersection(sK1,sK0),intersection(sK1,sK0)) | zex(difference(intersection(sK1,sK0),sK20))),
  inference(superposition,[],[f304,f5765])).
fof(f12390,plain,(
  ~cont(intersection(sK1,sK0),sK1) | ~p(sK1,sK1)),
  inference(resolution,[],[f4015,f10701])).
fof(f10701,plain,(
  ~p(intersection(sK1,sK0),sK1)),
  inference(backward_demodulation,[],[f10686,f6363])).
fof(f6363,plain,(
  ~p(intersection(sK1,sK0),difference(sK1,sK0))),
  inference(resolution,[],[f6281,f350])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (~po(X2,intersection(X0,X1)) | ~p(X2,difference(X0,X1))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | ~po(X2,intersection(X0,X1)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : (p(X2,difference(X0,X1)) => ~po(X2,intersection(X0,X1)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1] : ! [X2] : (p(X2,difference(X0,X1)) => ~po(X2,intersection(X0,X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X43] : ! [X44] : ! [X45] : (p(X45,difference(X43,X44)) => ~po(X45,intersection(X43,X44)))),
  file(unknown,unknown)).
fof(f6281,plain,(
  po(intersection(sK1,sK0),intersection(sK1,sK0))),
  inference(subsumption_resolution,[],[f6249,f5758])).
fof(f6249,plain,(
  zex(intersection(sK1,sK0)) | po(intersection(sK1,sK0),intersection(sK1,sK0))),
  inference(superposition,[],[f4001,f5765])).
fof(f4001,plain,(
  ( ! [X0,X1] : (po(X0,X0) | zex(difference(X0,X1))) )),
  inference(resolution,[],[f3961,f304])).
fof(f3961,plain,(
  ( ! [X4,X2,X3] : (~p(difference(X3,X4),X2) | po(X2,X3)) )),
  inference(subsumption_resolution,[],[f576,f1161])).
fof(f1161,plain,(
  ( ! [X12,X13] : (~p(X12,X13) | ~zex(X12)) )),
  inference(resolution,[],[f1139,f281])).
fof(f1139,plain,(
  ( ! [X15,X16] : (cont(X15,difference(X16,sK20)) | ~p(X15,X16)) )),
  inference(resolution,[],[f1085,f342])).
fof(f1085,plain,(
  ( ! [X2,X3] : (p(X3,difference(X2,sK20)) | ~p(X3,X2)) )),
  inference(subsumption_resolution,[],[f1084,f800])).
fof(f800,plain,(
  ( ! [X0] : (~po(X0,sK20)) )),
  inference(superposition,[],[f788,f417])).
fof(f417,plain,(
  ( ! [X0] : (difference(sK20,X0) = sK20) )),
  inference(resolution,[],[f405,f363])).
fof(f405,plain,(
  ( ! [X2,X3] : (~zex(X2) | difference(X2,X3) = sK20) )),
  inference(resolution,[],[f399,f346])).
fof(f346,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f255])).
fof(f255,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X28] : ! [X29] : (zex(difference(X28,X29)) <=> (zex(X28) | cont(X28,X29)))),
  file(unknown,unknown)).
fof(f399,plain,(
  ( ! [X0] : (~zex(X0) | sK20 = X0) )),
  inference(resolution,[],[f393,f363])).
fof(f393,plain,(
  ( ! [X0,X1] : (~zex(X1) | ~zex(X0) | X0 = X1) )),
  inference(resolution,[],[f364,f371])).
fof(f371,plain,(
  ( ! [X0,X1] : (~sQ1_eqProxy(X0,X1) | X0 = X1) )),
  introduced(equality_proxy_axiom,[])).
fof(f364,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X0,X1) | ~zex(X1) | ~zex(X0)) )),
  inference(equality_proxy_replacement,[],[f288])).
fof(f288,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X68] : ! [X69] : ((zex(X68) & zex(X69)) => X68 = X69)),
  file(unknown,unknown)).
fof(f788,plain,(
  ( ! [X0,X1] : (~po(X1,difference(X0,X1))) )),
  inference(subsumption_resolution,[],[f787,f584])).
fof(f584,plain,(
  ( ! [X2,X3] : (~po(X2,X3) | ~zex(X3)) )),
  inference(resolution,[],[f462,f280])).
fof(f280,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f149])).
fof(f462,plain,(
  ( ! [X4,X5] : (cont(sK17(X5,X4),X5) | ~po(X4,X5)) )),
  inference(resolution,[],[f334,f342])).
fof(f334,plain,(
  ( ! [X0,X1] : (p(sK17(X1,X0),X1) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1] : ((~po(X0,X1) | (p(sK17(X1,X0),X0) & p(sK17(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(skolemisation,[status(esa)],[f246])).
fof(f246,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(rectify,[],[f245])).
fof(f245,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X2] : (~p(X2,X0) | ~p(X2,X1)) | po(X0,X1)))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X88] : ! [X89] : (po(X88,X89) <=> ? [X90] : (p(X90,X88) & p(X90,X89)))),
  file(unknown,unknown)).
fof(f787,plain,(
  ( ! [X0,X1] : (~po(X1,difference(X0,X1)) | zex(difference(X0,X1))) )),
  inference(subsumption_resolution,[],[f298,f574])).
fof(f574,plain,(
  ( ! [X2,X3] : (~po(X2,X3) | ~zex(X2)) )),
  inference(resolution,[],[f459,f280])).
fof(f459,plain,(
  ( ! [X4,X5] : (cont(sK17(X5,X4),X4) | ~po(X4,X5)) )),
  inference(resolution,[],[f333,f342])).
fof(f333,plain,(
  ( ! [X0,X1] : (p(sK17(X1,X0),X0) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f247])).
fof(f298,plain,(
  ( ! [X0,X1] : (~po(X1,difference(X0,X1)) | zex(difference(X0,X1)) | zex(X1)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : (zex(X1) | zex(difference(X0,X1)) | ~po(X1,difference(X0,X1)))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1] : ((zex(X1) | zex(difference(X0,X1))) | ~po(X1,difference(X0,X1)))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ((~zex(X1) & ~zex(difference(X0,X1))) => ~po(X1,difference(X0,X1)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1] : ((~zex(X1) & ~zex(difference(X0,X1))) => ~po(X1,difference(X0,X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X36] : ! [X37] : ((~zex(X37) & ~zex(difference(X36,X37))) => ~po(X37,difference(X36,X37)))),
  file(unknown,unknown)).
fof(f1084,plain,(
  ( ! [X2,X3] : (po(X3,sK20) | p(X3,difference(X2,sK20)) | ~p(X3,X2)) )),
  inference(superposition,[],[f357,f862])).
fof(f862,plain,(
  ( ! [X15] : (intersection(X15,sK20) = sK20) )),
  inference(resolution,[],[f846,f363])).
fof(f846,plain,(
  ( ! [X0,X1] : (~zex(X1) | intersection(X0,X1) = sK20) )),
  inference(resolution,[],[f406,f466])).
fof(f466,plain,(
  ( ! [X2,X3] : (~c(X2,X3) | ~zex(X3)) )),
  inference(resolution,[],[f337,f280])).
fof(f337,plain,(
  ( ! [X0,X1] : (cont(sK18(X1,X0),X1) | ~c(X0,X1)) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1] : ((~c(X0,X1) | (cont(sK18(X1,X0),X0) & cont(sK18(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))),
  inference(skolemisation,[status(esa)],[f249])).
fof(f249,plain,(
  ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))),
  inference(rectify,[],[f248])).
fof(f248,plain,(
  ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | c(X0,X1)))),
  inference(nnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X103] : ! [X104] : (c(X103,X104) <=> ? [X105] : (cont(X105,X103) & cont(X105,X104)))),
  file(unknown,unknown)).
fof(f406,plain,(
  ( ! [X4,X5] : (c(X4,X5) | intersection(X4,X5) = sK20) )),
  inference(resolution,[],[f399,f301])).
fof(f301,plain,(
  ( ! [X0,X1] : (zex(intersection(X0,X1)) | c(X0,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : (c(X0,X1) | zex(intersection(X0,X1)))),
  inference(ennf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : (~c(X0,X1) => zex(intersection(X0,X1)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1] : (~c(X0,X1) => zex(intersection(X0,X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X49] : ! [X50] : (~c(X49,X50) => zex(intersection(X49,X50)))),
  file(unknown,unknown)).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (po(X2,intersection(X0,X1)) | p(X2,difference(X0,X1)) | ~p(X2,X0)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : (~p(X2,X0) | po(X2,intersection(X0,X1)) | p(X2,difference(X0,X1)))),
  inference(flattening,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((~p(X2,X0) | po(X2,intersection(X0,X1))) | p(X2,difference(X0,X1)))),
  inference(ennf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : ((p(X2,X0) & ~po(X2,intersection(X0,X1))) => p(X2,difference(X0,X1)))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((p(X2,X0) & ~po(X2,intersection(X0,X1))) => p(X2,difference(X0,X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X46] : ! [X47] : ! [X48] : ((p(X48,X46) & ~po(X48,intersection(X46,X47))) => p(X48,difference(X46,X47)))),
  file(unknown,unknown)).
fof(f576,plain,(
  ( ! [X4,X2,X3] : (po(X2,X3) | ~p(difference(X3,X4),X2) | zex(difference(X3,X4))) )),
  inference(resolution,[],[f335,f304])).
fof(f335,plain,(
  ( ! [X0,X3,X1] : (~p(X3,X1) | po(X0,X1) | ~p(X3,X0)) )),
  inference(cnf_transformation,[],[f247])).
fof(f10686,plain,(
  difference(sK1,sK0) = sK1),
  inference(resolution,[],[f10682,f285])).
fof(f10682,plain,(
  less(sK0,sK1)),
  inference(subsumption_resolution,[],[f10585,f10561])).
fof(f10561,plain,(
  ~less(sK1,sK0)),
  inference(resolution,[],[f10559,f331])).
fof(f331,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X75] : ! [X76] : (leq(X75,X76) <=> (less(X75,X76) | eqdim(X75,X76)))),
  file(unknown,unknown)).
fof(f10559,plain,(
  ~leq(sK1,sK0)),
  inference(backtracking_split_refutation,[],[f8646,f10253_D,f10556])).
fof(f10556,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f10555,f263])).
fof(f263,plain,(
  ~cont(sK1,sK0)),
  inference(cnf_transformation,[],[f202])).
fof(f10555,plain,(
  cont(sK1,sK0) | $spl32),
  inference(subsumption_resolution,[],[f10541,f262])).
fof(f10541,plain,(
  zex(sK1) | cont(sK1,sK0) | $spl32),
  inference(resolution,[],[f10540,f345])).
fof(f345,plain,(
  ( ! [X0,X1] : (~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f256])).
fof(f10540,plain,(
  zex(difference(sK1,sK0)) | $spl32),
  inference(resolution,[],[f10275,f277])).
fof(f277,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X113] : (~zex(X113) <=> cont(X113,X113))),
  file(unknown,unknown)).
fof(f10275,plain,(
  ~cont(difference(sK1,sK0),difference(sK1,sK0)) | $spl32),
  inference(resolution,[],[f10256,f4014])).
fof(f4014,plain,(
  ( ! [X0] : (p(X0,X0) | ~cont(X0,X0)) )),
  inference(resolution,[],[f4013,f279])).
fof(f279,plain,(
  ( ! [X0] : (~less(X0,X0)) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X62] : ~less(X62,X62)),
  file(unknown,unknown)).
fof(f4013,plain,(
  ( ! [X2,X3] : (less(X2,X3) | ~cont(X2,X3) | p(X2,X3)) )),
  inference(subsumption_resolution,[],[f630,f1809])).
fof(f1809,plain,(
  ( ! [X2,X3] : (~less(X3,X2) | ~cont(X2,X3)) )),
  inference(subsumption_resolution,[],[f1802,f286])).
fof(f286,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : (~less(X0,X1) | ~less(X1,X0))),
  inference(ennf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X63] : ! [X64] : (less(X63,X64) => ~less(X64,X63))),
  file(unknown,unknown)).
fof(f1802,plain,(
  ( ! [X2,X3] : (less(X2,X3) | ~cont(X2,X3) | ~less(X3,X2)) )),
  inference(resolution,[],[f571,f328])).
fof(f328,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | ~less(X1,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X73] : ! [X74] : (eqdim(X73,X74) <=> (~less(X73,X74) & ~less(X74,X73)))),
  file(unknown,unknown)).
fof(f571,plain,(
  ( ! [X4,X5] : (eqdim(X4,X5) | less(X4,X5) | ~cont(X4,X5)) )),
  inference(resolution,[],[f330,f287])).
fof(f330,plain,(
  ( ! [X0,X1] : (~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f244])).
fof(f630,plain,(
  ( ! [X2,X3] : (p(X2,X3) | ~cont(X2,X3) | less(X3,X2) | less(X2,X3)) )),
  inference(resolution,[],[f344,f329])).
fof(f329,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | less(X1,X0) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f344,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f254])).
fof(f10256,plain,(
  ( ! [X2] : (~p(X2,difference(sK1,sK0))) ) | $spl32),
  inference(resolution,[],[f10253,f1162])).
fof(f1162,plain,(
  ( ! [X14,X15] : (~zex(difference(X15,sK20)) | ~p(X14,X15)) )),
  inference(resolution,[],[f1139,f280])).
fof(f10253,plain,(
  ( ! [X0] : (zex(difference(difference(sK1,sK0),X0))) ) | $spl32),
  inference(cnf_transformation,[],[f10253_D])).
fof(f10253_D,plain,(
  ( ! [X0] : (zex(difference(difference(sK1,sK0),X0))) ) <=> ~$spl32),
  introduced(backtracking_splitting_component,[])).
fof(f8646,plain,(
  ( ! [X0] : (~leq(sK1,sK0) | zex(difference(difference(sK1,sK0),X0))) )),
  inference(resolution,[],[f6942,f523])).
fof(f523,plain,(
  ( ! [X6,X8,X7] : (p(difference(difference(X6,X7),X8),X6) | zex(difference(difference(X6,X7),X8))) )),
  inference(resolution,[],[f304,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (~p(X2,difference(X0,X1)) | p(X2,X0)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2] : (p(X2,difference(X0,X1)) => p(X2,X0))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : ! [X1] : ! [X2] : (p(X2,difference(X0,X1)) => p(X2,X0))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X40] : ! [X41] : ! [X42] : (p(X42,difference(X40,X41)) => p(X42,X40))),
  file(unknown,unknown)).
fof(f6942,plain,(
  ( ! [X0,X1] : (~p(difference(difference(X0,sK0),X1),sK1) | ~leq(X0,sK0)) )),
  inference(subsumption_resolution,[],[f6913,f1161])).
fof(f6913,plain,(
  ( ! [X0,X1] : (~leq(X0,sK0) | zex(difference(difference(X0,sK0),X1)) | ~p(difference(difference(X0,sK0),X1),sK1)) )),
  inference(resolution,[],[f1053,f264])).
fof(f1053,plain,(
  ( ! [X2,X0,X1] : (less(intersection(difference(difference(X0,X1),X2),X1),difference(difference(X0,X1),X2)) | ~leq(X0,X1) | zex(difference(difference(X0,X1),X2))) )),
  inference(resolution,[],[f353,f304])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~p(X2,difference(X0,X1)) | less(intersection(X2,X1),X2) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~p(X2,difference(X0,X1)) | less(intersection(X2,X1),X2))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~p(X2,difference(X0,X1))) | less(intersection(X2,X1),X2))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & p(X2,difference(X0,X1))) => less(intersection(X2,X1),X2))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & p(X2,difference(X0,X1))) => less(intersection(X2,X1),X2))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X25] : ! [X26] : ! [X27] : ((leq(X25,X26) & p(X27,difference(X25,X26))) => less(intersection(X27,X26),X27))),
  file(unknown,unknown)).
fof(f10585,plain,(
  less(sK0,sK1) | less(sK1,sK0)),
  inference(resolution,[],[f10560,f329])).
fof(f10560,plain,(
  ~eqdim(sK1,sK0)),
  inference(resolution,[],[f10559,f332])).
fof(f332,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f244])).
fof(f4015,plain,(
  ( ! [X1] : (p(intersection(X1,sK0),X1) | ~cont(intersection(X1,sK0),X1) | ~p(X1,sK1)) )),
  inference(resolution,[],[f4013,f264])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 904
Passive clauses: 4767
Generated clauses: 12466
Final active clauses: 664
Final passive clauses: 2808
Input formulas: 50
Initial clauses: 105

Duplicate literals: 229
Fw subsumption resolutions: 940
Fw demodulations: 1362
Bw demodulations: 32

Simple tautologies: 373
Equational tautologies: 26
Forward subsumptions: 5112

Binary resolution: 7456
Factoring: 3
Forward superposition: 912
Backward superposition: 1101
Self superposition: 1

Backtracking splits: 35
Backtracking splits refuted: 21
Backtracking splits refuted at zero level: 18

Memory used: 2174KB
Time elapsed: 0.156 s
------------------------------
% Success in time 5.295 s
