Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 29999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f809,plain,(
  $false),
  inference(subsumption_resolution,[],[f806,f413])).
fof(f413,plain,(
  atomicraycurve(sK0)),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  atomicraycurve(sK0) & ~icon(sK0)),
  inference(skolemisation,[status(esa)],[f229])).
fof(f229,plain,(
  ? [X0] : (atomicraycurve(X0) & ~icon(X0))),
  inference(ennf_transformation,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : (atomicraycurve(X0) => icon(X0))),
  file(unknown,unknown)).
fof(f806,plain,(
  ~atomicraycurve(sK0)),
  inference(resolution,[],[f802,f452])).
fof(f452,plain,(
  ( ! [X0,X1] : (~pp(X1,X0) | ~atomicraycurve(X0)) )),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0] : ((~atomicraycurve(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & (endpoint(sK10(X0),X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | sK10(X0) = X3)))) & (~curve(X0) | pp(sK11(X0),X0) | ! [X5] : (~endpoint(X5,X0) | (point(sK12(X5,X0)) & cont(sK12(X5,X0),X0) & sK12(X5,X0) != X5)) | atomicraycurve(X0)))),
  inference(skolemisation,[status(esa)],[f358])).
fof(f358,plain,(
  ! [X0] : ((~atomicraycurve(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (endpoint(X2,X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | X2 = X3)))) & (~curve(X0) | ? [X4] : pp(X4,X0) | ! [X5] : (~endpoint(X5,X0) | ? [X6] : (point(X6) & cont(X6,X0) & X5 != X6)) | atomicraycurve(X0)))),
  inference(rectify,[],[f357])).
fof(f357,plain,(
  ! [X0] : ((~atomicraycurve(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (endpoint(X2,X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | X2 = X3)))) & (~curve(X0) | ? [X1] : pp(X1,X0) | ! [X2] : (~endpoint(X2,X0) | ? [X3] : (point(X3) & cont(X3,X0) & X2 != X3)) | atomicraycurve(X0)))),
  inference(flattening,[],[f356])).
fof(f356,plain,(
  ! [X0] : ((~atomicraycurve(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (endpoint(X2,X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | X2 = X3)))) & ((~curve(X0) | ? [X1] : pp(X1,X0) | ! [X2] : (~endpoint(X2,X0) | ? [X3] : (point(X3) & cont(X3,X0) & X2 != X3))) | atomicraycurve(X0)))),
  inference(nnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0] : (atomicraycurve(X0) <=> (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (endpoint(X2,X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | X2 = X3))))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0] : (atomicraycurve(X0) <=> (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (endpoint(X2,X0) & ! [X3] : ((~point(X3) | ~cont(X3,X0)) | X2 = X3))))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : (atomicraycurve(X0) <=> (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (endpoint(X2,X0) & ! [X3] : ((point(X3) & cont(X3,X0)) => X2 = X3))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X67] : (atomicraycurve(X67) <=> (curve(X67) & ! [X68] : ~pp(X68,X67) & ? [X68] : (endpoint(X68,X67) & ! [X69] : ((point(X69) & cont(X69,X67)) => X68 = X69))))),
  file(unknown,unknown)).
fof(f802,plain,(
  pp(sK8(sK0),sK0)),
  inference(backtracking_split_refutation,[],[f792,f793_D,f801])).
fof(f801,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f798,f413])).
fof(f798,plain,(
  ~atomicraycurve(sK0) | $spl4),
  inference(resolution,[],[f794,f452])).
fof(f794,plain,(
  pp(sK6(sK0),sK0) | $spl4),
  inference(resolution,[],[f793,f435])).
fof(f435,plain,(
  ( ! [X0] : (con(X0) | pp(sK6(X0),X0)) )),
  inference(cnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0] : ((~con(X0) | ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1)))) & ((pp(sK6(X0),X0) & ~sc(sK6(X0),difference(X0,sK6(X0)))) | con(X0)))),
  inference(skolemisation,[status(esa)],[f341])).
fof(f341,plain,(
  ! [X0] : ((~con(X0) | ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1)))) & (? [X2] : (pp(X2,X0) & ~sc(X2,difference(X0,X2))) | con(X0)))),
  inference(rectify,[],[f340])).
fof(f340,plain,(
  ! [X0] : ((~con(X0) | ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1)))) & (? [X1] : (pp(X1,X0) & ~sc(X1,difference(X0,X1))) | con(X0)))),
  inference(nnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0] : (con(X0) <=> ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1))))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : (con(X0) <=> ! [X1] : (pp(X1,X0) => sc(X1,difference(X0,X1))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X95] : (con(X95) <=> ! [X96] : (pp(X96,X95) => sc(X96,difference(X95,X96))))),
  file(unknown,unknown)).
fof(f793,plain,(
  ~con(sK0) | $spl4),
  inference(cnf_transformation,[],[f793_D])).
fof(f793_D,plain,(
  ~con(sK0) <=> ~$spl4),
  introduced(backtracking_splitting_component,[])).
fof(f792,plain,(
  pp(sK8(sK0),sK0) | ~con(sK0)),
  inference(resolution,[],[f443,f414])).
fof(f414,plain,(
  ~icon(sK0)),
  inference(cnf_transformation,[],[f324])).
fof(f443,plain,(
  ( ! [X0] : (icon(X0) | pp(sK8(X0),X0) | ~con(X0)) )),
  inference(cnf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & (~con(X0) | (pp(sK8(X0),X0) & ~covers(X0,intersection(sK8(X0),difference(X0,sK8(X0))))) | icon(X0)))),
  inference(skolemisation,[status(esa)],[f348])).
fof(f348,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & (~con(X0) | ? [X2] : (pp(X2,X0) & ~covers(X0,intersection(X2,difference(X0,X2)))) | icon(X0)))),
  inference(rectify,[],[f347])).
fof(f347,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & (~con(X0) | ? [X1] : (pp(X1,X0) & ~covers(X0,intersection(X1,difference(X0,X1)))) | icon(X0)))),
  inference(flattening,[],[f346])).
fof(f346,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & ((~con(X0) | ? [X1] : (pp(X1,X0) & ~covers(X0,intersection(X1,difference(X0,X1))))) | icon(X0)))),
  inference(nnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0] : (icon(X0) <=> (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1))))))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0] : (icon(X0) <=> (con(X0) & ! [X1] : (pp(X1,X0) => covers(X0,intersection(X1,difference(X0,X1))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X55] : (icon(X55) <=> (con(X55) & ! [X56] : (pp(X56,X55) => covers(X55,intersection(X56,difference(X55,X56))))))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 156
Passive clauses: 312
Generated clauses: 385
Final active clauses: 154
Final passive clauses: 140
Input formulas: 84
Initial clauses: 175

Selected by SInE selection: 84
SInE iterations: 1

Duplicate literals: 1
Fw subsumption resolutions: 3

Simple tautologies: 8
Forward subsumptions: 65

Binary resolution: 199
Factoring: 1

Backtracking splits: 4
Backtracking splits refuted: 1
Backtracking splits refuted at zero level: 1

Memory used: 383KB
Time elapsed: 0.044 s
------------------------------
% Success in time 0.082 s
============================= vampire_lin32 ================================
execution finished: Mon Oct 07 16:06:14 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_mereotopology_codib/theorems/conversions/atomic_ray_curve_theorems_goal3.all.tptp
============================ end of footer ===========================
