Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 29999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2486
Passive clauses: 18811
Generated clauses: 45836
Input formulas: 68
Initial clauses: 139

Selected by SInE selection: 68
SInE iterations: 1

Duplicate literals: 167
Fw subsumption resolutions: 643

Simple tautologies: 137
Forward subsumptions: 25940

Binary resolution: 43994
Factoring: 98

Backtracking splits: 702
Backtracking splits refuted: 47
Backtracking splits refuted at zero level: 37

Memory used: 6140KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 29995 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 4385
Passive clauses: 33189
Generated clauses: 72524
Input formulas: 68
Initial clauses: 132

Unused predicate definitions: 4
Selected by SInE selection: 68
SInE iterations: 1

Duplicate literals: 2803
Fw subsumption resolutions: 4905

Simple tautologies: 229
Equational tautologies: 10
Forward subsumptions: 32988

Binary resolution: 46090
Factoring: 317
Forward superposition: 16256
Backward superposition: 2389

Unique components: 7

Memory used: 10874KB
Time elapsed: 1.400 s
------------------------------
% remaining time: 29981 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fsr=off:fde=none:gsp=input_only:lcm=reverse:nwc=10:ptb=off:ssec=off:stl=60:sos=on:sgo=on:sio=off:spl=backtracking:sfv=off:sp=reverse_arity_23 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 22941
Passive clauses: 79725
Generated clauses: 361932
Input formulas: 68
Initial clauses: 131

Unused predicate definitions: 4

Duplicate literals: 2840
Trivial inequalities: 12667
Bw subsumption resolutions: 10476
Fw demodulations: 81364
Condensations: 6

Simple tautologies: 2644
Equational tautologies: 1669
Forward subsumptions: 141940

Binary resolution: 129424
Factoring: 10
Forward superposition: 1403
Backward superposition: 21357

Backtracking splits: 16182
Backtracking splits refuted: 5104
Backtracking splits refuted at zero level: 283

Memory used: 40937KB
Time elapsed: 3.300 s
------------------------------
% remaining time: 29948 next slice time: 95
dis+2_4_bs=off:ep=on:nwc=1.5:nicw=on:ptb=off:ssec=off:sac=on:sio=off:spl=backtracking_70 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 10494
Passive clauses: 294365
Generated clauses: 584535
Input formulas: 68
Initial clauses: 132

Unused predicate definitions: 4

Duplicate literals: 12512
Fw subsumption resolutions: 29549
Fw demodulations: 2462
Bw demodulations: 660

Simple tautologies: 2129
Equational tautologies: 84
Forward subsumptions: 250373

Binary resolution: 533560
Factoring: 62
Forward superposition: 4242
Backward superposition: 4125

Backtracking splits: 152
Backtracking splits refuted: 47
Backtracking splits refuted at zero level: 37

Memory used: 97482KB
Time elapsed: 9.600 s
------------------------------
% remaining time: 29852 next slice time: 4
lrs-1003_12_bs=unit_only:bsr=on:bms=on:cond=fast:drc=off:ep=on:fde=none:lcm=predicate:nwc=10:nicw=on:ssec=off:stl=60:sgo=on:sio=off:spo=on:sfv=off:sp=occurrence:updr=off_2 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 1337
Passive clauses: 4066
Generated clauses: 14499
Input formulas: 68
Initial clauses: 138

Duplicate literals: 59
Fw subsumption resolutions: 282
Bw subsumption resolutions: 22
Condensations: 954

Simple tautologies: 38
Equational tautologies: 12
Forward subsumptions: 9079
Backward subsumptions: 8
Subsumed by BDD marking: 17

Binary resolution: 12857
Factoring: 205

Split clauses: 2
Split components: 4
Unique components: 2849
Introduced splitting names: 1
BDD propositional clauses: 1

Memory used: 3709KB
Time elapsed: 0.500 s
------------------------------
% remaining time: 29847 next slice time: 22
lrs+1010_50_bs=off:drc=off:flr=on:lcm=predicate:nwc=1.2:stl=60:sgo=on:sfv=off_16 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 5422
Passive clauses: 14408
Generated clauses: 270640
Input formulas: 68
Initial clauses: 130

Unused predicate definitions: 4

Duplicate literals: 15384
Fw subsumption resolutions: 20509
Fw literal rewrites: 21905

Simple tautologies: 3621
Equational tautologies: 4
Forward subsumptions: 195188

Binary resolution: 213216
Factoring: 7
Equality resolution: 1

Unique components: 7

Memory used: 12665KB
Time elapsed: 2.300 s
------------------------------
% remaining time: 29823 next slice time: 41
dis-1010_5_bd=off:bs=off:cond=fast:ep=on:fde=none:lcm=predicate:nwc=1.3:nicw=on:ptb=off:ssec=off:sos=on:sac=on:sgo=on:sio=off:spl=backtracking:updr=off_30 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 14266
Passive clauses: 75259
Generated clauses: 289149
Input formulas: 68
Initial clauses: 138

Duplicate literals: 9590
Fw subsumption resolutions: 7040
Fw demodulations: 5434
Condensations: 296

Simple tautologies: 189
Forward subsumptions: 196541

Binary resolution: 254904
Factoring: 1186
Backward superposition: 10456

Backtracking splits: 134
Backtracking splits refuted: 39
Backtracking splits refuted at zero level: 8

Memory used: 48229KB
Time elapsed: 4.200 s
------------------------------
% remaining time: 29781 next slice time: 237
lrs+1_2:1_bs=off:ep=on:gsp=input_only:lcm=predicate:nwc=1.1:nicw=on:ptb=off:ssec=off:stl=60:sos=on:sac=on:sgo=on:sio=off:spo=on:spl=backtracking:sp=occurrence:updr=off_175 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f18017,plain,(
  $false),
  inference(subsumption_resolution,[],[f18016,f17946])).
fof(f17946,plain,(
  eqdim(sK11(sK2,sK3),sK2)),
  inference(resolution,[],[f17937,f441])).
fof(f441,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f321])).
fof(f321,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X109] : ! [X110] : (p(X109,X110) <=> (cont(X109,X110) & eqdim(X109,X110)))),
  file(unknown,unknown)).
fof(f17937,plain,(
  p(sK11(sK2,sK3),sK2)),
  inference(backtracking_split_refutation,[],[f15103,f15189_D,f17936])).
fof(f17936,plain,(
  $false | $spl169),
  inference(subsumption_resolution,[],[f17920,f333])).
fof(f333,plain,(
  ~po(sK2,sK3)),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  curve(sK0) & (pp(sK2,sK0) & pp(sK3,sK0) & ~po(sK2,sK3) & endpoint(sK1,sK2) & endpoint(sK1,sK3)) & endpoint(sK1,sK0)),
  inference(skolemisation,[status(esa)],[f187])).
fof(f187,plain,(
  ? [X0,X1] : (curve(X0) & ? [X2,X3] : (pp(X2,X0) & pp(X3,X0) & ~po(X2,X3) & endpoint(X1,X2) & endpoint(X1,X3)) & endpoint(X1,X0))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ? [X0,X1] : ((curve(X0) & ? [X2,X3] : (pp(X2,X0) & pp(X3,X0) & ~po(X2,X3) & endpoint(X1,X2) & endpoint(X1,X3))) & endpoint(X1,X0))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ~! [X0,X1] : ((curve(X0) & ? [X2,X3] : (pp(X2,X0) & pp(X3,X0) & ~po(X2,X3) & endpoint(X1,X2) & endpoint(X1,X3))) => ~endpoint(X1,X0))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : ((curve(X0) & ? [X2] : ? [X3] : (pp(X2,X0) & pp(X3,X0) & ~po(X2,X3) & endpoint(X1,X2) & endpoint(X1,X3))) => ~endpoint(X1,X0))),
  file(unknown,unknown)).
fof(f17920,plain,(
  po(sK2,sK3) | $spl169),
  inference(resolution,[],[f15470,f15193])).
fof(f15193,plain,(
  ( ! [X3] : (~p(sK11(sK2,sK3),X3) | po(X3,sK3)) ) | $spl169),
  inference(resolution,[],[f15189,f419])).
fof(f419,plain,(
  ( ! [X0,X3,X1] : (~p(X3,X1) | po(X0,X1) | ~p(X3,X0)) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1] : ((~po(X0,X1) | (p(sK15(X1,X0),X0) & p(sK15(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(skolemisation,[status(esa)],[f306])).
fof(f306,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(rectify,[],[f305])).
fof(f305,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X2] : (~p(X2,X0) | ~p(X2,X1)) | po(X0,X1)))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0] : ! [X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X114] : ! [X115] : (po(X114,X115) <=> ? [X116] : (p(X116,X114) & p(X116,X115)))),
  file(unknown,unknown)).
fof(f15470,plain,(
  p(sK11(sK2,sK3),sK2) | $spl169),
  inference(subsumption_resolution,[],[f15463,f15098])).
fof(f15098,plain,(
  cont(sK11(sK2,sK3),sK2)),
  inference(backtracking_split_refutation,[],[f857,f2593_D,f15096])).
fof(f15096,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f15076,f2593])).
fof(f15076,plain,(
  ~sc(sK3,sK2)),
  inference(subsumption_resolution,[],[f15047,f530])).
fof(f530,plain,(
  cont(sK1,sK3)),
  inference(resolution,[],[f485,f382])).
fof(f382,plain,(
  ( ! [X0,X1] : (~bcont(X0,X1) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1] : (~bcont(X0,X1) | (cont(X0,X1) & inc(X0,X1)))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : (bcont(X0,X1) => (cont(X0,X1) & inc(X0,X1)))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1] : (bcont(X0,X1) => (cont(X0,X1) & inc(X0,X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X31] : ! [X32] : (bcont(X31,X32) => (cont(X31,X32) & inc(X31,X32)))),
  file(unknown,unknown)).
fof(f485,plain,(
  bcont(sK1,sK3)),
  inference(resolution,[],[f335,f429])).
fof(f429,plain,(
  ( ! [X0,X1] : (~endpoint(X0,X1) | bcont(X0,X1)) )),
  inference(cnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1] : ((~endpoint(X0,X1) | (point(X0) & curve(X1) & bcont(X0,X1))) & (~point(X0) | ~curve(X1) | ~bcont(X0,X1) | endpoint(X0,X1)))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X0,X1] : ((~endpoint(X0,X1) | (point(X0) & curve(X1) & bcont(X0,X1))) & ((~point(X0) | ~curve(X1) | ~bcont(X0,X1)) | endpoint(X0,X1)))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : (endpoint(X0,X1) <=> (point(X0) & curve(X1) & bcont(X0,X1)))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0] : ! [X1] : (endpoint(X0,X1) <=> (point(X0) & curve(X1) & bcont(X0,X1)))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X125] : ! [X126] : (endpoint(X125,X126) <=> (point(X125) & curve(X126) & bcont(X125,X126)))),
  file(unknown,unknown)).
fof(f335,plain,(
  endpoint(sK1,sK3)),
  inference(cnf_transformation,[],[f264])).
fof(f15047,plain,(
  ~cont(sK1,sK3) | ~sc(sK3,sK2)),
  inference(resolution,[],[f7822,f479])).
fof(f479,plain,(
  p(sK3,sK0)),
  inference(resolution,[],[f332,f414])).
fof(f414,plain,(
  ( ! [X0,X1] : (~pp(X0,X1) | p(X0,X1)) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1] : ((~pp(X0,X1) | (p(X0,X1) & X0 != X1)) & (~p(X0,X1) | X0 = X1 | pp(X0,X1)))),
  inference(flattening,[],[f303])).
fof(f303,plain,(
  ! [X0,X1] : ((~pp(X0,X1) | (p(X0,X1) & X0 != X1)) & ((~p(X0,X1) | X0 = X1) | pp(X0,X1)))),
  inference(nnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1] : (pp(X0,X1) <=> (p(X0,X1) & X0 != X1))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : (pp(X0,X1) <=> (p(X0,X1) & ~X0 = X1))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0] : ! [X1] : (pp(X0,X1) <=> (p(X0,X1) & ~X0 = X1))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X129] : ! [X130] : (pp(X129,X130) <=> (p(X129,X130) & ~X129 = X130))),
  file(unknown,unknown)).
fof(f332,plain,(
  pp(sK3,sK0)),
  inference(cnf_transformation,[],[f264])).
fof(f7822,plain,(
  ( ! [X0] : (~p(X0,sK0) | ~cont(sK1,X0) | ~sc(X0,sK2)) )),
  inference(subsumption_resolution,[],[f7805,f527])).
fof(f527,plain,(
  cont(sK1,sK2)),
  inference(resolution,[],[f482,f382])).
fof(f482,plain,(
  bcont(sK1,sK2)),
  inference(resolution,[],[f334,f429])).
fof(f334,plain,(
  endpoint(sK1,sK2)),
  inference(cnf_transformation,[],[f264])).
fof(f7805,plain,(
  ( ! [X0] : (~cont(sK1,X0) | ~cont(sK1,sK2) | ~p(X0,sK0) | ~sc(X0,sK2)) )),
  inference(resolution,[],[f1089,f478])).
fof(f478,plain,(
  p(sK2,sK0)),
  inference(resolution,[],[f331,f414])).
fof(f331,plain,(
  pp(sK2,sK0)),
  inference(cnf_transformation,[],[f264])).
fof(f1089,plain,(
  ( ! [X0,X1] : (~p(X0,sK0) | ~cont(sK1,X1) | ~cont(sK1,X0) | ~p(X1,sK0) | ~sc(X1,X0)) )),
  inference(subsumption_resolution,[],[f1085,f488])).
fof(f488,plain,(
  bcont(sK1,sK0)),
  inference(resolution,[],[f336,f429])).
fof(f336,plain,(
  endpoint(sK1,sK0)),
  inference(cnf_transformation,[],[f264])).
fof(f1085,plain,(
  ( ! [X0,X1] : (~bcont(sK1,sK0) | ~cont(sK1,X0) | ~cont(sK1,X1) | ~p(X0,sK0) | ~p(X1,sK0) | ~sc(X1,X0)) )),
  inference(resolution,[],[f986,f462])).
fof(f462,plain,(
  ( ! [X2,X0,X3,X1] : (~covers(X3,X2) | ~bcont(X2,X3) | ~cont(X2,X1) | ~cont(X2,X0) | ~p(X1,X3) | ~p(X0,X3) | ~sc(X0,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2,X3] : (~sc(X0,X1) | ~p(X0,X3) | ~p(X1,X3) | ~cont(X2,X0) | ~cont(X2,X1) | ~covers(X3,X2) | ~bcont(X2,X3))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2,X3] : ((~sc(X0,X1) | ~p(X0,X3) | ~p(X1,X3) | ~cont(X2,X0) | ~cont(X2,X1) | ~covers(X3,X2)) | ~bcont(X2,X3))),
  inference(ennf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : ((sc(X0,X1) & p(X0,X3) & p(X1,X3) & cont(X2,X0) & cont(X2,X1) & covers(X3,X2)) => ~bcont(X2,X3))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((sc(X0,X1) & p(X0,X3) & p(X1,X3) & cont(X2,X0) & cont(X2,X1) & covers(X3,X2)) => ~bcont(X2,X3))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X37] : ! [X38] : ! [X39] : ! [X40] : ((sc(X37,X38) & p(X37,X40) & p(X38,X40) & cont(X39,X37) & cont(X39,X38) & covers(X40,X39)) => ~bcont(X39,X40))),
  file(unknown,unknown)).
fof(f986,plain,(
  covers(sK0,sK1)),
  inference(resolution,[],[f546,f330])).
fof(f330,plain,(
  curve(sK0)),
  inference(cnf_transformation,[],[f264])).
fof(f546,plain,(
  ( ! [X0] : (~curve(X0) | covers(X0,sK1)) )),
  inference(resolution,[],[f525,f356])).
fof(f356,plain,(
  ( ! [X0,X1] : (~pointregion(X1) | covers(X0,X1) | ~curve(X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0] : ((~curve(X0) | ! [X1] : ((~pointregion(X1) | covers(X0,X1)) & (~covers(X0,X1) | pointregion(X1)))) & (((pointregion(sK8(X0)) | covers(X0,sK8(X0))) & (~pointregion(sK8(X0)) | ~covers(X0,sK8(X0)))) | curve(X0)))),
  inference(skolemisation,[status(esa)],[f283])).
fof(f283,plain,(
  ! [X0] : ((~curve(X0) | ! [X1] : ((~pointregion(X1) | covers(X0,X1)) & (~covers(X0,X1) | pointregion(X1)))) & (? [X2] : ((pointregion(X2) | covers(X0,X2)) & (~pointregion(X2) | ~covers(X0,X2))) | curve(X0)))),
  inference(rectify,[],[f282])).
fof(f282,plain,(
  ! [X0] : ((~curve(X0) | ! [X1] : ((~pointregion(X1) | covers(X0,X1)) & (~covers(X0,X1) | pointregion(X1)))) & (? [X1] : ((pointregion(X1) | covers(X0,X1)) & (~pointregion(X1) | ~covers(X0,X1))) | curve(X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (curve(X0) <=> ! [X1] : (pointregion(X1) <=> covers(X0,X1)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X107] : (curve(X107) <=> ! [X108] : (pointregion(X108) <=> covers(X107,X108)))),
  file(unknown,unknown)).
fof(f525,plain,(
  pointregion(sK1)),
  inference(resolution,[],[f491,f349])).
fof(f349,plain,(
  ( ! [X0] : (~mindim(X0) | pointregion(X0)) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0] : ((~pointregion(X0) | mindim(X0)) & (~mindim(X0) | pointregion(X0)))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (pointregion(X0) <=> mindim(X0))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X86] : (pointregion(X86) <=> mindim(X86))),
  file(unknown,unknown)).
fof(f491,plain,(
  mindim(sK1)),
  inference(resolution,[],[f480,f354])).
fof(f354,plain,(
  ( ! [X0] : (~point(X0) | mindim(X0)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0] : ((~point(X0) | (min(X0) & mindim(X0))) & (~min(X0) | ~mindim(X0) | point(X0)))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0] : ((~point(X0) | (min(X0) & mindim(X0))) & ((~min(X0) | ~mindim(X0)) | point(X0)))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : (point(X0) <=> (min(X0) & mindim(X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X72] : (point(X72) <=> (min(X72) & mindim(X72)))),
  file(unknown,unknown)).
fof(f480,plain,(
  point(sK1)),
  inference(resolution,[],[f334,f427])).
fof(f427,plain,(
  ( ! [X0,X1] : (~endpoint(X0,X1) | point(X0)) )),
  inference(cnf_transformation,[],[f313])).
fof(f2593,plain,(
  sc(sK3,sK2) | $spl46),
  inference(cnf_transformation,[],[f2593_D])).
fof(f2593_D,plain,(
  sc(sK3,sK2) <=> ~$spl46),
  introduced(backtracking_splitting_component,[])).
fof(f857,plain,(
  cont(sK11(sK2,sK3),sK2) | sc(sK3,sK2)),
  inference(resolution,[],[f739,f397])).
fof(f397,plain,(
  ( ! [X0,X1] : (~c(X0,X1) | cont(sK11(X1,X0),X1) | sc(X0,X1)) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1] : (~c(X0,X1) | (cont(sK11(X1,X0),X0) & cont(sK11(X1,X0),X1) & (p(sK11(X1,X0),X0) | p(sK11(X1,X0),X1))) | sc(X0,X1))),
  inference(skolemisation,[status(esa)],[f228])).
fof(f228,plain,(
  ! [X0,X1] : (~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1) & (p(X2,X0) | p(X2,X1))) | sc(X0,X1))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1) & (p(X2,X0) | p(X2,X1)))) | sc(X0,X1))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : ((c(X0,X1) & ! [X2] : (~cont(X2,X0) | ~cont(X2,X1) | (~p(X2,X0) & ~p(X2,X1)))) => sc(X0,X1))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0] : ! [X1] : ((c(X0,X1) & ! [X2] : (~cont(X2,X0) | ~cont(X2,X1) | (~p(X2,X0) & ~p(X2,X1)))) => sc(X0,X1))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X60] : ! [X61] : ((c(X60,X61) & ! [X62] : (~cont(X62,X60) | ~cont(X62,X61) | (~p(X62,X60) & ~p(X62,X61)))) => sc(X60,X61))),
  file(unknown,unknown)).
fof(f739,plain,(
  c(sK3,sK2)),
  inference(resolution,[],[f578,f530])).
fof(f578,plain,(
  ( ! [X0] : (~cont(sK1,X0) | c(X0,sK2)) )),
  inference(resolution,[],[f527,f413])).
fof(f413,plain,(
  ( ! [X0,X3,X1] : (~cont(X3,X1) | c(X0,X1) | ~cont(X3,X0)) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1] : ((~c(X0,X1) | (cont(sK14(X1,X0),X0) & cont(sK14(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))),
  inference(skolemisation,[status(esa)],[f301])).
fof(f301,plain,(
  ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))),
  inference(rectify,[],[f300])).
fof(f300,plain,(
  ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | c(X0,X1)))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0] : ! [X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X49] : ! [X50] : (c(X49,X50) <=> ? [X51] : (cont(X51,X49) & cont(X51,X50)))),
  file(unknown,unknown)).
fof(f15463,plain,(
  p(sK11(sK2,sK3),sK2) | ~cont(sK11(sK2,sK3),sK2) | $spl169),
  inference(resolution,[],[f15293,f442])).
fof(f442,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f322])).
fof(f15293,plain,(
  eqdim(sK11(sK2,sK3),sK2) | $spl169),
  inference(resolution,[],[f15197,f1229])).
fof(f1229,plain,(
  ( ! [X2] : (~eqdim(X2,sK3) | eqdim(X2,sK2)) )),
  inference(resolution,[],[f1047,f458])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (~eqdim(X1,X2) | eqdim(X0,X2) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : (~eqdim(X0,X1) | ~eqdim(X1,X2) | eqdim(X0,X2))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : ((~eqdim(X0,X1) | ~eqdim(X1,X2)) | eqdim(X0,X2))),
  inference(ennf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X14] : ! [X15] : ! [X16] : ((eqdim(X14,X15) & eqdim(X15,X16)) => eqdim(X14,X16))),
  file(unknown,unknown)).
fof(f1047,plain,(
  eqdim(sK3,sK2)),
  inference(resolution,[],[f625,f514])).
fof(f514,plain,(
  eqdim(sK3,sK0)),
  inference(resolution,[],[f479,f441])).
fof(f625,plain,(
  ( ! [X2] : (~eqdim(X2,sK0) | eqdim(X2,sK2)) )),
  inference(resolution,[],[f549,f458])).
fof(f549,plain,(
  eqdim(sK0,sK2)),
  inference(resolution,[],[f500,f386])).
fof(f386,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X12] : ! [X13] : (eqdim(X12,X13) => eqdim(X13,X12))),
  file(unknown,unknown)).
fof(f500,plain,(
  eqdim(sK2,sK0)),
  inference(resolution,[],[f478,f441])).
fof(f15197,plain,(
  eqdim(sK11(sK2,sK3),sK3) | $spl169),
  inference(resolution,[],[f15189,f441])).
fof(f15189,plain,(
  p(sK11(sK2,sK3),sK3) | $spl169),
  inference(cnf_transformation,[],[f15189_D])).
fof(f15189_D,plain,(
  p(sK11(sK2,sK3),sK3) <=> ~$spl169),
  introduced(backtracking_splitting_component,[])).
fof(f15103,plain,(
  p(sK11(sK2,sK3),sK2) | p(sK11(sK2,sK3),sK3)),
  inference(backtracking_split_refutation,[],[f858,f3214_D,f15100])).
fof(f15100,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f15076,f3214])).
fof(f3214,plain,(
  sc(sK3,sK2) | $spl76),
  inference(cnf_transformation,[],[f3214_D])).
fof(f3214_D,plain,(
  sc(sK3,sK2) <=> ~$spl76),
  introduced(backtracking_splitting_component,[])).
fof(f858,plain,(
  p(sK11(sK2,sK3),sK2) | p(sK11(sK2,sK3),sK3) | sc(sK3,sK2)),
  inference(resolution,[],[f739,f398])).
fof(f398,plain,(
  ( ! [X0,X1] : (~c(X0,X1) | p(sK11(X1,X0),X1) | p(sK11(X1,X0),X0) | sc(X0,X1)) )),
  inference(cnf_transformation,[],[f291])).
fof(f18016,plain,(
  ~eqdim(sK11(sK2,sK3),sK2)),
  inference(backtracking_split_refutation,[],[f15152,f15674_D,f18012])).
fof(f18012,plain,(
  $false | $spl180),
  inference(subsumption_resolution,[],[f18001,f17971])).
fof(f17971,plain,(
  ~eqdim(sK11(sK2,sK3),sK3)),
  inference(backtracking_split_refutation,[],[f15138,f15669_D,f17969])).
fof(f17969,plain,(
  $false | $spl179),
  inference(subsumption_resolution,[],[f17968,f15078])).
fof(f15078,plain,(
  cont(sK11(sK2,sK3),sK3)),
  inference(backtracking_split_refutation,[],[f856,f2582_D,f15077])).
fof(f15077,plain,(
  $false | $spl45),
  inference(subsumption_resolution,[],[f15076,f2582])).
fof(f2582,plain,(
  sc(sK3,sK2) | $spl45),
  inference(cnf_transformation,[],[f2582_D])).
fof(f2582_D,plain,(
  sc(sK3,sK2) <=> ~$spl45),
  introduced(backtracking_splitting_component,[])).
fof(f856,plain,(
  cont(sK11(sK2,sK3),sK3) | sc(sK3,sK2)),
  inference(resolution,[],[f739,f396])).
fof(f396,plain,(
  ( ! [X0,X1] : (~c(X0,X1) | cont(sK11(X1,X0),X0) | sc(X0,X1)) )),
  inference(cnf_transformation,[],[f291])).
fof(f17968,plain,(
  ~cont(sK11(sK2,sK3),sK3) | $spl179),
  inference(subsumption_resolution,[],[f17961,f17938])).
fof(f17938,plain,(
  ~p(sK11(sK2,sK3),sK3)),
  inference(backtracking_split_refutation,[],[f15103,f15189_D,f17936])).
fof(f17961,plain,(
  p(sK11(sK2,sK3),sK3) | ~cont(sK11(sK2,sK3),sK3) | $spl179),
  inference(resolution,[],[f15669,f442])).
fof(f15669,plain,(
  eqdim(sK11(sK2,sK3),sK3) | $spl179),
  inference(cnf_transformation,[],[f15669_D])).
fof(f15669_D,plain,(
  eqdim(sK11(sK2,sK3),sK3) <=> ~$spl179),
  introduced(backtracking_splitting_component,[])).
fof(f15138,plain,(
  less(sK11(sK2,sK3),sK3) | eqdim(sK11(sK2,sK3),sK3)),
  inference(resolution,[],[f15078,f388])).
fof(f388,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | less(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1] : (~cont(X0,X1) | less(X0,X1) | eqdim(X0,X1))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : (~cont(X0,X1) | (less(X0,X1) | eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : (cont(X0,X1) => (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f2])).
fof(f2,axiom,(
  ! [X4] : ! [X5] : (cont(X4,X5) => (less(X4,X5) | eqdim(X4,X5)))),
  file(unknown,unknown)).
fof(f18001,plain,(
  eqdim(sK11(sK2,sK3),sK3) | $spl180),
  inference(resolution,[],[f15674,f1269])).
fof(f1269,plain,(
  ( ! [X2] : (~eqdim(X2,sK2) | eqdim(X2,sK3)) )),
  inference(resolution,[],[f1065,f458])).
fof(f1065,plain,(
  eqdim(sK2,sK3)),
  inference(resolution,[],[f650,f500])).
fof(f650,plain,(
  ( ! [X2] : (~eqdim(X2,sK0) | eqdim(X2,sK3)) )),
  inference(resolution,[],[f567,f458])).
fof(f567,plain,(
  eqdim(sK0,sK3)),
  inference(resolution,[],[f514,f386])).
fof(f15674,plain,(
  eqdim(sK11(sK2,sK3),sK2) | $spl180),
  inference(cnf_transformation,[],[f15674_D])).
fof(f15674_D,plain,(
  eqdim(sK11(sK2,sK3),sK2) <=> ~$spl180),
  introduced(backtracking_splitting_component,[])).
fof(f15152,plain,(
  less(sK11(sK2,sK3),sK2) | eqdim(sK11(sK2,sK3),sK2)),
  inference(resolution,[],[f15098,f388])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 3440
Passive clauses: 10189
Generated clauses: 17818
Final active clauses: 2849
Final passive clauses: 4581
Input formulas: 68
Initial clauses: 139

Duplicate literals: 19
Fw subsumption resolutions: 513
Fw demodulations: 223
Bw demodulations: 1122

Simple tautologies: 108
Forward subsumptions: 6632

Binary resolution: 15253

Backtracking splits: 232
Backtracking splits refuted: 89
Backtracking splits refuted at zero level: 30

Memory used: 4605KB
Time elapsed: 0.208 s
------------------------------
% Success in time 22.062 s
============================= vampire_lin32 ================================
execution finished: Thu Oct 10 11:56:57 2013
total CPU time used: 20
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_mereotopology_codib/theorems/conversions/endpoint_theorems_goal2.all.tptp
============================ end of footer ===========================
