Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1200 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2132
Passive clauses: 20695
Generated clauses: 34811
Input formulas: 39
Initial clauses: 89

Selected by SInE selection: 39
SInE iterations: 1

Duplicate literals: 126
Fw subsumption resolutions: 1114

Simple tautologies: 71
Forward subsumptions: 12027

Binary resolution: 31984
Factoring: 79

Backtracking splits: 1084
Backtracking splits refuted: 134
Backtracking splits refuted at zero level: 32

Memory used: 6780KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f338,plain,(
  $false),
  inference(subsumption_resolution,[],[f336,f321])).
fof(f321,plain,(
  ~zex(sK0)),
  inference(resolution,[],[f278,f301])).
fof(f301,plain,(
  ~zex(difference(sK0,sK0))),
  inference(forward_demodulation,[],[f208,f207])).
fof(f207,plain,(
  sK0 = sK1),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  sK0 = sK1 & ~zex(difference(sK0,sK1))),
  inference(skolemisation,[status(esa)],[f109])).
fof(f109,plain,(
  ? [X0,X1] : (X0 = X1 & ~zex(difference(X0,X1)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1] : (X0 = X1 => zex(difference(X0,X1)))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : (X0 = X1 => zex(difference(X0,X1)))),
  file(unknown,unknown)).
fof(f208,plain,(
  ~zex(difference(sK0,sK1))),
  inference(cnf_transformation,[],[f152])).
fof(f278,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f201])).
fof(f201,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X18] : ! [X19] : (zex(difference(X18,X19)) <=> (zex(X18) | cont(X18,X19)))),
  file(unknown,unknown)).
fof(f336,plain,(
  zex(sK0)),
  inference(resolution,[],[f335,f221])).
fof(f221,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X84] : (~zex(X84) <=> cont(X84,X84))),
  file(unknown,unknown)).
fof(f335,plain,(
  ~cont(sK0,sK0)),
  inference(resolution,[],[f279,f301])).
fof(f279,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f202])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 55
Passive clauses: 115
Generated clauses: 126
Final active clauses: 55
Final passive clauses: 58
Input formulas: 39
Initial clauses: 88

Selected by SInE selection: 39
SInE iterations: 1

Fw subsumption resolutions: 1
Fw demodulations: 1

Simple tautologies: 6
Equational tautologies: 1
Forward subsumptions: 2

Binary resolution: 35
Factoring: 1

Unique components: 2

Memory used: 127KB
Time elapsed: 0.035 s
------------------------------
% Success in time 0.454 s
