Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f3161,plain,(
  $false),
  inference(subsumption_resolution,[],[f3151,f301])).
fof(f301,plain,(
  zex(sK0)),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  zex(sK0) & zex(sK1) & ~zex(sum(sK0,sK1))),
  inference(skolemisation,[status(esa)],[f167])).
fof(f167,plain,(
  ? [X0,X1] : (zex(X0) & zex(X1) & ~zex(sum(X0,X1)))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ? [X0,X1] : ((zex(X0) & zex(X1)) & ~zex(sum(X0,X1)))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ~! [X0,X1] : ((zex(X0) & zex(X1)) => zex(sum(X0,X1)))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => zex(sum(X0,X1)))),
  file(unknown,unknown)).
fof(f3151,plain,(
  ~zex(sK0)),
  inference(resolution,[],[f3143,f320])).
fof(f320,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X140] : ! [X141] : (zex(X140) => (~cont(X141,X140) & ~cont(X140,X141)))),
  file(unknown,unknown)).
fof(f3143,plain,(
  cont(sum(sK0,sK1),sK0)),
  inference(subsumption_resolution,[],[f3119,f302])).
fof(f302,plain,(
  zex(sK1)),
  inference(cnf_transformation,[],[f238])).
fof(f3119,plain,(
  cont(sum(sK0,sK1),sK0) | ~zex(sK1)),
  inference(resolution,[],[f1419,f429])).
fof(f429,plain,(
  cont(sum(sK0,sK1),sum(sK0,sK1))),
  inference(resolution,[],[f316,f303])).
fof(f303,plain,(
  ~zex(sum(sK0,sK1))),
  inference(cnf_transformation,[],[f238])).
fof(f316,plain,(
  ( ! [X0] : (zex(X0) | cont(X0,X0)) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X134] : (~zex(X134) <=> cont(X134,X134))),
  file(unknown,unknown)).
fof(f1419,plain,(
  ( ! [X19,X17,X18] : (~cont(X17,sum(X18,X19)) | cont(X17,X18) | ~zex(X19)) )),
  inference(resolution,[],[f407,f320])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : (~cont(X2,sum(X0,X1)) | cont(X2,X0) | cont(difference(X2,X0),X1))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((~cont(X2,sum(X0,X1)) | cont(X2,X0)) | cont(difference(X2,X0),X1))),
  inference(ennf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X9] : ! [X10] : ! [X11] : ((cont(X11,sum(X9,X10)) & ~cont(X11,X9)) => cont(difference(X11,X9),X10))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 424
Passive clauses: 1633
Generated clauses: 2850
Final active clauses: 421
Final passive clauses: 1077
Input formulas: 58
Initial clauses: 116

Selected by SInE selection: 58
SInE iterations: 1

Duplicate literals: 16
Fw subsumption resolutions: 80

Simple tautologies: 30
Forward subsumptions: 1091

Binary resolution: 2528
Factoring: 8

Backtracking splits: 83
Backtracking splits refuted: 10
Backtracking splits refuted at zero level: 1

Memory used: 895KB
Time elapsed: 0.030 s
------------------------------
% Success in time 0.057 s
