Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f1536,plain,(
  $false),
  inference(subsumption_resolution,[],[f1526,f348])).
fof(f348,plain,(
  ~p(difference(sK2,sK0),sK2)),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  eqdim(sK0,sK1) & cont(sK2,sum(sK0,sK1)) & ~cont(sK2,sK0) & ~cont(sK2,sK1) & ~zex(sK0) & ~zex(sK1) & ~p(difference(sK2,sK0),sK2)),
  inference(skolemisation,[status(esa)],[f193])).
fof(f193,plain,(
  ? [X0,X1,X2] : (eqdim(X0,X1) & cont(X2,sum(X0,X1)) & ~cont(X2,X0) & ~cont(X2,X1) & ~zex(X0) & ~zex(X1) & ~p(difference(X2,X0),X2))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ? [X0,X1,X2] : ((eqdim(X0,X1) & cont(X2,sum(X0,X1)) & ~cont(X2,X0) & ~cont(X2,X1) & ~zex(X0) & ~zex(X1)) & ~p(difference(X2,X0),X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ~! [X0,X1,X2] : ((eqdim(X0,X1) & cont(X2,sum(X0,X1)) & ~cont(X2,X0) & ~cont(X2,X1) & ~zex(X0) & ~zex(X1)) => p(difference(X2,X0),X2))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ~! [X0] : ! [X1] : ! [X2] : ((eqdim(X0,X1) & cont(X2,sum(X0,X1)) & ~cont(X2,X0) & ~cont(X2,X1) & ~zex(X0) & ~zex(X1)) => p(difference(X2,X0),X2))),
  inference(rectify,[],[f10])).
fof(f10,negated_conjecture,(
  ~! [X22] : ! [X23] : ! [X24] : ((eqdim(X22,X23) & cont(X24,sum(X22,X23)) & ~cont(X24,X22) & ~cont(X24,X23) & ~zex(X22) & ~zex(X23)) => p(difference(X24,X22),X24))),
  file(unknown,unknown)).
fof(f1526,plain,(
  p(difference(sK2,sK0),sK2)),
  inference(resolution,[],[f1516,f398])).
fof(f398,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : ! [X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X138] : ! [X139] : (~zex(difference(X138,X139)) => p(difference(X138,X139),X138))),
  file(unknown,unknown)).
fof(f1516,plain,(
  ~zex(difference(sK2,sK0))),
  inference(resolution,[],[f1373,f344])).
fof(f344,plain,(
  ~cont(sK2,sK0)),
  inference(cnf_transformation,[],[f278])).
fof(f1373,plain,(
  ( ! [X14] : (cont(sK2,X14) | ~zex(difference(sK2,X14))) )),
  inference(resolution,[],[f442,f501])).
fof(f501,plain,(
  ~zex(sK2)),
  inference(resolution,[],[f368,f343])).
fof(f343,plain,(
  cont(sK2,sum(sK0,sK1))),
  inference(cnf_transformation,[],[f278])).
fof(f368,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X163] : ! [X164] : (zex(X163) => (~cont(X164,X163) & ~cont(X163,X164)))),
  file(unknown,unknown)).
fof(f442,plain,(
  ( ! [X0,X1] : (zex(X0) | cont(X0,X1) | ~zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f335])).
fof(f335,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X48] : ! [X49] : (zex(difference(X48,X49)) <=> (zex(X48) | cont(X48,X49)))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 295
Passive clauses: 607
Generated clauses: 1183
Final active clauses: 295
Final passive clauses: 284
Input formulas: 67
Initial clauses: 132

Selected by SInE selection: 67
SInE iterations: 1

Duplicate literals: 22
Fw subsumption resolutions: 18

Simple tautologies: 13
Forward subsumptions: 518

Binary resolution: 986
Factoring: 2

Backtracking splits: 23

Memory used: 383KB
Time elapsed: 0.034 s
------------------------------
% Success in time 0.057 s
