Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f11802,plain,(
  $false),
  inference(subsumption_resolution,[],[f11800,f644])).
fof(f644,plain,(
  p(sK1,sK1)),
  inference(subsumption_resolution,[],[f635,f535])).
fof(f535,plain,(
  ( ! [X0] : (eqdim(X0,X0)) )),
  inference(subsumption_resolution,[],[f523,f262])).
fof(f262,plain,(
  ( ! [X0] : (~less(X0,X0)) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X53] : ~less(X53,X53)),
  file(unknown,unknown)).
fof(f523,plain,(
  ( ! [X0] : (eqdim(X0,X0) | less(X0,X0)) )),
  inference(resolution,[],[f306,f262])).
fof(f306,plain,(
  ( ! [X0,X1] : (less(X1,X0) | eqdim(X0,X1) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X64] : ! [X65] : (eqdim(X64,X65) <=> (~less(X64,X65) & ~less(X65,X64)))),
  file(unknown,unknown)).
fof(f635,plain,(
  ~eqdim(sK1,sK1) | p(sK1,sK1)),
  inference(resolution,[],[f321,f353])).
fof(f353,plain,(
  cont(sK1,sK1)),
  inference(resolution,[],[f260,f244])).
fof(f244,plain,(
  ~zex(sK1)),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ~zex(sK0) & ~zex(sK1) & ~p(sK1,sK0) & ~eqdim(sK0,sK1) & ! [X2] : (~p(X2,sK1) | po(X2,sK0))),
  inference(skolemisation,[status(esa)],[f134])).
fof(f134,plain,(
  ? [X0,X1] : (~zex(X0) & ~zex(X1) & ~p(X1,X0) & ~eqdim(X0,X1) & ! [X2] : (~p(X2,X1) | po(X2,X0)))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ? [X0,X1] : ((~zex(X0) & ~zex(X1) & ~p(X1,X0) & ~eqdim(X0,X1)) & ! [X2] : (~p(X2,X1) | po(X2,X0)))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ~! [X0,X1] : ((~zex(X0) & ~zex(X1) & ~p(X1,X0) & ~eqdim(X0,X1)) => ? [X2] : (p(X2,X1) & ~po(X2,X0)))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : ((~zex(X0) & ~zex(X1) & ~p(X1,X0) & ~eqdim(X0,X1)) => ? [X2] : (p(X2,X1) & ~po(X2,X0)))),
  file(unknown,unknown)).
fof(f260,plain,(
  ( ! [X0] : (zex(X0) | cont(X0,X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X104] : (~zex(X104) <=> cont(X104,X104))),
  file(unknown,unknown)).
fof(f321,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X75] : ! [X76] : (p(X75,X76) <=> (cont(X75,X76) & eqdim(X75,X76)))),
  file(unknown,unknown)).
fof(f11800,plain,(
  ~p(sK1,sK1)),
  inference(resolution,[],[f11799,f247])).
fof(f247,plain,(
  ( ! [X2] : (po(X2,sK0) | ~p(X2,sK1)) )),
  inference(cnf_transformation,[],[f187])).
fof(f11799,plain,(
  ~po(sK1,sK0)),
  inference(resolution,[],[f11748,f311])).
fof(f311,plain,(
  ( ! [X0,X1] : (p(sK14(X1,X0),X1) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1] : ((~po(X0,X1) | (p(sK14(X1,X0),X0) & p(sK14(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(skolemisation,[status(esa)],[f228])).
fof(f228,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(rectify,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X2] : (~p(X2,X0) | ~p(X2,X1)) | po(X0,X1)))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : ! [X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X79] : ! [X80] : (po(X79,X80) <=> ? [X81] : (p(X81,X79) & p(X81,X80)))),
  file(unknown,unknown)).
fof(f11748,plain,(
  ~p(sK14(sK0,sK1),sK0)),
  inference(resolution,[],[f11725,f320])).
fof(f320,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f236])).
fof(f11725,plain,(
  ~eqdim(sK14(sK0,sK1),sK0)),
  inference(resolution,[],[f11713,f304])).
fof(f304,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f224])).
fof(f11713,plain,(
  less(sK14(sK0,sK1),sK0)),
  inference(backtracking_split_refutation,[],[f1185,f1196_D,f11711])).
fof(f11711,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f11709,f644])).
fof(f11709,plain,(
  ~p(sK1,sK1) | $spl12),
  inference(resolution,[],[f11708,f247])).
fof(f11708,plain,(
  ~po(sK1,sK0) | $spl12),
  inference(resolution,[],[f11055,f310])).
fof(f310,plain,(
  ( ! [X0,X1] : (p(sK14(X1,X0),X0) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f229])).
fof(f11055,plain,(
  ~p(sK14(sK0,sK1),sK1) | $spl12),
  inference(resolution,[],[f10700,f1320])).
fof(f1320,plain,(
  eqdim(sK0,sK14(sK0,sK1)) | $spl12),
  inference(backtracking_split_refutation,[],[f1276,f1278_D,f1319])).
fof(f1319,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f1317,f644])).
fof(f1317,plain,(
  ~p(sK1,sK1) | $spl20),
  inference(resolution,[],[f1316,f247])).
fof(f1316,plain,(
  ~po(sK1,sK0) | $spl20),
  inference(resolution,[],[f1305,f311])).
fof(f1305,plain,(
  ~p(sK14(sK0,sK1),sK0) | $spl20),
  inference(resolution,[],[f1292,f320])).
fof(f1292,plain,(
  ~eqdim(sK14(sK0,sK1),sK0) | $spl20),
  inference(resolution,[],[f1278,f305])).
fof(f305,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f224])).
fof(f1278,plain,(
  less(sK0,sK14(sK0,sK1)) | $spl20),
  inference(cnf_transformation,[],[f1278_D])).
fof(f1278_D,plain,(
  less(sK0,sK14(sK0,sK1)) <=> ~$spl20),
  introduced(backtracking_splitting_component,[])).
fof(f1276,plain,(
  eqdim(sK0,sK14(sK0,sK1)) | less(sK0,sK14(sK0,sK1)) | $spl12),
  inference(resolution,[],[f1196,f306])).
fof(f10700,plain,(
  ( ! [X2] : (~eqdim(sK0,X2) | ~p(X2,sK1)) )),
  inference(resolution,[],[f10656,f320])).
fof(f10656,plain,(
  ( ! [X87] : (~eqdim(sK0,X87) | ~eqdim(X87,sK1)) )),
  inference(resolution,[],[f2042,f4626])).
fof(f4626,plain,(
  less(sK1,sK0)),
  inference(backtracking_split_refutation,[],[f1178,f1192_D,f4624])).
fof(f4624,plain,(
  $false | $spl9),
  inference(subsumption_resolution,[],[f4622,f644])).
fof(f4622,plain,(
  ~p(sK1,sK1) | $spl9),
  inference(resolution,[],[f4621,f247])).
fof(f4621,plain,(
  ~po(sK1,sK0) | $spl9),
  inference(resolution,[],[f4620,f310])).
fof(f4620,plain,(
  ~p(sK14(sK0,sK1),sK1) | $spl9),
  inference(resolution,[],[f4611,f320])).
fof(f4611,plain,(
  ~eqdim(sK14(sK0,sK1),sK1) | $spl9),
  inference(resolution,[],[f4546,f536])).
fof(f536,plain,(
  ( ! [X6,X5] : (eqdim(X5,X6) | ~eqdim(X6,X5)) )),
  inference(subsumption_resolution,[],[f526,f305])).
fof(f526,plain,(
  ( ! [X6,X5] : (eqdim(X5,X6) | less(X5,X6) | ~eqdim(X6,X5)) )),
  inference(resolution,[],[f306,f304])).
fof(f4546,plain,(
  ~eqdim(sK1,sK14(sK0,sK1)) | $spl9),
  inference(resolution,[],[f2057,f1203])).
fof(f1203,plain,(
  less(sK0,sK1) | $spl9),
  inference(subsumption_resolution,[],[f1202,f246])).
fof(f246,plain,(
  ~eqdim(sK0,sK1)),
  inference(cnf_transformation,[],[f187])).
fof(f1202,plain,(
  eqdim(sK0,sK1) | less(sK0,sK1) | $spl9),
  inference(resolution,[],[f1192,f306])).
fof(f2057,plain,(
  ( ! [X42] : (~less(sK0,X42) | ~eqdim(X42,sK14(sK0,sK1))) )),
  inference(resolution,[],[f698,f1321])).
fof(f1321,plain,(
  ~less(sK0,sK14(sK0,sK1))),
  inference(backtracking_split_refutation,[],[f1276,f1278_D,f1319])).
fof(f698,plain,(
  ( ! [X12,X13,X11] : (less(X11,X12) | ~less(X11,X13) | ~eqdim(X13,X12)) )),
  inference(resolution,[],[f331,f315])).
fof(f315,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X66] : ! [X67] : (leq(X66,X67) <=> (less(X66,X67) | eqdim(X66,X67)))),
  file(unknown,unknown)).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (~leq(X1,X2) | less(X0,X2) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2] : (~less(X0,X1) | ~leq(X1,X2) | less(X0,X2))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : ((~less(X0,X1) | ~leq(X1,X2)) | less(X0,X2))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X56] : ! [X57] : ! [X58] : ((less(X56,X57) & leq(X57,X58)) => less(X56,X58))),
  file(unknown,unknown)).
fof(f1192,plain,(
  ~less(sK1,sK0) | $spl9),
  inference(cnf_transformation,[],[f1192_D])).
fof(f1192_D,plain,(
  ~less(sK1,sK0) <=> ~$spl9),
  introduced(backtracking_splitting_component,[])).
fof(f1178,plain,(
  ( ! [X22] : (~zex(difference(sK0,X22)) | ~less(sK1,sK0) | ~mindim(X22)) )),
  inference(resolution,[],[f718,f378])).
fof(f378,plain,(
  ( ! [X1] : (~less(sK1,X1) | ~mindim(X1)) )),
  inference(resolution,[],[f257,f244])).
fof(f257,plain,(
  ( ! [X0,X1] : (zex(X1) | ~less(X1,X0) | ~mindim(X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1)))) & (zex(X0) | (less(sK5(X0),X0) & ~zex(sK5(X0))) | mindim(X0)))),
  inference(skolemisation,[status(esa)],[f201])).
fof(f201,plain,(
  ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1)))) & (zex(X0) | ? [X2] : (less(X2,X0) & ~zex(X2)) | mindim(X0)))),
  inference(rectify,[],[f200])).
fof(f200,plain,(
  ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1)))) & (zex(X0) | ? [X1] : (less(X1,X0) & ~zex(X1)) | mindim(X0)))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1)))) & ((zex(X0) | ? [X1] : (less(X1,X0) & ~zex(X1))) | mindim(X0)))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : (mindim(X0) <=> (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1))))),
  inference(ennf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (mindim(X0) <=> (~zex(X0) & ! [X1] : (less(X1,X0) => zex(X1))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X70] : (mindim(X70) <=> (~zex(X70) & ! [X71] : (less(X71,X70) => zex(X71))))),
  file(unknown,unknown)).
fof(f718,plain,(
  ( ! [X0,X1] : (less(X1,X0) | ~zex(difference(sK0,X0)) | ~less(X1,sK0)) )),
  inference(resolution,[],[f677,f331])).
fof(f677,plain,(
  ( ! [X1] : (leq(sK0,X1) | ~zex(difference(sK0,X1))) )),
  inference(resolution,[],[f663,f270])).
fof(f270,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1)) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X50] : ! [X51] : (cont(X50,X51) => leq(X50,X51))),
  file(unknown,unknown)).
fof(f663,plain,(
  ( ! [X0] : (cont(sK0,X0) | ~zex(difference(sK0,X0))) )),
  inference(resolution,[],[f322,f243])).
fof(f243,plain,(
  ~zex(sK0)),
  inference(cnf_transformation,[],[f187])).
fof(f322,plain,(
  ( ! [X0,X1] : (zex(X0) | cont(X0,X1) | ~zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X19] : ! [X20] : (zex(difference(X19,X20)) <=> (zex(X19) | cont(X19,X20)))),
  file(unknown,unknown)).
fof(f2042,plain,(
  ( ! [X6,X7,X5] : (~less(X5,X6) | ~eqdim(X6,X7) | ~eqdim(X7,X5)) )),
  inference(resolution,[],[f698,f305])).
fof(f1196,plain,(
  ~less(sK14(sK0,sK1),sK0) | $spl12),
  inference(cnf_transformation,[],[f1196_D])).
fof(f1196_D,plain,(
  ~less(sK14(sK0,sK1),sK0) <=> ~$spl12),
  introduced(backtracking_splitting_component,[])).
fof(f1185,plain,(
  ( ! [X28] : (~zex(difference(sK0,X28)) | ~less(sK14(sK0,sK1),sK0) | ~mindim(X28)) )),
  inference(resolution,[],[f718,f658])).
fof(f658,plain,(
  ( ! [X1] : (~less(sK14(sK0,sK1),X1) | ~mindim(X1)) )),
  inference(resolution,[],[f653,f257])).
fof(f653,plain,(
  ~zex(sK14(sK0,sK1))),
  inference(resolution,[],[f644,f463])).
fof(f463,plain,(
  ( ! [X0] : (~p(X0,sK1) | ~zex(sK14(sK0,X0))) )),
  inference(resolution,[],[f458,f247])).
fof(f458,plain,(
  ( ! [X2,X3] : (~po(X2,X3) | ~zex(sK14(X3,X2))) )),
  inference(resolution,[],[f310,f370])).
fof(f370,plain,(
  ( ! [X4,X3] : (~p(X3,X4) | ~zex(X3)) )),
  inference(resolution,[],[f319,f264])).
fof(f264,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X110] : ! [X111] : (zex(X110) => (~cont(X111,X110) & ~cont(X110,X111)))),
  file(unknown,unknown)).
fof(f319,plain,(
  ( ! [X0,X1] : (cont(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f236])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1082
Passive clauses: 6333
Generated clauses: 11624
Final active clauses: 899
Final passive clauses: 3782
Input formulas: 47
Initial clauses: 101

Selected by SInE selection: 47
SInE iterations: 1

Duplicate literals: 73
Fw subsumption resolutions: 434

Simple tautologies: 51
Forward subsumptions: 4480

Binary resolution: 10476
Factoring: 32

Backtracking splits: 337
Backtracking splits refuted: 89
Backtracking splits refuted at zero level: 32

Memory used: 2174KB
Time elapsed: 0.109 s
------------------------------
% Success in time 0.132 s
