Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 1983
Passive clauses: 21158
Generated clauses: 32880
Input formulas: 59
Initial clauses: 117

Selected by SInE selection: 59
SInE iterations: 1

Duplicate literals: 220
Fw subsumption resolutions: 1993

Simple tautologies: 187
Forward subsumptions: 8125

Binary resolution: 29023
Factoring: 55

Backtracking splits: 1276
Backtracking splits refuted: 180
Backtracking splits refuted at zero level: 45

Memory used: 6780KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f9336,plain,(
  $false),
  inference(subsumption_resolution,[],[f9328,f306])).
fof(f306,plain,(
  eqdim(sK1,sK0)),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  eqdim(sK1,sK0) & ~zex(sK0) & ~p(sK0,intersection(sK0,sum(sK0,sK1)))),
  inference(skolemisation,[status(esa)],[f170])).
fof(f170,plain,(
  ? [X0,X1] : (eqdim(X1,X0) & ~zex(X0) & ~p(X0,intersection(X0,sum(X0,X1))))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ? [X0,X1] : ((eqdim(X1,X0) & ~zex(X0)) & ~p(X0,intersection(X0,sum(X0,X1))))),
  inference(ennf_transformation,[],[f61])).
fof(f61,plain,(
  ~! [X0,X1] : ((eqdim(X1,X0) & ~zex(X0)) => p(X0,intersection(X0,sum(X0,X1))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ~! [X0] : ! [X1] : ((eqdim(X1,X0) & ~zex(X0)) => p(X0,intersection(X0,sum(X0,X1))))),
  inference(rectify,[],[f2])).
fof(f2,negated_conjecture,(
  ~! [X2] : ! [X3] : ((eqdim(X3,X2) & ~zex(X2)) => p(X2,intersection(X2,sum(X2,X3))))),
  file(unknown,unknown)).
fof(f9328,plain,(
  ~eqdim(sK1,sK0)),
  inference(resolution,[],[f9327,f382])).
fof(f382,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X80] : ! [X81] : (leq(X80,X81) <=> (less(X80,X81) | eqdim(X80,X81)))),
  file(unknown,unknown)).
fof(f9327,plain,(
  ~leq(sK1,sK0)),
  inference(subsumption_resolution,[],[f9322,f307])).
fof(f307,plain,(
  ~zex(sK0)),
  inference(cnf_transformation,[],[f243])).
fof(f9322,plain,(
  ~leq(sK1,sK0) | zex(sK0)),
  inference(resolution,[],[f9320,f321])).
fof(f321,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X136] : (~zex(X136) <=> cont(X136,X136))),
  file(unknown,unknown)).
fof(f9320,plain,(
  ~cont(sK0,sK0) | ~leq(sK1,sK0)),
  inference(duplicate_literal_removal,[],[f9306])).
fof(f9306,plain,(
  ~cont(sK0,sK0) | ~cont(sK0,sK0) | ~leq(sK1,sK0)),
  inference(resolution,[],[f9277,f1172])).
fof(f1172,plain,(
  ( ! [X6,X8,X7] : (cont(X8,sum(X7,X6)) | ~cont(X8,X7) | ~leq(X6,X7)) )),
  inference(superposition,[],[f405,f324])).
fof(f324,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = sum(X1,X0)) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : sum(X0,X1) = sum(X1,X0)),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X4] : ! [X5] : sum(X4,X5) = sum(X5,X4)),
  file(unknown,unknown)).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (cont(X2,sum(X0,X1)) | ~cont(X2,X1) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X1) | cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X1)) | cont(X2,sum(X0,X1)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & cont(X2,X1)) => cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & cont(X2,X1)) => cont(X2,sum(X0,X1)))),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X8] : ! [X9] : ! [X10] : ((leq(X8,X9) & cont(X10,X9)) => cont(X10,sum(X8,X9)))),
  file(unknown,unknown)).
fof(f9277,plain,(
  ~cont(sK0,sum(sK0,sK1)) | ~cont(sK0,sK0)),
  inference(subsumption_resolution,[],[f2349,f5442])).
fof(f5442,plain,(
  ( ! [X0] : (~less(sK0,intersection(sK0,X0))) )),
  inference(resolution,[],[f5403,f1804])).
fof(f1804,plain,(
  ( ! [X0] : (less(sK1,X0) | ~less(sK0,X0)) )),
  inference(subsumption_resolution,[],[f958,f906])).
fof(f906,plain,(
  ( ! [X1] : (~less(sK0,X1) | ~less(X1,sK1)) )),
  inference(resolution,[],[f902,f381])).
fof(f381,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f287])).
fof(f902,plain,(
  ( ! [X19] : (~leq(X19,sK1) | ~less(sK0,X19)) )),
  inference(resolution,[],[f408,f443])).
fof(f443,plain,(
  ~less(sK0,sK1)),
  inference(resolution,[],[f387,f306])).
fof(f387,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | ~less(X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f291])).
fof(f291,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X78] : ! [X79] : (eqdim(X78,X79) <=> (~less(X78,X79) & ~less(X79,X78)))),
  file(unknown,unknown)).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (less(X0,X2) | ~leq(X1,X2) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : (~less(X0,X1) | ~leq(X1,X2) | less(X0,X2))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((~less(X0,X1) | ~leq(X1,X2)) | less(X0,X2))),
  inference(ennf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X70] : ! [X71] : ! [X72] : ((less(X70,X71) & leq(X71,X72)) => less(X70,X72))),
  file(unknown,unknown)).
fof(f958,plain,(
  ( ! [X0] : (~less(sK0,X0) | less(sK1,X0) | less(X0,sK1)) )),
  inference(resolution,[],[f905,f388])).
fof(f388,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | less(X1,X0) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f292])).
fof(f905,plain,(
  ( ! [X0] : (~eqdim(X0,sK1) | ~less(sK0,X0)) )),
  inference(resolution,[],[f902,f382])).
fof(f5403,plain,(
  ( ! [X0] : (~less(sK1,intersection(sK0,X0))) )),
  inference(subsumption_resolution,[],[f5362,f3086])).
fof(f3086,plain,(
  ( ! [X47] : (~less(sK1,X47) | ~zex(X47)) )),
  inference(subsumption_resolution,[],[f3074,f307])).
fof(f3074,plain,(
  ( ! [X47] : (~less(sK1,X47) | ~zex(X47) | zex(sK0)) )),
  inference(resolution,[],[f3010,f464])).
fof(f464,plain,(
  ( ! [X4,X3] : (~less(X3,X4) | ~zex(X4) | zex(X3)) )),
  inference(resolution,[],[f335,f332])).
fof(f332,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : (~less(X0,X1) | ~less(X1,X0))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X68] : ! [X69] : (less(X68,X69) => ~less(X69,X68))),
  file(unknown,unknown)).
fof(f335,plain,(
  ( ! [X0,X1] : (less(X0,X1) | zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : (~zex(X0) | zex(X1) | less(X0,X1))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((~zex(X0) | zex(X1)) | less(X0,X1))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X75] : ! [X76] : ((zex(X75) & ~zex(X76)) => less(X75,X76))),
  file(unknown,unknown)).
fof(f3010,plain,(
  ( ! [X0] : (less(sK0,X0) | ~less(sK1,X0)) )),
  inference(subsumption_resolution,[],[f986,f932])).
fof(f932,plain,(
  ( ! [X1] : (~less(sK1,X1) | ~less(X1,sK0)) )),
  inference(resolution,[],[f903,f381])).
fof(f903,plain,(
  ( ! [X20] : (~leq(X20,sK0) | ~less(sK1,X20)) )),
  inference(resolution,[],[f408,f441])).
fof(f441,plain,(
  ~less(sK1,sK0)),
  inference(resolution,[],[f386,f306])).
fof(f386,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f292])).
fof(f986,plain,(
  ( ! [X0] : (~less(sK1,X0) | less(sK0,X0) | less(X0,sK0)) )),
  inference(resolution,[],[f931,f388])).
fof(f931,plain,(
  ( ! [X0] : (~eqdim(X0,sK0) | ~less(sK1,X0)) )),
  inference(resolution,[],[f903,f382])).
fof(f5362,plain,(
  ( ! [X0] : (~less(sK1,intersection(sK0,X0)) | zex(intersection(sK0,X0))) )),
  inference(resolution,[],[f5357,f321])).
fof(f5357,plain,(
  ( ! [X6,X7] : (~cont(X6,intersection(sK0,X7)) | ~less(sK1,X6)) )),
  inference(subsumption_resolution,[],[f5320,f325])).
fof(f325,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X142] : ! [X143] : (zex(X142) => (~cont(X143,X142) & ~cont(X142,X143)))),
  file(unknown,unknown)).
fof(f5320,plain,(
  ( ! [X6,X7] : (~less(sK1,X6) | ~cont(X6,intersection(sK0,X7)) | zex(intersection(sK0,X7))) )),
  inference(resolution,[],[f1001,f354])).
fof(f354,plain,(
  ( ! [X0,X1] : (cont(intersection(X0,X1),X0) | zex(intersection(X0,X1))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))),
  inference(ennf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X37] : ! [X38] : (~zex(intersection(X37,X38)) => cont(intersection(X37,X38),X37))),
  file(unknown,unknown)).
fof(f1001,plain,(
  ( ! [X0,X1] : (~cont(X1,sK0) | ~less(sK1,X0) | ~cont(X0,X1)) )),
  inference(resolution,[],[f933,f409])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))),
  inference(flattening,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))),
  inference(ennf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X139] : ! [X140] : ! [X141] : ((cont(X139,X140) & cont(X140,X141)) => cont(X139,X141))),
  file(unknown,unknown)).
fof(f933,plain,(
  ( ! [X2] : (~cont(X2,sK0) | ~less(sK1,X2)) )),
  inference(resolution,[],[f903,f333])).
fof(f333,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X51] : ! [X52] : (cont(X51,X52) => leq(X51,X52))),
  file(unknown,unknown)).
fof(f2349,plain,(
  ~cont(sK0,sum(sK0,sK1)) | ~cont(sK0,sK0) | less(sK0,intersection(sK0,sum(sK0,sK1)))),
  inference(subsumption_resolution,[],[f2337,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0)) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))),
  inference(ennf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X39] : ! [X40] : ! [X41] : ((cont(X41,X39) & cont(X41,X40)) => leq(X41,intersection(X39,X40)))),
  file(unknown,unknown)).
fof(f2337,plain,(
  ~cont(sK0,sum(sK0,sK1)) | ~cont(sK0,sK0) | less(sK0,intersection(sK0,sum(sK0,sK1))) | ~leq(sK0,intersection(sK0,sum(sK0,sK1)))),
  inference(resolution,[],[f1838,f380])).
fof(f380,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | less(X0,X1) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f287])).
fof(f1838,plain,(
  ~eqdim(sK0,intersection(sK0,sum(sK0,sK1))) | ~cont(sK0,sum(sK0,sK1)) | ~cont(sK0,sK0)),
  inference(resolution,[],[f414,f308])).
fof(f308,plain,(
  ~p(sK0,intersection(sK0,sum(sK0,sK1)))),
  inference(cnf_transformation,[],[f243])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (p(X2,intersection(X0,X1)) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0)) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X42] : ! [X43] : ! [X44] : ((cont(X44,X42) & cont(X44,X43) & eqdim(X44,intersection(X42,X43))) <=> p(X44,intersection(X42,X43)))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 843
Passive clauses: 4914
Generated clauses: 9023
Final active clauses: 843
Final passive clauses: 3974
Input formulas: 59
Initial clauses: 116

Selected by SInE selection: 59
SInE iterations: 1

Duplicate literals: 101
Fw subsumption resolutions: 480

Simple tautologies: 166
Equational tautologies: 8
Forward subsumptions: 3449

Binary resolution: 5146
Factoring: 37
Forward superposition: 3054
Backward superposition: 95

Unique components: 3

Memory used: 2174KB
Time elapsed: 0.139 s
------------------------------
% Success in time 0.573 s
