Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 30000 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f4358,plain,(
  $false),
  inference(subsumption_resolution,[],[f4355,f4335])).
fof(f4335,plain,(
  sP0(sK0,sK0)),
  inference(backtracking_split_refutation,[],[f4232,f4235_D,f4333])).
fof(f4333,plain,(
  $false | $spl153),
  inference(subsumption_resolution,[],[f4332,f116])).
fof(f116,plain,(
  inc(sK0,sK0)),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  inc(sK0,sK0)),
  inference(skolemisation,[status(esa)],[f62])).
fof(f62,plain,(
  ? [X0] : inc(X0,X0)),
  inference(ennf_transformation,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ~inc(X0,X0)),
  file(unknown,unknown)).
fof(f4332,plain,(
  ~inc(sK0,sK0) | $spl153),
  inference(subsumption_resolution,[],[f4329,f4235])).
fof(f4329,plain,(
  sP0(sK0,sK0) | ~inc(sK0,sK0) | $spl153),
  inference(resolution,[],[f4254,f153])).
fof(f153,plain,(
  ( ! [X0,X1] : (less(sK4(X1,X0),X1) | sP0(X1,X0) | ~inc(X0,X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | (less(sK4(X1,X0),X1) & cont(sK4(X1,X0),X1) & p(sK4(X1,X0),X0))) & ((~sP0(X1,X0) & ! [X3] : (~less(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))),
  inference(skolemisation,[status(esa)],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | ? [X2] : (less(X2,X1) & cont(X2,X1) & p(X2,X0))) & ((~sP0(X1,X0) & ! [X3] : (~less(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))),
  inference(rectify,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | ? [X3] : (less(X3,X1) & cont(X3,X1) & p(X3,X0))) & ((~sP0(X1,X0) & ! [X3] : (~less(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : ((~inc(X0,X1) | (sP0(X1,X0) | ? [X3] : (less(X3,X1) & cont(X3,X1) & p(X3,X0)))) & ((~sP0(X1,X0) & ! [X3] : (~less(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : (inc(X0,X1) <=> (sP0(X1,X0) | ? [X3] : (less(X3,X1) & cont(X3,X1) & p(X3,X0))))),
  inference(definition_folding,[],[f87])).
fof(f87,plain,(
  ! [X1,X0] : (sP0(X1,X0) <=> ? [X2] : (less(X2,X0) & cont(X2,X0) & p(X2,X1)))),
  introduced(predicate_definition_introduction,[])).
fof(f4254,plain,(
  ~less(sK4(sK0,sK0),sK0) | $spl153),
  inference(subsumption_resolution,[],[f4244,f116])).
fof(f4244,plain,(
  ~inc(sK0,sK0) | ~less(sK4(sK0,sK0),sK0) | $spl153),
  inference(resolution,[],[f4235,f351])).
fof(f351,plain,(
  ( ! [X6,X5] : (sP0(X5,X6) | ~inc(X6,X5) | ~less(sK4(X5,X6),X6)) )),
  inference(resolution,[],[f155,f178])).
fof(f178,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | ~less(X0,X1)) )),
  inference(resolution,[],[f162,f147])).
fof(f147,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X37] : ! [X38] : (p(X37,X38) <=> (cont(X37,X38) & eqdim(X37,X38)))),
  file(unknown,unknown)).
fof(f162,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : ((~less(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | less(X0,X1)))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((~less(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | less(X0,X1)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : (less(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : (less(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X35] : ! [X36] : (less(X35,X36) <=> (leq(X35,X36) & ~eqdim(X35,X36)))),
  file(unknown,unknown)).
fof(f155,plain,(
  ( ! [X0,X1] : (p(sK4(X1,X0),X0) | sP0(X1,X0) | ~inc(X0,X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f4235,plain,(
  ~sP0(sK0,sK0) | $spl153),
  inference(cnf_transformation,[],[f4235_D])).
fof(f4235_D,plain,(
  ~sP0(sK0,sK0) <=> ~$spl153),
  introduced(backtracking_splitting_component,[])).
fof(f4232,plain,(
  ( ! [X0] : (~mindim(X0) | inc(sK0,X0) | ~cont(sK3(sK0,sK0),X0) | ~sP0(sK0,sK0)) ) | $spl146),
  inference(resolution,[],[f4146,f876])).
fof(f876,plain,(
  ( ! [X4,X2,X3] : (less(sK3(X2,X3),X4) | inc(X2,X4) | ~cont(sK3(X2,X3),X4) | ~sP0(X3,X2)) )),
  inference(subsumption_resolution,[],[f863,f149])).
fof(f149,plain,(
  ( ! [X0,X1] : (less(sK3(X1,X0),X1) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (less(sK3(X1,X0),X1) & cont(sK3(X1,X0),X1) & p(sK3(X1,X0),X0))) & (! [X3] : (~less(X3,X1) | ~cont(X3,X1) | ~p(X3,X0)) | sP0(X0,X1)))),
  inference(skolemisation,[status(esa)],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ? [X2] : (less(X2,X1) & cont(X2,X1) & p(X2,X0))) & (! [X3] : (~less(X3,X1) | ~cont(X3,X1) | ~p(X3,X0)) | sP0(X0,X1)))),
  inference(rectify,[],[f104])).
fof(f104,plain,(
  ! [X1,X0] : ((~sP0(X1,X0) | ? [X2] : (less(X2,X0) & cont(X2,X0) & p(X2,X1))) & (! [X2] : (~less(X2,X0) | ~cont(X2,X0) | ~p(X2,X1)) | sP0(X1,X0)))),
  inference(nnf_transformation,[],[f87])).
fof(f863,plain,(
  ( ! [X4,X2,X3] : (~cont(sK3(X2,X3),X4) | inc(X2,X4) | ~less(sK3(X2,X3),X2) | less(sK3(X2,X3),X4) | ~sP0(X3,X2)) )),
  inference(resolution,[],[f459,f150])).
fof(f150,plain,(
  ( ! [X0,X1] : (cont(sK3(X1,X0),X1) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f106])).
fof(f459,plain,(
  ( ! [X4,X5,X3] : (~cont(X3,X5) | ~cont(X3,X4) | inc(X4,X5) | ~less(X3,X4) | less(X3,X5)) )),
  inference(resolution,[],[f376,f206])).
fof(f206,plain,(
  ( ! [X0,X1] : (p(X0,X1) | ~cont(X0,X1) | less(X0,X1)) )),
  inference(duplicate_literal_removal,[],[f204])).
fof(f204,plain,(
  ( ! [X0,X1] : (p(X0,X1) | ~cont(X0,X1) | less(X0,X1) | ~cont(X0,X1)) )),
  inference(resolution,[],[f148,f135])).
fof(f135,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | less(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : (~cont(X0,X1) | less(X0,X1) | eqdim(X0,X1))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : (~cont(X0,X1) | (less(X0,X1) | eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : (cont(X0,X1) => (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X41] : ! [X42] : (cont(X41,X42) => (less(X41,X42) | eqdim(X41,X42)))),
  file(unknown,unknown)).
fof(f148,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f103])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (~p(X0,X1) | ~cont(X0,X2) | ~less(X0,X2) | inc(X2,X1)) )),
  inference(resolution,[],[f152,f156])).
fof(f156,plain,(
  ( ! [X0,X1] : (~sP0(X1,X0) | inc(X0,X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f152,plain,(
  ( ! [X0,X3,X1] : (sP0(X0,X1) | ~p(X3,X0) | ~cont(X3,X1) | ~less(X3,X1)) )),
  inference(cnf_transformation,[],[f106])).
fof(f4146,plain,(
  ( ! [X3] : (~less(sK3(sK0,sK0),X3) | ~mindim(X3)) ) | $spl146),
  inference(subsumption_resolution,[],[f4136,f1034])).
fof(f1034,plain,(
  s(sK3(sK0,sK0))),
  inference(backtracking_split_refutation,[],[f894,f898_D,f1032])).
fof(f1032,plain,(
  $false | $spl27),
  inference(subsumption_resolution,[],[f1031,f116])).
fof(f1031,plain,(
  ~inc(sK0,sK0) | $spl27),
  inference(subsumption_resolution,[],[f1030,f953])).
fof(f953,plain,(
  ~sP0(sK0,sK0) | $spl27),
  inference(resolution,[],[f917,f151])).
fof(f151,plain,(
  ( ! [X0,X1] : (p(sK3(X1,X0),X0) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f106])).
fof(f917,plain,(
  ( ! [X6] : (~p(sK3(sK0,sK0),X6)) ) | $spl27),
  inference(resolution,[],[f901,f177])).
fof(f177,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~p(X1,X0)) )),
  inference(resolution,[],[f159,f147])).
fof(f159,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | leq(X1,X0)) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X39] : ! [X40] : (eqdim(X39,X40) <=> (leq(X39,X40) & leq(X40,X39)))),
  file(unknown,unknown)).
fof(f901,plain,(
  ( ! [X0] : (~leq(X0,sK3(sK0,sK0))) ) | $spl27),
  inference(resolution,[],[f898,f137])).
fof(f137,plain,(
  ( ! [X0,X1] : (s(X1) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : (~leq(X0,X1) | s(X1))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : (leq(X0,X1) => s(X1))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) => s(X1))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X20] : ! [X21] : (leq(X20,X21) => s(X21))),
  file(unknown,unknown)).
fof(f1030,plain,(
  sP0(sK0,sK0) | ~inc(sK0,sK0) | $spl27),
  inference(resolution,[],[f970,f153])).
fof(f970,plain,(
  ~less(sK4(sK0,sK0),sK0) | $spl27),
  inference(subsumption_resolution,[],[f962,f116])).
fof(f962,plain,(
  ~inc(sK0,sK0) | ~less(sK4(sK0,sK0),sK0) | $spl27),
  inference(resolution,[],[f953,f351])).
fof(f898,plain,(
  ~s(sK3(sK0,sK0)) | $spl27),
  inference(cnf_transformation,[],[f898_D])).
fof(f898_D,plain,(
  ~s(sK3(sK0,sK0)) <=> ~$spl27),
  introduced(backtracking_splitting_component,[])).
fof(f894,plain,(
  less(sK2(sK3(sK0,sK0)),sK3(sK0,sK0)) | mindim(sK3(sK0,sK0)) | ~s(sK3(sK0,sK0)) | $spl24),
  inference(resolution,[],[f891,f127])).
fof(f127,plain,(
  ( ! [X0] : (zex(X0) | less(sK2(X0),X0) | mindim(X0) | ~s(X0)) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0] : ((~mindim(X0) | (s(X0) & ~zex(X0) & ! [X1] : (~s(X1) | ~less(X1,X0) | zex(X1)))) & (~s(X0) | zex(X0) | (s(sK2(X0)) & less(sK2(X0),X0) & ~zex(sK2(X0))) | mindim(X0)))),
  inference(skolemisation,[status(esa)],[f96])).
fof(f96,plain,(
  ! [X0] : ((~mindim(X0) | (s(X0) & ~zex(X0) & ! [X1] : (~s(X1) | ~less(X1,X0) | zex(X1)))) & (~s(X0) | zex(X0) | ? [X2] : (s(X2) & less(X2,X0) & ~zex(X2)) | mindim(X0)))),
  inference(rectify,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((~mindim(X0) | (s(X0) & ~zex(X0) & ! [X1] : (~s(X1) | ~less(X1,X0) | zex(X1)))) & (~s(X0) | zex(X0) | ? [X1] : (s(X1) & less(X1,X0) & ~zex(X1)) | mindim(X0)))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((~mindim(X0) | (s(X0) & ~zex(X0) & ! [X1] : (~s(X1) | ~less(X1,X0) | zex(X1)))) & ((~s(X0) | zex(X0) | ? [X1] : (s(X1) & less(X1,X0) & ~zex(X1))) | mindim(X0)))),
  inference(nnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0] : (mindim(X0) <=> (s(X0) & ~zex(X0) & ! [X1] : (~s(X1) | ~less(X1,X0) | zex(X1))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : (mindim(X0) <=> (s(X0) & ~zex(X0) & ! [X1] : ((~s(X1) | ~less(X1,X0)) | zex(X1))))),
  inference(ennf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : (mindim(X0) <=> (s(X0) & ~zex(X0) & ! [X1] : ((s(X1) & less(X1,X0)) => zex(X1))))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X5] : (mindim(X5) <=> (s(X5) & ~zex(X5) & ! [X6] : ((s(X6) & less(X6,X5)) => zex(X6))))),
  file(unknown,unknown)).
fof(f891,plain,(
  ~zex(sK3(sK0,sK0)) | $spl24),
  inference(cnf_transformation,[],[f891_D])).
fof(f891_D,plain,(
  ~zex(sK3(sK0,sK0)) <=> ~$spl24),
  introduced(backtracking_splitting_component,[])).
fof(f4136,plain,(
  ( ! [X3] : (~less(sK3(sK0,sK0),X3) | ~s(sK3(sK0,sK0)) | ~mindim(X3)) ) | $spl146),
  inference(resolution,[],[f4127,f125])).
fof(f125,plain,(
  ( ! [X0,X1] : (zex(X1) | ~less(X1,X0) | ~s(X1) | ~mindim(X0)) )),
  inference(cnf_transformation,[],[f97])).
fof(f4127,plain,(
  ~zex(sK3(sK0,sK0)) | $spl146),
  inference(cnf_transformation,[],[f4127_D])).
fof(f4127_D,plain,(
  ~zex(sK3(sK0,sK0)) <=> ~$spl146),
  introduced(backtracking_splitting_component,[])).
fof(f4355,plain,(
  ~sP0(sK0,sK0)),
  inference(resolution,[],[f4340,f149])).
fof(f4340,plain,(
  ~less(sK3(sK0,sK0),sK0)),
  inference(resolution,[],[f4335,f188])).
fof(f188,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | ~less(sK3(X1,X0),X0)) )),
  inference(resolution,[],[f151,f178])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 603
Passive clauses: 2199
Generated clauses: 4316
Final active clauses: 397
Final passive clauses: 1021
Input formulas: 23
Initial clauses: 54

Selected by SInE selection: 23
SInE iterations: 1

Duplicate literals: 57
Fw subsumption resolutions: 430

Simple tautologies: 69
Forward subsumptions: 1551

Binary resolution: 3358
Factoring: 52

Backtracking splits: 160
Backtracking splits refuted: 65
Backtracking splits refuted at zero level: 5

Memory used: 895KB
Time elapsed: 0.050 s
------------------------------
% Success in time 0.073 s
============================= vampire_lin32 ================================
execution finished: Tue Oct 08 20:36:16 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_space_codi/theorems/conversions/inc_theorems_goal1.all.tptp
============================ end of footer ===========================
