Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 30000 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2313
Passive clauses: 11090
Generated clauses: 32116
Input formulas: 24
Initial clauses: 52

Selected by SInE selection: 24
SInE iterations: 1

Duplicate literals: 548
Fw subsumption resolutions: 1424

Simple tautologies: 226
Forward subsumptions: 18667

Binary resolution: 27460
Factoring: 277

Backtracking splits: 1037
Backtracking splits refuted: 486
Backtracking splits refuted at zero level: 41

Memory used: 4733KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 29995 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f10433,plain,(
  $false),
  inference(subsumption_resolution,[],[f10432,f115])).
fof(f115,plain,(
  po(sK0,sK1)),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  po(sK0,sK1) & ~eqdim(sK0,sK1)),
  inference(skolemisation,[status(esa)],[f66])).
fof(f66,plain,(
  ? [X0,X1] : (po(X0,X1) & ~eqdim(X0,X1))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1] : (po(X0,X1) => eqdim(X0,X1))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : (po(X0,X1) => eqdim(X0,X1))),
  file(unknown,unknown)).
fof(f10432,plain,(
  ~po(sK0,sK1)),
  inference(duplicate_literal_removal,[],[f10424])).
fof(f10424,plain,(
  ~po(sK0,sK1) | ~po(sK0,sK1)),
  inference(resolution,[],[f10398,f148])).
fof(f148,plain,(
  ( ! [X0,X1] : (p(sK5(X1,X0),X0) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((~po(X0,X1) | (p(sK5(X1,X0),X0) & p(sK5(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(skolemisation,[status(esa)],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(rectify,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X2] : (~p(X2,X0) | ~p(X2,X1)) | po(X0,X1)))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X0] : ! [X1] : (po(X0,X1) <=> ? [X6] : (p(X6,X0) & p(X6,X1)))),
  file(unknown,unknown)).
fof(f10398,plain,(
  ( ! [X3] : (~p(sK5(sK1,X3),sK0) | ~po(X3,sK1)) )),
  inference(resolution,[],[f10380,f146])).
fof(f146,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X14] : ! [X15] : (p(X14,X15) <=> (cont(X14,X15) & eqdim(X14,X15)))),
  file(unknown,unknown)).
fof(f10380,plain,(
  ( ! [X1] : (~eqdim(sK5(sK1,X1),sK0) | ~po(X1,sK1)) )),
  inference(resolution,[],[f10370,f149])).
fof(f149,plain,(
  ( ! [X0,X1] : (p(sK5(X1,X0),X1) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f109])).
fof(f10370,plain,(
  ( ! [X3] : (~p(X3,sK1) | ~eqdim(X3,sK0)) )),
  inference(resolution,[],[f10366,f146])).
fof(f10366,plain,(
  ( ! [X0] : (~eqdim(X0,sK1) | ~eqdim(X0,sK0)) )),
  inference(duplicate_literal_removal,[],[f10363])).
fof(f10363,plain,(
  ( ! [X0] : (~eqdim(X0,sK1) | ~eqdim(X0,sK0) | ~eqdim(X0,sK0)) )),
  inference(factoring,[],[f8151])).
fof(f8151,plain,(
  ( ! [X10,X11] : (~eqdim(X11,sK1) | ~eqdim(X11,sK0) | ~eqdim(X10,sK1) | ~eqdim(X10,sK0)) )),
  inference(resolution,[],[f4458,f154])).
fof(f154,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16] : ! [X17] : (eqdim(X16,X17) <=> (leq(X16,X17) & leq(X17,X16)))),
  file(unknown,unknown)).
fof(f4458,plain,(
  ( ! [X10,X11] : (~leq(X10,sK0) | ~eqdim(X10,sK1) | ~eqdim(X11,sK0) | ~eqdim(X11,sK1)) )),
  inference(resolution,[],[f2309,f154])).
fof(f2309,plain,(
  ( ! [X8,X9] : (~leq(X9,sK1) | ~eqdim(X8,sK1) | ~leq(X8,sK0) | ~eqdim(X9,sK0)) )),
  inference(resolution,[],[f1262,f155])).
fof(f155,plain,(
  ( ! [X0,X1] : (leq(X1,X0) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f1262,plain,(
  ( ! [X0,X1] : (~leq(sK0,X1) | ~eqdim(X0,sK1) | ~leq(X1,sK1) | ~leq(X0,sK0)) )),
  inference(resolution,[],[f605,f158])).
fof(f158,plain,(
  ( ! [X2,X0,X1] : (leq(X0,X2) | ~leq(X1,X2) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~leq(X1,X2)) | leq(X0,X2))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & leq(X1,X2)) => leq(X0,X2))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & leq(X1,X2)) => leq(X0,X2))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X39] : ! [X40] : ! [X41] : ((leq(X39,X40) & leq(X40,X41)) => leq(X39,X41))),
  file(unknown,unknown)).
fof(f605,plain,(
  ( ! [X4] : (~leq(sK0,sK1) | ~leq(X4,sK0) | ~eqdim(X4,sK1)) )),
  inference(resolution,[],[f384,f155])).
fof(f384,plain,(
  ( ! [X7] : (~leq(sK1,X7) | ~leq(X7,sK0) | ~leq(sK0,sK1)) )),
  inference(resolution,[],[f158,f336])).
fof(f336,plain,(
  ~leq(sK1,sK0) | ~leq(sK0,sK1)),
  inference(resolution,[],[f156,f116])).
fof(f116,plain,(
  ~eqdim(sK0,sK1)),
  inference(cnf_transformation,[],[f91])).
fof(f156,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f113])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1358
Passive clauses: 3422
Generated clauses: 10316
Final active clauses: 1358
Final passive clauses: 1718
Input formulas: 24
Initial clauses: 47

Unused predicate definitions: 2
Selected by SInE selection: 24
SInE iterations: 1

Duplicate literals: 798
Fw subsumption resolutions: 764

Simple tautologies: 96
Equational tautologies: 23
Forward subsumptions: 5608

Binary resolution: 4821
Factoring: 343
Forward superposition: 3062
Backward superposition: 534
Self superposition: 4

Unique components: 2

Memory used: 1407KB
Time elapsed: 0.155 s
------------------------------
% Success in time 0.59 s
============================= vampire_lin32 ================================
execution finished: Tue Oct 08 20:45:38 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_space_codi/theorems/conversions/po_theorems_goal3.all.tptp
============================ end of footer ===========================
