Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 30000 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2506
Passive clauses: 12180
Generated clauses: 33674
Input formulas: 24
Initial clauses: 56

Selected by SInE selection: 24
SInE iterations: 1

Duplicate literals: 998
Fw subsumption resolutions: 3236

Simple tautologies: 742
Forward subsumptions: 16531

Binary resolution: 25891
Factoring: 486

Backtracking splits: 1397
Backtracking splits refuted: 452
Backtracking splits refuted at zero level: 99

Memory used: 4605KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 29995 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f2616,plain,(
  $false),
  inference(subsumption_resolution,[],[f2615,f115])).
fof(f115,plain,(
  sc(sK0,sK1)),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  sc(sK0,sK1) & ~sc(sK1,sK0)),
  inference(skolemisation,[status(esa)],[f65])).
fof(f65,plain,(
  ? [X0,X1] : (sc(X0,X1) & ~sc(X1,X0))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1] : (sc(X0,X1) => sc(X1,X0))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : (sc(X0,X1) => sc(X1,X0))),
  file(unknown,unknown)).
fof(f2615,plain,(
  ~sc(sK0,sK1)),
  inference(duplicate_literal_removal,[],[f2607])).
fof(f2607,plain,(
  ~sc(sK0,sK1) | ~sc(sK0,sK1)),
  inference(resolution,[],[f2533,f151])).
fof(f151,plain,(
  ( ! [X0,X1] : (cont(sK5(X1,X0),X0) | ~sc(X0,X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | ((cont(sK5(X1,X0),X0) & cont(sK5(X1,X0),X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | (cont(sK6(X1,X0),X0) & cont(sK6(X1,X0),X1) & (~less(sK6(X1,X0),X0) | ~less(sK6(X1,X0),X1))) | sc(X0,X1)))),
  inference(skolemisation,[status(esa)],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | ? [X5] : (cont(X5,X0) & cont(X5,X1) & (~less(X5,X0) | ~less(X5,X1))) | sc(X0,X1)))),
  inference(rectify,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | ? [X3] : (cont(X3,X0) & cont(X3,X1) & (~less(X3,X0) | ~less(X3,X1))) | sc(X0,X1)))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & ((! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | ? [X3] : (cont(X3,X0) & cont(X3,X1) & (~less(X3,X0) | ~less(X3,X1)))) | sc(X0,X1)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((~cont(X3,X0) | ~cont(X3,X1)) | (less(X3,X0) & less(X3,X1)))))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((cont(X3,X0) & cont(X3,X1)) => (less(X3,X0) & less(X3,X1)))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((cont(X3,X0) & cont(X3,X1)) => (less(X3,X0) & less(X3,X1)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X43] : ! [X44] : (sc(X43,X44) <=> (? [X45] : (cont(X45,X43) & cont(X45,X44)) & ! [X45] : ((cont(X45,X43) & cont(X45,X44)) => (less(X45,X43) & less(X45,X44)))))),
  file(unknown,unknown)).
fof(f2533,plain,(
  ( ! [X5] : (~cont(sK5(sK1,X5),sK0) | ~sc(X5,sK1)) )),
  inference(resolution,[],[f2527,f152])).
fof(f152,plain,(
  ( ! [X0,X1] : (cont(sK5(X1,X0),X1) | ~sc(X0,X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f2527,plain,(
  ( ! [X55] : (~cont(X55,sK1) | ~cont(X55,sK0)) )),
  inference(subsumption_resolution,[],[f2514,f115])).
fof(f2514,plain,(
  ( ! [X55] : (~cont(X55,sK1) | ~cont(X55,sK0) | ~sc(sK0,sK1)) )),
  inference(resolution,[],[f2302,f116])).
fof(f116,plain,(
  ~sc(sK1,sK0)),
  inference(cnf_transformation,[],[f92])).
fof(f2302,plain,(
  ( ! [X2,X0,X1] : (sc(X1,X0) | ~cont(X2,X1) | ~cont(X2,X0) | ~sc(X0,X1)) )),
  inference(subsumption_resolution,[],[f2301,f156])).
fof(f156,plain,(
  ( ! [X4,X0,X1] : (cont(sK6(X1,X0),X1) | sc(X0,X1) | ~cont(X4,X1) | ~cont(X4,X0)) )),
  inference(cnf_transformation,[],[f113])).
fof(f2301,plain,(
  ( ! [X2,X0,X1] : (~cont(sK6(X0,X1),X0) | ~sc(X0,X1) | ~cont(X2,X1) | ~cont(X2,X0) | sc(X1,X0)) )),
  inference(subsumption_resolution,[],[f2300,f155])).
fof(f155,plain,(
  ( ! [X4,X0,X1] : (cont(sK6(X1,X0),X0) | sc(X0,X1) | ~cont(X4,X1) | ~cont(X4,X0)) )),
  inference(cnf_transformation,[],[f113])).
fof(f2300,plain,(
  ( ! [X2,X0,X1] : (~cont(sK6(X0,X1),X1) | ~cont(sK6(X0,X1),X0) | ~sc(X0,X1) | ~cont(X2,X1) | ~cont(X2,X0) | sc(X1,X0)) )),
  inference(duplicate_literal_removal,[],[f2290])).
fof(f2290,plain,(
  ( ! [X2,X0,X1] : (~cont(sK6(X0,X1),X1) | ~cont(sK6(X0,X1),X0) | ~sc(X0,X1) | sc(X1,X0) | ~sc(X0,X1) | sc(X1,X0) | ~cont(X2,X0) | ~cont(X2,X1)) )),
  inference(resolution,[],[f1245,f155])).
fof(f1245,plain,(
  ( ! [X2,X0,X1] : (~cont(sK6(X0,X1),X2) | ~cont(sK6(X0,X1),X1) | ~cont(sK6(X0,X1),X0) | ~sc(X0,X2) | sc(X1,X0) | ~sc(X0,X1)) )),
  inference(factoring,[],[f543])).
fof(f543,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cont(sK6(X1,X0),X4) | ~cont(X2,X1) | ~cont(X2,X0) | ~cont(sK6(X1,X0),X3) | ~sc(X1,X3) | sc(X0,X1) | ~sc(X4,X0)) )),
  inference(subsumption_resolution,[],[f539,f155])).
fof(f539,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sc(X0,X1) | ~cont(X2,X1) | ~cont(X2,X0) | ~cont(sK6(X1,X0),X3) | ~sc(X1,X3) | ~cont(sK6(X1,X0),X0) | ~cont(sK6(X1,X0),X4) | ~sc(X4,X0)) )),
  inference(resolution,[],[f424,f154])).
fof(f154,plain,(
  ( ! [X0,X3,X1] : (less(X3,X1) | ~cont(X3,X1) | ~cont(X3,X0) | ~sc(X0,X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f424,plain,(
  ( ! [X6,X4,X7,X5] : (~less(sK6(X5,X4),X4) | sc(X4,X5) | ~cont(X6,X5) | ~cont(X6,X4) | ~cont(sK6(X5,X4),X7) | ~sc(X5,X7)) )),
  inference(subsumption_resolution,[],[f418,f156])).
fof(f418,plain,(
  ( ! [X6,X4,X7,X5] : (sc(X4,X5) | ~less(sK6(X5,X4),X4) | ~cont(X6,X5) | ~cont(X6,X4) | ~cont(sK6(X5,X4),X7) | ~cont(sK6(X5,X4),X5) | ~sc(X5,X7)) )),
  inference(resolution,[],[f157,f153])).
fof(f153,plain,(
  ( ! [X0,X3,X1] : (less(X3,X0) | ~cont(X3,X1) | ~cont(X3,X0) | ~sc(X0,X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f157,plain,(
  ( ! [X4,X0,X1] : (~less(sK6(X1,X0),X1) | sc(X0,X1) | ~less(sK6(X1,X0),X0) | ~cont(X4,X1) | ~cont(X4,X0)) )),
  inference(cnf_transformation,[],[f113])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 368
Passive clauses: 1025
Generated clauses: 2499
Final active clauses: 368
Final passive clauses: 517
Input formulas: 24
Initial clauses: 48

Unused predicate definitions: 3
Selected by SInE selection: 24
SInE iterations: 1

Duplicate literals: 152
Fw subsumption resolutions: 268

Simple tautologies: 74
Equational tautologies: 6
Forward subsumptions: 1112

Binary resolution: 1518
Factoring: 137
Forward superposition: 368
Backward superposition: 14

Unique components: 2

Memory used: 511KB
Time elapsed: 0.042 s
------------------------------
% Success in time 0.481 s
============================= vampire_lin32 ================================
execution finished: Tue Oct 08 20:43:25 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_space_codi/theorems/conversions/sc_theorems_goal6.all.tptp
============================ end of footer ===========================
