Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 30000 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f1303,plain,(
  $false),
  inference(subsumption_resolution,[],[f1299,f120])).
fof(f120,plain,(
  sc(sK0,sK1)),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  sc(sK0,sK1) & (cont(sK2,sK0) & p(sK2,sK1))),
  inference(skolemisation,[status(esa)],[f66])).
fof(f66,plain,(
  ? [X0,X1] : (sc(X0,X1) & ? [X2] : (cont(X2,X0) & p(X2,X1)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1] : (sc(X0,X1) => ~? [X2] : (cont(X2,X0) & p(X2,X1)))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : (sc(X0,X1) => ~? [X2] : (cont(X2,X0) & p(X2,X1)))),
  file(unknown,unknown)).
fof(f1299,plain,(
  ~sc(sK0,sK1)),
  inference(resolution,[],[f1297,f121])).
fof(f121,plain,(
  cont(sK2,sK0)),
  inference(cnf_transformation,[],[f93])).
fof(f1297,plain,(
  ( ! [X6] : (~cont(sK2,X6) | ~sc(X6,sK1)) )),
  inference(resolution,[],[f552,f122])).
fof(f122,plain,(
  p(sK2,sK1)),
  inference(cnf_transformation,[],[f93])).
fof(f552,plain,(
  ( ! [X6,X4,X5] : (~p(X4,X5) | ~sc(X6,X5) | ~cont(X4,X6)) )),
  inference(subsumption_resolution,[],[f537,f155])).
fof(f155,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X18] : ! [X19] : (p(X18,X19) <=> (cont(X18,X19) & eqdim(X18,X19)))),
  file(unknown,unknown)).
fof(f537,plain,(
  ( ! [X6,X4,X5] : (~cont(X4,X5) | ~cont(X4,X6) | ~sc(X6,X5) | ~p(X4,X5)) )),
  inference(resolution,[],[f161,f189])).
fof(f189,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~p(X0,X1)) )),
  inference(resolution,[],[f169,f156])).
fof(f156,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f169,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ((~less(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | less(X0,X1)))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ((~less(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | less(X0,X1)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : (less(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : (less(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X41] : ! [X42] : (less(X41,X42) <=> (leq(X41,X42) & ~eqdim(X41,X42)))),
  file(unknown,unknown)).
fof(f161,plain,(
  ( ! [X0,X3,X1] : (less(X3,X1) | ~cont(X3,X1) | ~cont(X3,X0) | ~sc(X0,X1)) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | ((cont(sK6(X1,X0),X0) & cont(sK6(X1,X0),X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | (cont(sK7(X1,X0),X0) & cont(sK7(X1,X0),X1) & (~less(sK7(X1,X0),X0) | ~less(sK7(X1,X0),X1))) | sc(X0,X1)))),
  inference(skolemisation,[status(esa)],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | ? [X5] : (cont(X5,X0) & cont(X5,X1) & (~less(X5,X0) | ~less(X5,X1))) | sc(X0,X1)))),
  inference(rectify,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | ? [X3] : (cont(X3,X0) & cont(X3,X1) & (~less(X3,X0) | ~less(X3,X1))) | sc(X0,X1)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & ((! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | ? [X3] : (cont(X3,X0) & cont(X3,X1) & (~less(X3,X0) | ~less(X3,X1)))) | sc(X0,X1)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((~cont(X3,X0) | ~cont(X3,X1)) | (less(X3,X0) & less(X3,X1)))))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((cont(X3,X0) & cont(X3,X1)) => (less(X3,X0) & less(X3,X1)))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((cont(X3,X0) & cont(X3,X1)) => (less(X3,X0) & less(X3,X1)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X45] : ! [X46] : (sc(X45,X46) <=> (? [X47] : (cont(X47,X45) & cont(X47,X46)) & ! [X47] : ((cont(X47,X45) & cont(X47,X46)) => (less(X47,X45) & less(X47,X46)))))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 292
Passive clauses: 648
Generated clauses: 1183
Final active clauses: 239
Final passive clauses: 262
Input formulas: 24
Initial clauses: 57

Selected by SInE selection: 24
SInE iterations: 1

Duplicate literals: 8
Fw subsumption resolutions: 113

Simple tautologies: 29
Forward subsumptions: 361

Binary resolution: 897
Factoring: 14

Backtracking splits: 64
Backtracking splits refuted: 14
Backtracking splits refuted at zero level: 6

Memory used: 383KB
Time elapsed: 0.013 s
------------------------------
% Success in time 0.027 s
============================= vampire_lin32 ================================
execution finished: Tue Oct 08 20:43:11 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_space_codi/theorems/conversions/sc_theorems_goal2.all.tptp
============================ end of footer ===========================
