Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 30000 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f1308,plain,(
  $false),
  inference(subsumption_resolution,[],[f1306,f119])).
fof(f119,plain,(
  sc(sK0,sK1)),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  sc(sK0,sK1) & (p(sK2,sK0) & cont(sK2,sK1))),
  inference(skolemisation,[status(esa)],[f65])).
fof(f65,plain,(
  ? [X0,X1] : (sc(X0,X1) & ? [X2] : (p(X2,X0) & cont(X2,X1)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1] : (sc(X0,X1) => ~? [X2] : (p(X2,X0) & cont(X2,X1)))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : (sc(X0,X1) => ~? [X2] : (p(X2,X0) & cont(X2,X1)))),
  file(unknown,unknown)).
fof(f1306,plain,(
  ~sc(sK0,sK1)),
  inference(resolution,[],[f1266,f120])).
fof(f120,plain,(
  p(sK2,sK0)),
  inference(cnf_transformation,[],[f92])).
fof(f1266,plain,(
  ( ! [X9] : (~p(sK2,X9) | ~sc(X9,sK1)) )),
  inference(resolution,[],[f531,f121])).
fof(f121,plain,(
  cont(sK2,sK1)),
  inference(cnf_transformation,[],[f92])).
fof(f531,plain,(
  ( ! [X6,X4,X5] : (~cont(X4,X5) | ~sc(X6,X5) | ~p(X4,X6)) )),
  inference(subsumption_resolution,[],[f517,f154])).
fof(f154,plain,(
  ( ! [X0,X1] : (cont(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X17] : ! [X18] : (p(X17,X18) <=> (cont(X17,X18) & eqdim(X17,X18)))),
  file(unknown,unknown)).
fof(f517,plain,(
  ( ! [X6,X4,X5] : (~cont(X4,X5) | ~cont(X4,X6) | ~sc(X6,X5) | ~p(X4,X6)) )),
  inference(resolution,[],[f159,f190])).
fof(f190,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~p(X0,X1)) )),
  inference(resolution,[],[f168,f155])).
fof(f155,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f109])).
fof(f168,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ((~less(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | less(X0,X1)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ((~less(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | less(X0,X1)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : (less(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : (less(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X40] : ! [X41] : (less(X40,X41) <=> (leq(X40,X41) & ~eqdim(X40,X41)))),
  file(unknown,unknown)).
fof(f159,plain,(
  ( ! [X0,X3,X1] : (less(X3,X0) | ~cont(X3,X1) | ~cont(X3,X0) | ~sc(X0,X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | ((cont(sK6(X1,X0),X0) & cont(sK6(X1,X0),X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | (cont(sK7(X1,X0),X0) & cont(sK7(X1,X0),X1) & (~less(sK7(X1,X0),X0) | ~less(sK7(X1,X0),X1))) | sc(X0,X1)))),
  inference(skolemisation,[status(esa)],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | ? [X5] : (cont(X5,X0) & cont(X5,X1) & (~less(X5,X0) | ~less(X5,X1))) | sc(X0,X1)))),
  inference(rectify,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | ? [X3] : (cont(X3,X0) & cont(X3,X1) & (~less(X3,X0) | ~less(X3,X1))) | sc(X0,X1)))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & ((! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | ? [X3] : (cont(X3,X0) & cont(X3,X1) & (~less(X3,X0) | ~less(X3,X1)))) | sc(X0,X1)))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1)))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((~cont(X3,X0) | ~cont(X3,X1)) | (less(X3,X0) & less(X3,X1)))))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((cont(X3,X0) & cont(X3,X1)) => (less(X3,X0) & less(X3,X1)))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((cont(X3,X0) & cont(X3,X1)) => (less(X3,X0) & less(X3,X1)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X44] : ! [X45] : (sc(X44,X45) <=> (? [X46] : (cont(X46,X44) & cont(X46,X45)) & ! [X46] : ((cont(X46,X44) & cont(X46,X45)) => (less(X46,X44) & less(X46,X45)))))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 292
Passive clauses: 666
Generated clauses: 1192
Final active clauses: 222
Final passive clauses: 264
Input formulas: 24
Initial clauses: 57

Selected by SInE selection: 24
SInE iterations: 1

Duplicate literals: 13
Fw subsumption resolutions: 117

Simple tautologies: 29
Forward subsumptions: 353

Binary resolution: 899
Factoring: 13

Backtracking splits: 53
Backtracking splits refuted: 18
Backtracking splits refuted at zero level: 6

Memory used: 383KB
Time elapsed: 0.014 s
------------------------------
% Success in time 0.027 s
============================= vampire_lin32 ================================
execution finished: Tue Oct 08 20:43:13 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_space_codi/theorems/conversions/sc_theorems_goal3.all.tptp
============================ end of footer ===========================
