Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 29999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f809,plain,(
  $false),
  inference(subsumption_resolution,[],[f806,f413])).
fof(f413,plain,(
  atomicraycurve(sK0)),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  atomicraycurve(sK0) & ~icon(sK0)),
  inference(skolemisation,[status(esa)],[f229])).
fof(f229,plain,(
  ? [X0] : (atomicraycurve(X0) & ~icon(X0))),
  inference(ennf_transformation,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : (atomicraycurve(X0) => icon(X0))),
  file(unknown,unknown)).
fof(f806,plain,(
  ~atomicraycurve(sK0)),
  inference(resolution,[],[f801,f449])).
fof(f449,plain,(
  ( ! [X0,X1] : (~pp(X1,X0) | ~atomicraycurve(X0)) )),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0] : ((~atomicraycurve(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & (endpoint(sK10(X0),X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | sK10(X0) = X3)))) & (~curve(X0) | pp(sK11(X0),X0) | ! [X5] : (~endpoint(X5,X0) | (point(sK12(X5,X0)) & cont(sK12(X5,X0),X0) & sK12(X5,X0) != X5)) | atomicraycurve(X0)))),
  inference(skolemisation,[status(esa)],[f356])).
fof(f356,plain,(
  ! [X0] : ((~atomicraycurve(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (endpoint(X2,X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | X2 = X3)))) & (~curve(X0) | ? [X4] : pp(X4,X0) | ! [X5] : (~endpoint(X5,X0) | ? [X6] : (point(X6) & cont(X6,X0) & X5 != X6)) | atomicraycurve(X0)))),
  inference(rectify,[],[f355])).
fof(f355,plain,(
  ! [X0] : ((~atomicraycurve(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (endpoint(X2,X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | X2 = X3)))) & (~curve(X0) | ? [X1] : pp(X1,X0) | ! [X2] : (~endpoint(X2,X0) | ? [X3] : (point(X3) & cont(X3,X0) & X2 != X3)) | atomicraycurve(X0)))),
  inference(flattening,[],[f354])).
fof(f354,plain,(
  ! [X0] : ((~atomicraycurve(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (endpoint(X2,X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | X2 = X3)))) & ((~curve(X0) | ? [X1] : pp(X1,X0) | ! [X2] : (~endpoint(X2,X0) | ? [X3] : (point(X3) & cont(X3,X0) & X2 != X3))) | atomicraycurve(X0)))),
  inference(nnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0] : (atomicraycurve(X0) <=> (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (endpoint(X2,X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | X2 = X3))))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0] : (atomicraycurve(X0) <=> (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (endpoint(X2,X0) & ! [X3] : ((~point(X3) | ~cont(X3,X0)) | X2 = X3))))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : (atomicraycurve(X0) <=> (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (endpoint(X2,X0) & ! [X3] : ((point(X3) & cont(X3,X0)) => X2 = X3))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X20] : (atomicraycurve(X20) <=> (curve(X20) & ! [X21] : ~pp(X21,X20) & ? [X21] : (endpoint(X21,X20) & ! [X22] : ((point(X22) & cont(X22,X20)) => X21 = X22))))),
  file(unknown,unknown)).
fof(f801,plain,(
  pp(sK8(sK0),sK0)),
  inference(backtracking_split_refutation,[],[f790,f791_D,f800])).
fof(f800,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f797,f413])).
fof(f797,plain,(
  ~atomicraycurve(sK0) | $spl4),
  inference(resolution,[],[f792,f449])).
fof(f792,plain,(
  pp(sK6(sK0),sK0) | $spl4),
  inference(resolution,[],[f791,f433])).
fof(f433,plain,(
  ( ! [X0] : (con(X0) | pp(sK6(X0),X0)) )),
  inference(cnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0] : ((~con(X0) | ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1)))) & ((pp(sK6(X0),X0) & ~sc(sK6(X0),difference(X0,sK6(X0)))) | con(X0)))),
  inference(skolemisation,[status(esa)],[f340])).
fof(f340,plain,(
  ! [X0] : ((~con(X0) | ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1)))) & (? [X2] : (pp(X2,X0) & ~sc(X2,difference(X0,X2))) | con(X0)))),
  inference(rectify,[],[f339])).
fof(f339,plain,(
  ! [X0] : ((~con(X0) | ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1)))) & (? [X1] : (pp(X1,X0) & ~sc(X1,difference(X0,X1))) | con(X0)))),
  inference(nnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0] : (con(X0) <=> ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1))))),
  inference(ennf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : (con(X0) <=> ! [X1] : (pp(X1,X0) => sc(X1,difference(X0,X1))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X24] : (con(X24) <=> ! [X25] : (pp(X25,X24) => sc(X25,difference(X24,X25))))),
  file(unknown,unknown)).
fof(f791,plain,(
  ~con(sK0) | $spl4),
  inference(cnf_transformation,[],[f791_D])).
fof(f791_D,plain,(
  ~con(sK0) <=> ~$spl4),
  introduced(backtracking_splitting_component,[])).
fof(f790,plain,(
  pp(sK8(sK0),sK0) | ~con(sK0)),
  inference(resolution,[],[f441,f414])).
fof(f414,plain,(
  ~icon(sK0)),
  inference(cnf_transformation,[],[f324])).
fof(f441,plain,(
  ( ! [X0] : (icon(X0) | pp(sK8(X0),X0) | ~con(X0)) )),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & (~con(X0) | (pp(sK8(X0),X0) & ~covers(X0,intersection(sK8(X0),difference(X0,sK8(X0))))) | icon(X0)))),
  inference(skolemisation,[status(esa)],[f347])).
fof(f347,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & (~con(X0) | ? [X2] : (pp(X2,X0) & ~covers(X0,intersection(X2,difference(X0,X2)))) | icon(X0)))),
  inference(rectify,[],[f346])).
fof(f346,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & (~con(X0) | ? [X1] : (pp(X1,X0) & ~covers(X0,intersection(X1,difference(X0,X1)))) | icon(X0)))),
  inference(flattening,[],[f345])).
fof(f345,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & ((~con(X0) | ? [X1] : (pp(X1,X0) & ~covers(X0,intersection(X1,difference(X0,X1))))) | icon(X0)))),
  inference(nnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0] : (icon(X0) <=> (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1))))))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : (icon(X0) <=> (con(X0) & ! [X1] : (pp(X1,X0) => covers(X0,intersection(X1,difference(X0,X1))))))),
  inference(rectify,[],[f6])).
fof(f6,axiom,(
  ! [X8] : (icon(X8) <=> (con(X8) & ! [X9] : (pp(X9,X8) => covers(X8,intersection(X9,difference(X8,X9))))))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 164
Passive clauses: 305
Generated clauses: 386
Final active clauses: 161
Final passive clauses: 125
Input formulas: 84
Initial clauses: 175

Selected by SInE selection: 84
SInE iterations: 1

Duplicate literals: 2
Fw subsumption resolutions: 4

Simple tautologies: 8
Forward subsumptions: 68

Binary resolution: 195
Factoring: 2

Backtracking splits: 4
Backtracking splits refuted: 2
Backtracking splits refuted at zero level: 2

Memory used: 383KB
Time elapsed: 0.012 s
------------------------------
% Success in time 0.04 s
============================= vampire_lin32 ================================
execution finished: Tue Oct 08 20:02:45 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_mereotopology_codib/theorems/conversions/atomic_ray_curve_theorems_goal3.all.tptp
============================ end of footer ===========================
