Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 29999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f901,plain,(
  $false),
  inference(subsumption_resolution,[],[f900,f417])).
fof(f417,plain,(
  atomiccurvesegment(sK0)),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  atomiccurvesegment(sK0) & ~icon(sK0)),
  inference(skolemisation,[status(esa)],[f230])).
fof(f230,plain,(
  ? [X0] : (atomiccurvesegment(X0) & ~icon(X0))),
  inference(ennf_transformation,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : (atomiccurvesegment(X0) => icon(X0))),
  file(unknown,unknown)).
fof(f900,plain,(
  ~atomiccurvesegment(sK0)),
  inference(resolution,[],[f892,f462])).
fof(f462,plain,(
  ( ! [X0] : (sP0(X0) | ~atomiccurvesegment(X0)) )),
  inference(cnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X0] : ((~atomiccurvesegment(X0) | sP0(X0)) & (~sP0(X0) | atomiccurvesegment(X0)))),
  inference(nnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0] : (atomiccurvesegment(X0) <=> sP0(X0))),
  inference(definition_folding,[],[f323])).
fof(f323,plain,(
  ! [X0] : (sP0(X0) <=> (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2,X3] : (endpoint(X2,X0) & endpoint(X3,X0) & X2 != X3 & ! [X4] : (~point(X4) | ~cont(X4,X0) | X2 = X4 | X3 = X4))))),
  introduced(predicate_definition_introduction,[])).
fof(f892,plain,(
  ~sP0(sK0)),
  inference(resolution,[],[f885,f453])).
fof(f453,plain,(
  ( ! [X0,X1] : (~pp(X1,X0) | ~sP0(X0)) )),
  inference(cnf_transformation,[],[f360])).
fof(f360,plain,(
  ! [X0] : ((~sP0(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & (endpoint(sK10(X0),X0) & endpoint(sK11(X0),X0) & sK10(X0) != sK11(X0) & ! [X4] : (~point(X4) | ~cont(X4,X0) | sK10(X0) = X4 | sK11(X0) = X4)))) & (~curve(X0) | pp(sK12(X0),X0) | ! [X6,X7] : (~endpoint(X6,X0) | ~endpoint(X7,X0) | X6 = X7 | (point(sK13(X7,X6,X0)) & cont(sK13(X7,X6,X0),X0) & sK13(X7,X6,X0) != X6 & sK13(X7,X6,X0) != X7)) | sP0(X0)))),
  inference(skolemisation,[status(esa)],[f359])).
fof(f359,plain,(
  ! [X0] : ((~sP0(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2,X3] : (endpoint(X2,X0) & endpoint(X3,X0) & X2 != X3 & ! [X4] : (~point(X4) | ~cont(X4,X0) | X2 = X4 | X3 = X4)))) & (~curve(X0) | ? [X5] : pp(X5,X0) | ! [X6,X7] : (~endpoint(X6,X0) | ~endpoint(X7,X0) | X6 = X7 | ? [X8] : (point(X8) & cont(X8,X0) & X6 != X8 & X7 != X8)) | sP0(X0)))),
  inference(rectify,[],[f358])).
fof(f358,plain,(
  ! [X0] : ((~sP0(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2,X3] : (endpoint(X2,X0) & endpoint(X3,X0) & X2 != X3 & ! [X4] : (~point(X4) | ~cont(X4,X0) | X2 = X4 | X3 = X4)))) & (~curve(X0) | ? [X1] : pp(X1,X0) | ! [X2,X3] : (~endpoint(X2,X0) | ~endpoint(X3,X0) | X2 = X3 | ? [X4] : (point(X4) & cont(X4,X0) & X2 != X4 & X3 != X4)) | sP0(X0)))),
  inference(flattening,[],[f357])).
fof(f357,plain,(
  ! [X0] : ((~sP0(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2,X3] : (endpoint(X2,X0) & endpoint(X3,X0) & X2 != X3 & ! [X4] : (~point(X4) | ~cont(X4,X0) | X2 = X4 | X3 = X4)))) & ((~curve(X0) | ? [X1] : pp(X1,X0) | ! [X2,X3] : (~endpoint(X2,X0) | ~endpoint(X3,X0) | X2 = X3 | ? [X4] : (point(X4) & cont(X4,X0) & X2 != X4 & X3 != X4))) | sP0(X0)))),
  inference(nnf_transformation,[],[f323])).
fof(f885,plain,(
  pp(sK8(sK0),sK0)),
  inference(backtracking_split_refutation,[],[f871,f872_D,f884])).
fof(f884,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f883,f417])).
fof(f883,plain,(
  ~atomiccurvesegment(sK0) | $spl4),
  inference(resolution,[],[f877,f462])).
fof(f877,plain,(
  ~sP0(sK0) | $spl4),
  inference(resolution,[],[f873,f453])).
fof(f873,plain,(
  pp(sK6(sK0),sK0) | $spl4),
  inference(resolution,[],[f872,f437])).
fof(f437,plain,(
  ( ! [X0] : (con(X0) | pp(sK6(X0),X0)) )),
  inference(cnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0] : ((~con(X0) | ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1)))) & ((pp(sK6(X0),X0) & ~sc(sK6(X0),difference(X0,sK6(X0)))) | con(X0)))),
  inference(skolemisation,[status(esa)],[f343])).
fof(f343,plain,(
  ! [X0] : ((~con(X0) | ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1)))) & (? [X2] : (pp(X2,X0) & ~sc(X2,difference(X0,X2))) | con(X0)))),
  inference(rectify,[],[f342])).
fof(f342,plain,(
  ! [X0] : ((~con(X0) | ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1)))) & (? [X1] : (pp(X1,X0) & ~sc(X1,difference(X0,X1))) | con(X0)))),
  inference(nnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0] : (con(X0) <=> ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1))))),
  inference(ennf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : (con(X0) <=> ! [X1] : (pp(X1,X0) => sc(X1,difference(X0,X1))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X131] : (con(X131) <=> ! [X132] : (pp(X132,X131) => sc(X132,difference(X131,X132))))),
  file(unknown,unknown)).
fof(f872,plain,(
  ~con(sK0) | $spl4),
  inference(cnf_transformation,[],[f872_D])).
fof(f872_D,plain,(
  ~con(sK0) <=> ~$spl4),
  introduced(backtracking_splitting_component,[])).
fof(f871,plain,(
  pp(sK8(sK0),sK0) | ~con(sK0)),
  inference(resolution,[],[f445,f418])).
fof(f418,plain,(
  ~icon(sK0)),
  inference(cnf_transformation,[],[f327])).
fof(f445,plain,(
  ( ! [X0] : (icon(X0) | pp(sK8(X0),X0) | ~con(X0)) )),
  inference(cnf_transformation,[],[f351])).
fof(f351,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & (~con(X0) | (pp(sK8(X0),X0) & ~covers(X0,intersection(sK8(X0),difference(X0,sK8(X0))))) | icon(X0)))),
  inference(skolemisation,[status(esa)],[f350])).
fof(f350,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & (~con(X0) | ? [X2] : (pp(X2,X0) & ~covers(X0,intersection(X2,difference(X0,X2)))) | icon(X0)))),
  inference(rectify,[],[f349])).
fof(f349,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & (~con(X0) | ? [X1] : (pp(X1,X0) & ~covers(X0,intersection(X1,difference(X0,X1)))) | icon(X0)))),
  inference(flattening,[],[f348])).
fof(f348,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & ((~con(X0) | ? [X1] : (pp(X1,X0) & ~covers(X0,intersection(X1,difference(X0,X1))))) | icon(X0)))),
  inference(nnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0] : (icon(X0) <=> (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1))))))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : (icon(X0) <=> (con(X0) & ! [X1] : (pp(X1,X0) => covers(X0,intersection(X1,difference(X0,X1))))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X155] : (icon(X155) <=> (con(X155) & ! [X156] : (pp(X156,X155) => covers(X155,intersection(X156,difference(X155,X156))))))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 191
Passive clauses: 380
Generated clauses: 470
Final active clauses: 187
Final passive clauses: 171
Input formulas: 84
Initial clauses: 180

Selected by SInE selection: 84
SInE iterations: 1

Duplicate literals: 3
Fw subsumption resolutions: 5

Simple tautologies: 11
Forward subsumptions: 76

Binary resolution: 273
Factoring: 2

Backtracking splits: 5
Backtracking splits refuted: 1
Backtracking splits refuted at zero level: 1

Memory used: 383KB
Time elapsed: 0.008 s
------------------------------
% Success in time 0.025 s
============================= vampire_lin32 ================================
execution finished: Tue Oct 08 20:02:16 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_mereotopology_codib/theorems/conversions/atomic_curve_segment_theorems_goal3.all.tptp
============================ end of footer ===========================
