Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 59999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\GitHub\colore\ontologies\simple_features\conversions\sfc_fol_nontrivial.all
Refutation found. Thanks to Tanya!
9920. $false (2:0) [subsumption resolution 9919,7990]
7990. ~sf_contains(sK104,sK103) (2:3) [resolution 7984,1231]
1231. cont(X0,X1) | ~sf_contains(X0,X1) (0:6) [cnf transformation 854]
854. ! [X0,X1] : ((~sf_contains(X0,X1) | (sf_geometry(X0) & sf_geometry(X1) & cont(X0,X1))) & (~sf_geometry(X0) | ~sf_geometry(X1) | ~cont(X0,X1) | sf_contains(X0,X1)))[flattening 853]
853. ! [X0,X1] : ((~sf_contains(X0,X1) | (sf_geometry(X0) & sf_geometry(X1) & cont(X0,X1))) & ((~sf_geometry(X0) | ~sf_geometry(X1) | ~cont(X0,X1)) | sf_contains(X0,X1)))[nnf transformation 346]
346. ! [X0,X1] : (sf_contains(X0,X1) <=> (sf_geometry(X0) & sf_geometry(X1) & cont(X0,X1)))[flattening 345]
345. ! [X0] : ! [X1] : (sf_contains(X0,X1) <=> (sf_geometry(X0) & sf_geometry(X1) & cont(X0,X1)))[rectify 170]
170. ! [X162] : ! [X163] : (sf_contains(X162,X163) <=> (sf_geometry(X162) & sf_geometry(X163) & cont(X162,X163)))[input]
7984. ~cont(sK104,sK103) (1:3) [backtracking split refutation 5178,5270,7983]
7983. $false {659} (2:0) [subsumption resolution 7982,2022]
2022. sf_contains(sK103,sK104) (1:3) [resolution 1217,1382]
1382. sf_equals(sK103,sK104) (0:3) [cnf transformation 967]
967. sf_equals(sK103,sK104) & sK103 != sK104[skolemisation 709]
709. ? [X0,X1] : (sf_equals(X0,X1) & X0 != X1)[ennf transformation 471]
471. ? [X0,X1] : (sf_equals(X0,X1) & ~X0 = X1)[flattening 470]
470. ? [X0] : ? [X1] : (sf_equals(X0,X1) & ~X0 = X1)[rectify 97]
97. ? [X132] : ? [X133] : (sf_equals(X132,X133) & ~X132 = X133)[input]
1217. ~sf_equals(X0,X1) | sf_contains(X0,X1) (0:6) [cnf transformation 847]
847. ! [X0,X1] : ((~sf_equals(X0,X1) | (sf_contains(X0,X1) & sf_within(X0,X1))) & (~sf_contains(X0,X1) | ~sf_within(X0,X1) | sf_equals(X0,X1)))[flattening 846]
846. ! [X0,X1] : ((~sf_equals(X0,X1) | (sf_contains(X0,X1) & sf_within(X0,X1))) & ((~sf_contains(X0,X1) | ~sf_within(X0,X1)) | sf_equals(X0,X1)))[nnf transformation 338]
338. ! [X0,X1] : (sf_equals(X0,X1) <=> (sf_contains(X0,X1) & sf_within(X0,X1)))[flattening 337]
337. ! [X0] : ! [X1] : (sf_equals(X0,X1) <=> (sf_contains(X0,X1) & sf_within(X0,X1)))[rectify 193]
193. ! [X162] : ! [X163] : (sf_equals(X162,X163) <=> (sf_contains(X162,X163) & sf_within(X162,X163)))[input]
7982. ~sf_contains(sK103,sK104) {659} (2:3) [resolution 5270,1231]
5270. ~cont(sK103,sK104) {659} (1:3) [backtracking splitting component]
5178. ~cont(sK104,sK103) | ~cont(sK103,sK104) (1:6) [resolution 1478,1482]
1482. ~sQ1_eqProxy(sK103,sK104) (0:3) [equality proxy replacement 1383]
1383. sK103 != sK104 (0:3) [cnf transformation 967]
1478. sQ1_eqProxy(X0,X1) | ~cont(X1,X0) | ~cont(X0,X1) (0:9) [equality proxy replacement 1179]
1179. X0 = X1 | ~cont(X1,X0) | ~cont(X0,X1) (0:9) [cnf transformation 638]
638. ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1)[flattening 637]
637. ! [X0,X1] : ((~cont(X0,X1) | ~cont(X1,X0)) | X0 = X1)[ennf transformation 324]
324. ! [X0,X1] : ((cont(X0,X1) & cont(X1,X0)) => X0 = X1)[flattening 323]
323. ! [X0] : ! [X1] : ((cont(X0,X1) & cont(X1,X0)) => X0 = X1)[rectify 140]
140. ! [X162] : ! [X163] : ((cont(X162,X163) & cont(X163,X162)) => X162 = X163)[input]
9919. sf_contains(sK104,sK103) (2:3) [resolution 2027,1382]
2027. ~sf_equals(X1,X0) | sf_contains(X0,X1) (1:6) [resolution 1223,1218]
1218. sf_within(X0,X1) | ~sf_equals(X0,X1) (0:6) [cnf transformation 847]
1223. ~sf_within(X1,X0) | sf_contains(X0,X1) (0:6) [cnf transformation 850]
850. ! [X0,X1] : ((~sf_within(X1,X0) | sf_contains(X0,X1)) & (~sf_contains(X0,X1) | sf_within(X1,X0)))[nnf transformation 342]
342. ! [X0,X1] : (sf_within(X1,X0) <=> sf_contains(X0,X1))[flattening 341]
341. ! [X0] : ! [X1] : (sf_within(X1,X0) <=> sf_contains(X0,X1))[rectify 187]
187. ! [X162] : ! [X163] : (sf_within(X163,X162) <=> sf_contains(X162,X163))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 2043
Passive clauses: 5987
Generated clauses: 8858
Final active clauses: 2023
Final passive clauses: 3478
Input formulas: 225
Initial clauses: 463

Selected by SInE selection: 225
SInE iterations: 1

Duplicate literals: 27
Fw subsumption resolutions: 453

Simple tautologies: 57
Forward subsumptions: 1610

Binary resolution: 6806
Factoring: 2

Backtracking splits: 1075
Backtracking splits refuted: 19
Backtracking splits refuted at zero level: 14

Memory used: 2302KB
Time elapsed: 0.367 s
------------------------------
% Success in time 0.449 s

========================== MACLEOD SUMMARY ===========================
reasoner: C:/Users/torsten/macleod/vampire_win
status: 
execution finished: Sun Apr 25 16:29:22 2021
total CPU time used: 0
The command was "C:/Users/torsten/macleod/vampire_win --mode casc --proof tptp -t 6000 --input_file C:\Users\torsten\GitHub\colore\ontologies\simple_features\conversions\sfc_fol_nontrivial.all.tptp"
============================ end of footer ===========================
