Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_cont\theorems\conversions\cont_basic_ext_goal1.all
Refutation found. Thanks to Tanya!
114. $false (1:0) [subsumption resolution 111,109]
109. cont(sK0,sK1) (1:3) [backtracking split refutation 46,48,107]
107. $false {2} (2:0) [subsumption resolution 106,54]
54. ~cont(sK0,sK0) {2} (2:3) [resolution 48,29]
29. cont(X2,sK1) | ~cont(X2,sK0) (0:6) [cnf transformation 27]
27. ! [X2] : ((~cont(X2,sK0) | cont(X2,sK1)) & (~cont(X2,sK1) | cont(X2,sK0))) & sK0 != sK1[skolemisation 26]
26. ? [X0,X1] : (! [X2] : ((~cont(X2,X0) | cont(X2,X1)) & (~cont(X2,X1) | cont(X2,X0))) & X0 != X1)[nnf transformation 17]
17. ? [X0,X1] : (! [X2] : (cont(X2,X0) <=> cont(X2,X1)) & X0 != X1)[ennf transformation 7]
7. ~! [X0,X1] : (! [X2] : (cont(X2,X0) <=> cont(X2,X1)) => X0 = X1)[flattening 1]
1. ~! [X0] : ! [X1] : (! [X2] : (cont(X2,X0) <=> cont(X2,X1)) => X0 = X1)[negated conjecture]
106. cont(sK0,sK0) (2:3) [resolution 93,32]
32. zex(X0) | cont(X0,X0) (0:5) [cnf transformation 28]
28. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[nnf transformation 18]
18. ! [X0] : (~zex(X0) <=> cont(X0,X0))[ennf transformation 8]
8. ! [X0] : (~zex(X0) <=> cont(X0,X0))[rectify 2]
2. ! [X2] : (~zex(X2) <=> cont(X2,X2))[input]
93. ~zex(sK0) (1:2) [backtracking split refutation 51,82,92]
92. $false {8} (3:0) [resolution 82,83]
83. cont(sK1,sK1) (2:3) [resolution 80,32]
80. ~zex(sK1) (1:2) [backtracking split refutation 50,52,79]
79. $false {3} (1:0) [subsumption resolution 78,68]
68. zex(sK0) {3} (1:2) [backtracking split refutation 61,63,66]
66. $false {3, 5} (2:0) [subsumption resolution 65,52]
65. cont(sK0,sK0) {5} (2:3) [resolution 63,32]
63. ~zex(sK0) {5} (1:2) [backtracking splitting component]
61. ~zex(sK1) | ~zex(sK0) (1:4) [resolution 40,39]
39. ~sQ1_eqProxy(sK0,sK1) (0:3) [equality proxy replacement 31]
31. sK0 != sK1 (0:3) [cnf transformation 27]
40. sQ1_eqProxy(X0,X1) | ~zex(X1) | ~zex(X0) (0:7) [equality proxy replacement 36]
36. X0 = X1 | ~zex(X1) | ~zex(X0) (0:7) [cnf transformation 21]
21. ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)[flattening 20]
20. ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)[ennf transformation 12]
12. ! [X0,X1] : ((zex(X0) & zex(X1)) => X0 = X1)[flattening 11]
11. ! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => X0 = X1)[rectify 6]
6. ! [X10] : ! [X11] : ((zex(X10) & zex(X11)) => X10 = X11)[input]
78. ~zex(sK0) {3} (1:2) [backtracking split refutation 51,53,77]
77. $false {3, 4} (2:0) [subsumption resolution 76,53]
76. cont(sK1,sK1) {3} (2:3) [resolution 67,32]
67. ~zex(sK1) {3} (1:2) [backtracking split refutation 61,63,66]
53. ~cont(X1,sK1) {4} (1:3) [backtracking splitting component]
52. ~cont(X0,sK0) {3} (1:3) [backtracking splitting component]
50. ~zex(sK1) | ~cont(X0,sK0) (1:5) [resolution 34,29]
34. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 19]
19. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 10]
10. ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))[flattening 9]
9. ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))[rectify 5]
5. ! [X8] : ! [X9] : (zex(X8) => (~cont(X9,X8) & ~cont(X8,X9)))[input]
82. ~cont(X1,sK1) {8} (1:3) [backtracking splitting component]
51. ~zex(sK0) | ~cont(X1,sK1) (1:5) [resolution 34,30]
30. cont(X2,sK0) | ~cont(X2,sK1) (0:6) [cnf transformation 27]
48. ~cont(sK0,sK1) {2} (1:3) [backtracking splitting component]
46. ~zex(sK0) | ~cont(sK0,sK1) (1:5) [resolution 33,30]
33. ~cont(X0,X0) | ~zex(X0) (0:5) [cnf transformation 28]
111. ~cont(sK0,sK1) (1:3) [subsumption resolution 85,96]
96. cont(sK1,sK0) (1:3) [backtracking split refutation 45,47,94]
94. $false {1} (2:0) [subsumption resolution 49,83]
49. ~cont(sK1,sK1) {1} (2:3) [resolution 47,30]
47. ~cont(sK1,sK0) {1} (1:3) [backtracking splitting component]
45. ~zex(sK1) | ~cont(sK1,sK0) (1:5) [resolution 33,29]
85. ~cont(sK1,sK0) | ~cont(sK0,sK1) (1:6) [resolution 41,39]
41. sQ1_eqProxy(X0,X1) | ~cont(X1,X0) | ~cont(X0,X1) (0:9) [equality proxy replacement 37]
37. X0 = X1 | ~cont(X1,X0) | ~cont(X0,X1) (0:9) [cnf transformation 23]
23. ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1)[flattening 22]
22. ! [X0,X1] : ((~cont(X0,X1) | ~cont(X1,X0)) | X0 = X1)[ennf transformation 14]
14. ! [X0,X1] : ((cont(X0,X1) & cont(X1,X0)) => X0 = X1)[flattening 13]
13. ! [X0] : ! [X1] : ((cont(X0,X1) & cont(X1,X0)) => X0 = X1)[rectify 3]
3. ! [X3] : ! [X4] : ((cont(X3,X4) & cont(X4,X3)) => X3 = X4)[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 29
Passive clauses: 49
Generated clauses: 100
Final active clauses: 19
Final passive clauses: 4
Input formulas: 6
Initial clauses: 13

Selected by SInE selection: 6
SInE iterations: 1

Fw subsumption resolutions: 9

Forward subsumptions: 35

Binary resolution: 42

Backtracking splits: 8
Backtracking splits refuted: 7
Backtracking splits refuted at zero level: 4

Memory used: 127KB
Time elapsed: 0.007 s
------------------------------
% Success in time 0.053 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sat Jun 23 11:24:07 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_cont\theorems\conversions\cont_basic_ext_goal1.all.tptp"
Input read from 
============================ end of footer ===========================
