Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 29999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f934,plain,(
  $false),
  inference(subsumption_resolution,[],[f933,f458])).
fof(f458,plain,(
  atomicloopcurve(sK0)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  atomicloopcurve(sK0) & ~icon(sK0)),
  inference(skolemisation,[status(esa)],[f247])).
fof(f247,plain,(
  ? [X0] : (atomicloopcurve(X0) & ~icon(X0))),
  inference(ennf_transformation,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : (atomicloopcurve(X0) => icon(X0))),
  file(unknown,unknown)).
fof(f933,plain,(
  ~atomicloopcurve(sK0)),
  inference(resolution,[],[f928,f530])).
fof(f530,plain,(
  ( ! [X0] : (sP1(X0) | ~atomicloopcurve(X0)) )),
  inference(cnf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X0] : ((~atomicloopcurve(X0) | sP1(X0)) & (~sP1(X0) | atomicloopcurve(X0)))),
  inference(nnf_transformation,[],[f354])).
fof(f354,plain,(
  ! [X0] : (atomicloopcurve(X0) <=> sP1(X0))),
  inference(definition_folding,[],[f353])).
fof(f353,plain,(
  ! [X0] : (sP1(X0) <=> (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (point(X2) & cont(X2,X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | X2 = X3)) & ! [X4] : ~endpoint(X4,X0)))),
  introduced(predicate_definition_introduction,[])).
fof(f928,plain,(
  ~sP1(sK0)),
  inference(resolution,[],[f925,f522])).
fof(f522,plain,(
  ( ! [X0,X1] : (~pp(X1,X0) | ~sP1(X0)) )),
  inference(cnf_transformation,[],[f405])).
fof(f405,plain,(
  ! [X0] : ((~sP1(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & (point(sK18(X0)) & cont(sK18(X0),X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | sK18(X0) = X3)) & ! [X4] : ~endpoint(X4,X0))) & (~curve(X0) | pp(sK19(X0),X0) | ! [X6] : (~point(X6) | ~cont(X6,X0) | (point(sK20(X6,X0)) & cont(sK20(X6,X0),X0) & sK20(X6,X0) != X6)) | endpoint(sK21(X0),X0) | sP1(X0)))),
  inference(skolemisation,[status(esa)],[f404])).
fof(f404,plain,(
  ! [X0] : ((~sP1(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (point(X2) & cont(X2,X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | X2 = X3)) & ! [X4] : ~endpoint(X4,X0))) & (~curve(X0) | ? [X5] : pp(X5,X0) | ! [X6] : (~point(X6) | ~cont(X6,X0) | ? [X7] : (point(X7) & cont(X7,X0) & X6 != X7)) | ? [X8] : endpoint(X8,X0) | sP1(X0)))),
  inference(rectify,[],[f403])).
fof(f403,plain,(
  ! [X0] : ((~sP1(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (point(X2) & cont(X2,X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | X2 = X3)) & ! [X4] : ~endpoint(X4,X0))) & (~curve(X0) | ? [X1] : pp(X1,X0) | ! [X2] : (~point(X2) | ~cont(X2,X0) | ? [X3] : (point(X3) & cont(X3,X0) & X2 != X3)) | ? [X4] : endpoint(X4,X0) | sP1(X0)))),
  inference(flattening,[],[f402])).
fof(f402,plain,(
  ! [X0] : ((~sP1(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2] : (point(X2) & cont(X2,X0) & ! [X3] : (~point(X3) | ~cont(X3,X0) | X2 = X3)) & ! [X4] : ~endpoint(X4,X0))) & ((~curve(X0) | ? [X1] : pp(X1,X0) | ! [X2] : (~point(X2) | ~cont(X2,X0) | ? [X3] : (point(X3) & cont(X3,X0) & X2 != X3)) | ? [X4] : endpoint(X4,X0)) | sP1(X0)))),
  inference(nnf_transformation,[],[f353])).
fof(f925,plain,(
  pp(sK14(sK0),sK0)),
  inference(backtracking_split_refutation,[],[f911,f912_D,f924])).
fof(f924,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f923,f458])).
fof(f923,plain,(
  ~atomicloopcurve(sK0) | $spl6),
  inference(resolution,[],[f917,f530])).
fof(f917,plain,(
  ~sP1(sK0) | $spl6),
  inference(resolution,[],[f913,f522])).
fof(f913,plain,(
  pp(sK12(sK0),sK0) | $spl6),
  inference(resolution,[],[f912,f496])).
fof(f496,plain,(
  ( ! [X0] : (con(X0) | pp(sK12(X0),X0)) )),
  inference(cnf_transformation,[],[f380])).
fof(f380,plain,(
  ! [X0] : ((~con(X0) | ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1)))) & ((pp(sK12(X0),X0) & ~sc(sK12(X0),difference(X0,sK12(X0)))) | con(X0)))),
  inference(skolemisation,[status(esa)],[f379])).
fof(f379,plain,(
  ! [X0] : ((~con(X0) | ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1)))) & (? [X2] : (pp(X2,X0) & ~sc(X2,difference(X0,X2))) | con(X0)))),
  inference(rectify,[],[f378])).
fof(f378,plain,(
  ! [X0] : ((~con(X0) | ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1)))) & (? [X1] : (pp(X1,X0) & ~sc(X1,difference(X0,X1))) | con(X0)))),
  inference(nnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0] : (con(X0) <=> ! [X1] : (~pp(X1,X0) | sc(X1,difference(X0,X1))))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : (con(X0) <=> ! [X1] : (pp(X1,X0) => sc(X1,difference(X0,X1))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X26] : (con(X26) <=> ! [X27] : (pp(X27,X26) => sc(X27,difference(X26,X27))))),
  file(unknown,unknown)).
fof(f912,plain,(
  ~con(sK0) | $spl6),
  inference(cnf_transformation,[],[f912_D])).
fof(f912_D,plain,(
  ~con(sK0) <=> ~$spl6),
  introduced(backtracking_splitting_component,[])).
fof(f911,plain,(
  pp(sK14(sK0),sK0) | ~con(sK0)),
  inference(resolution,[],[f504,f459])).
fof(f459,plain,(
  ~icon(sK0)),
  inference(cnf_transformation,[],[f357])).
fof(f504,plain,(
  ( ! [X0] : (icon(X0) | pp(sK14(X0),X0) | ~con(X0)) )),
  inference(cnf_transformation,[],[f387])).
fof(f387,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & (~con(X0) | (pp(sK14(X0),X0) & ~covers(X0,intersection(sK14(X0),difference(X0,sK14(X0))))) | icon(X0)))),
  inference(skolemisation,[status(esa)],[f386])).
fof(f386,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & (~con(X0) | ? [X2] : (pp(X2,X0) & ~covers(X0,intersection(X2,difference(X0,X2)))) | icon(X0)))),
  inference(rectify,[],[f385])).
fof(f385,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & (~con(X0) | ? [X1] : (pp(X1,X0) & ~covers(X0,intersection(X1,difference(X0,X1)))) | icon(X0)))),
  inference(flattening,[],[f384])).
fof(f384,plain,(
  ! [X0] : ((~icon(X0) | (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1)))))) & ((~con(X0) | ? [X1] : (pp(X1,X0) & ~covers(X0,intersection(X1,difference(X0,X1))))) | icon(X0)))),
  inference(nnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0] : (icon(X0) <=> (con(X0) & ! [X1] : (~pp(X1,X0) | covers(X0,intersection(X1,difference(X0,X1))))))),
  inference(ennf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : (icon(X0) <=> (con(X0) & ! [X1] : (pp(X1,X0) => covers(X0,intersection(X1,difference(X0,X1))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X33] : (icon(X33) <=> (con(X33) & ! [X34] : (pp(X34,X33) => covers(X33,intersection(X34,difference(X33,X34))))))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 200
Passive clauses: 364
Generated clauses: 454
Final active clauses: 196
Final passive clauses: 142
Input formulas: 91
Initial clauses: 207

Selected by SInE selection: 91
SInE iterations: 1

Duplicate literals: 2
Fw subsumption resolutions: 5

Simple tautologies: 11
Forward subsumptions: 78

Binary resolution: 227
Factoring: 1

Backtracking splits: 6
Backtracking splits refuted: 3
Backtracking splits refuted at zero level: 3

Memory used: 383KB
Time elapsed: 0.007 s
------------------------------
% Success in time 0.023 s
============================= vampire_lin32 ================================
execution finished: Fri Oct 11 16:15:29 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_mereotopology_codib/theorems/conversions/atomic_loop_curve_theorems_goal3.all.tptp
============================ end of footer ===========================
