Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 29999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f664,plain,(
  $false),
  inference(subsumption_resolution,[],[f662,f416])).
fof(f416,plain,(
  atomiccurvesegment(sK0)),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  atomiccurvesegment(sK0) & closed(sK0)),
  inference(skolemisation,[status(esa)],[f229])).
fof(f229,plain,(
  ? [X0] : (atomiccurvesegment(X0) & closed(X0))),
  inference(ennf_transformation,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : (atomiccurvesegment(X0) => ~closed(X0))),
  file(unknown,unknown)).
fof(f662,plain,(
  ~atomiccurvesegment(sK0)),
  inference(resolution,[],[f661,f417])).
fof(f417,plain,(
  closed(sK0)),
  inference(cnf_transformation,[],[f326])).
fof(f661,plain,(
  ( ! [X0] : (~closed(X0) | ~atomiccurvesegment(X0)) )),
  inference(resolution,[],[f659,f462])).
fof(f462,plain,(
  ( ! [X0] : (sP0(X0) | ~atomiccurvesegment(X0)) )),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0] : ((~atomiccurvesegment(X0) | sP0(X0)) & (~sP0(X0) | atomiccurvesegment(X0)))),
  inference(nnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0] : (atomiccurvesegment(X0) <=> sP0(X0))),
  inference(definition_folding,[],[f322])).
fof(f322,plain,(
  ! [X0] : (sP0(X0) <=> (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2,X3] : (endpoint(X2,X0) & endpoint(X3,X0) & X2 != X3 & ! [X4] : (~point(X4) | ~cont(X4,X0) | X2 = X4 | X3 = X4))))),
  introduced(predicate_definition_introduction,[])).
fof(f659,plain,(
  ( ! [X0] : (~sP0(X0) | ~closed(X0)) )),
  inference(resolution,[],[f658,f454])).
fof(f454,plain,(
  ( ! [X0] : (endpoint(sK9(X0),X0) | ~sP0(X0)) )),
  inference(cnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0] : ((~sP0(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & (endpoint(sK9(X0),X0) & endpoint(sK10(X0),X0) & sK9(X0) != sK10(X0) & ! [X4] : (~point(X4) | ~cont(X4,X0) | sK9(X0) = X4 | sK10(X0) = X4)))) & (~curve(X0) | pp(sK11(X0),X0) | ! [X6,X7] : (~endpoint(X6,X0) | ~endpoint(X7,X0) | X6 = X7 | (point(sK12(X7,X6,X0)) & cont(sK12(X7,X6,X0),X0) & sK12(X7,X6,X0) != X6 & sK12(X7,X6,X0) != X7)) | sP0(X0)))),
  inference(skolemisation,[status(esa)],[f357])).
fof(f357,plain,(
  ! [X0] : ((~sP0(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2,X3] : (endpoint(X2,X0) & endpoint(X3,X0) & X2 != X3 & ! [X4] : (~point(X4) | ~cont(X4,X0) | X2 = X4 | X3 = X4)))) & (~curve(X0) | ? [X5] : pp(X5,X0) | ! [X6,X7] : (~endpoint(X6,X0) | ~endpoint(X7,X0) | X6 = X7 | ? [X8] : (point(X8) & cont(X8,X0) & X6 != X8 & X7 != X8)) | sP0(X0)))),
  inference(rectify,[],[f356])).
fof(f356,plain,(
  ! [X0] : ((~sP0(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2,X3] : (endpoint(X2,X0) & endpoint(X3,X0) & X2 != X3 & ! [X4] : (~point(X4) | ~cont(X4,X0) | X2 = X4 | X3 = X4)))) & (~curve(X0) | ? [X1] : pp(X1,X0) | ! [X2,X3] : (~endpoint(X2,X0) | ~endpoint(X3,X0) | X2 = X3 | ? [X4] : (point(X4) & cont(X4,X0) & X2 != X4 & X3 != X4)) | sP0(X0)))),
  inference(flattening,[],[f355])).
fof(f355,plain,(
  ! [X0] : ((~sP0(X0) | (curve(X0) & ! [X1] : ~pp(X1,X0) & ? [X2,X3] : (endpoint(X2,X0) & endpoint(X3,X0) & X2 != X3 & ! [X4] : (~point(X4) | ~cont(X4,X0) | X2 = X4 | X3 = X4)))) & ((~curve(X0) | ? [X1] : pp(X1,X0) | ! [X2,X3] : (~endpoint(X2,X0) | ~endpoint(X3,X0) | X2 = X3 | ? [X4] : (point(X4) & cont(X4,X0) & X2 != X4 & X3 != X4))) | sP0(X0)))),
  inference(nnf_transformation,[],[f322])).
fof(f658,plain,(
  ( ! [X4,X5] : (~endpoint(X4,X5) | ~closed(X5)) )),
  inference(resolution,[],[f522,f464])).
fof(f464,plain,(
  ( ! [X0,X1] : (~bcont(X1,X0) | ~closed(X0)) )),
  inference(cnf_transformation,[],[f362])).
fof(f362,plain,(
  ! [X0] : ((~closed(X0) | ! [X1] : ~bcont(X1,X0)) & (bcont(sK13(X0),X0) | closed(X0)))),
  inference(skolemisation,[status(esa)],[f361])).
fof(f361,plain,(
  ! [X0] : ((~closed(X0) | ! [X1] : ~bcont(X1,X0)) & (? [X2] : bcont(X2,X0) | closed(X0)))),
  inference(rectify,[],[f360])).
fof(f360,plain,(
  ! [X0] : ((~closed(X0) | ! [X1] : ~bcont(X1,X0)) & (? [X1] : bcont(X1,X0) | closed(X0)))),
  inference(nnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0] : (closed(X0) <=> ! [X1] : ~bcont(X1,X0))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : (closed(X0) <=> ! [X1] : ~bcont(X1,X0))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X133] : (closed(X133) <=> ! [X134] : ~bcont(X134,X133))),
  file(unknown,unknown)).
fof(f522,plain,(
  ( ! [X0,X1] : (bcont(X0,X1) | ~endpoint(X0,X1)) )),
  inference(cnf_transformation,[],[f377])).
fof(f377,plain,(
  ! [X0,X1] : ((~endpoint(X0,X1) | (point(X0) & curve(X1) & bcont(X0,X1))) & (~point(X0) | ~curve(X1) | ~bcont(X0,X1) | endpoint(X0,X1)))),
  inference(flattening,[],[f376])).
fof(f376,plain,(
  ! [X0,X1] : ((~endpoint(X0,X1) | (point(X0) & curve(X1) & bcont(X0,X1))) & ((~point(X0) | ~curve(X1) | ~bcont(X0,X1)) | endpoint(X0,X1)))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : (endpoint(X0,X1) <=> (point(X0) & curve(X1) & bcont(X0,X1)))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0] : ! [X1] : (endpoint(X0,X1) <=> (point(X0) & curve(X1) & bcont(X0,X1)))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X108] : ! [X109] : (endpoint(X108,X109) <=> (point(X108) & curve(X109) & bcont(X108,X109)))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 83
Passive clauses: 207
Generated clauses: 233
Final active clauses: 83
Final passive clauses: 123
Input formulas: 84
Initial clauses: 180

Selected by SInE selection: 84
SInE iterations: 1

Fw subsumption resolutions: 1

Simple tautologies: 5
Forward subsumptions: 19

Binary resolution: 51
Factoring: 1

Memory used: 255KB
Time elapsed: 0.008 s
------------------------------
% Success in time 0.038 s
============================= vampire_lin32 ================================
execution finished: Mon Oct 07 16:13:15 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_mereotopology_codib/theorems/conversions/atomic_curve_segment_theorems_goal1.all.tptp
============================ end of footer ===========================
