Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_dim\theorems\conversions\dim_prime_linear_bounded_discrete_theorems_goal3.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 501
Passive clauses: 2883
Generated clauses: 5695
Final active clauses: 191
Final passive clauses: 1100
Input formulas: 17
Initial clauses: 37

Selected by SInE selection: 17
SInE iterations: 1

Duplicate literals: 41
Fw subsumption resolutions: 334

Simple tautologies: 100
Forward subsumptions: 2253

Binary resolution: 4839
Factoring: 60

Backtracking splits: 176
Backtracking splits refuted: 66

Memory used: 895KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_dim\theorems\conversions\dim_prime_linear_bounded_discrete_theorems_goal3.all
Refutation found. Thanks to Tanya!
10667. $false (8:0) [subsumption resolution 10666,79]
79. eqdim(sK0,sK1) (0:3) [cnf transformation 58]
58. eqdim(sK0,sK1) & eqdim(sK1,sK2) & ~eqdim(sK0,sK2)[skolemisation 45]
45. ? [X0,X1,X2] : (eqdim(X0,X1) & eqdim(X1,X2) & ~eqdim(X0,X2))[flattening 44]
44. ? [X0,X1,X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) & ~eqdim(X0,X2))[ennf transformation 18]
18. ~! [X0,X1,X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))[negated conjecture]
10666. ~eqdim(sK0,sK1) (8:3) [subsumption resolution 10652,102]
102. leq(X0,X1) | ~eqdim(X0,X1) (0:6) [cnf transformation 76]
76. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 75]
75. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[nnf transformation 39]
39. ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[flattening 38]
38. ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[rectify 7]
7. ! [X11] : ! [X12] : (eqdim(X11,X12) <=> (leq(X11,X12) & leq(X12,X11)))[input]
10652. ~leq(sK0,sK1) | ~eqdim(sK0,sK1) (8:6) [resolution 10640,103]
103. leq(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 76]
10640. ~leq(sK1,sK0) | ~leq(sK0,sK1) (7:6) [resolution 10593,104]
104. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 76]
10593. ~eqdim(sK1,sK0) (6:3) [resolution 10590,80]
80. eqdim(sK1,sK2) (0:3) [cnf transformation 58]
10590. ~eqdim(X13,sK2) | ~eqdim(X13,sK0) (5:6) [subsumption resolution 10578,825]
825. ~lt(X17,X18) | ~eqdim(X19,X18) | ~eqdim(X19,X17) (3:9) [resolution 504,102]
504. ~leq(X20,X21) | ~lt(X21,X22) | ~eqdim(X20,X22) (2:9) [resolution 246,103]
246. ~leq(X2,X0) | ~leq(X0,X1) | ~lt(X1,X2) (1:9) [resolution 105,96]
96. ~leq(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 70]
70. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 69]
69. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[nnf transformation 54]
54. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~leq(X1,X0)))[ennf transformation 35]
35. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~leq(X1,X0)))[flattening 34]
34. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~leq(X1,X0)))[rectify 4]
4. ! [X1] : ! [X2] : (lt(X1,X2) <=> (leq(X1,X2) & ~leq(X2,X1)))[input]
105. leq(X0,X2) | ~leq(X1,X2) | ~leq(X0,X1) (0:9) [cnf transformation 57]
57. ! [X0,X1,X2] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2))[flattening 56]
56. ! [X0,X1,X2] : ((~leq(X0,X1) | ~leq(X1,X2)) | leq(X0,X2))[ennf transformation 41]
41. ! [X0,X1,X2] : ((leq(X0,X1) & leq(X1,X2)) => leq(X0,X2))[flattening 40]
40. ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & leq(X1,X2)) => leq(X0,X2))[rectify 15]
15. ! [X25] : ! [X26] : ! [X27] : ((leq(X25,X26) & leq(X26,X27)) => leq(X25,X27))[input]
10578. ~eqdim(X13,sK0) | ~eqdim(X13,sK2) | lt(sK2,sK0) (5:9) [resolution 4314,160]
160. leq(X1,X0) | lt(X0,X1) (0:6) [subsumption resolution 97,92]
92. leq(X1,X0) | leq(X0,X1) (0:6) [cnf transformation 25]
25. ! [X0,X1] : (leq(X0,X1) | leq(X1,X0))[flattening 24]
24. ! [X0] : ! [X1] : (leq(X0,X1) | leq(X1,X0))[rectify 14]
14. ! [X23] : ! [X24] : (leq(X23,X24) | leq(X24,X23))[input]
97. lt(X0,X1) | leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 70]
4314. ~leq(sK0,sK2) | ~eqdim(X43,sK0) | ~eqdim(X43,sK2) (4:9) [resolution 825,231]
231. lt(sK0,sK2) | ~leq(sK0,sK2) (2:6) [resolution 221,160]
221. ~leq(sK2,sK0) | ~leq(sK0,sK2) (1:6) [resolution 104,81]
81. ~eqdim(sK0,sK2) (0:3) [cnf transformation 58]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 588
Passive clauses: 3502
Generated clauses: 10588
Final active clauses: 588
Final passive clauses: 2649
Input formulas: 17
Initial clauses: 31

Unused predicate definitions: 2
Selected by SInE selection: 17
SInE iterations: 1

Duplicate literals: 472
Fw subsumption resolutions: 704

Simple tautologies: 435
Equational tautologies: 1
Forward subsumptions: 5739

Binary resolution: 7897
Factoring: 50
Forward superposition: 1406
Backward superposition: 42

Unique components: 3

Memory used: 1151KB
Time elapsed: 0.706 s
------------------------------
% Success in time 1.23 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Wed Jan 23 15:09:56 2019
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_dim\theorems\conversions\dim_prime_linear_bounded_discrete_theorems_goal3.all.tptp"
Input read from 
============================ end of footer ===========================
