Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_dim\theorems\conversions\dim_prime_linear_bounded_discrete_theorems_goal6.all
Refutation found. Thanks to Tanya!
861. $false (3:0) [subsumption resolution 860,81]
81. ~lt(sK0,sK1) (0:3) [cnf transformation 57]
57. ~lt(sK0,sK1) & ~lt(sK1,sK0) & ~eqdim(sK0,sK1)[skolemisation 43]
43. ? [X0,X1] : (~lt(X0,X1) & ~lt(X1,X0) & ~eqdim(X0,X1))[ennf transformation 18]
18. ~! [X0,X1] : (lt(X0,X1) | lt(X1,X0) | eqdim(X0,X1))[flattening 1]
1. ~! [X0] : ! [X1] : (lt(X0,X1) | lt(X1,X0) | eqdim(X0,X1))[negated conjecture]
860. lt(sK0,sK1) (3:3) [subsumption resolution 852,82]
82. ~lt(sK1,sK0) (0:3) [cnf transformation 57]
852. lt(sK1,sK0) | lt(sK0,sK1) (3:6) [resolution 357,83]
83. ~eqdim(sK0,sK1) (0:3) [cnf transformation 57]
357. eqdim(X17,X16) | lt(X16,X17) | lt(X17,X16) (2:9) [resolution 166,165]
165. leq(X1,X0) | lt(X0,X1) (0:6) [subsumption resolution 107,94]
94. leq(X1,X0) | leq(X0,X1) (0:6) [cnf transformation 25]
25. ! [X0,X1] : (leq(X0,X1) | leq(X1,X0))[flattening 24]
24. ! [X0] : ! [X1] : (leq(X0,X1) | leq(X1,X0))[rectify 14]
14. ! [X22] : ! [X23] : (leq(X22,X23) | leq(X23,X22))[input]
107. lt(X0,X1) | leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 74]
74. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 73]
73. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[nnf transformation 53]
53. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~leq(X1,X0)))[ennf transformation 36]
36. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~leq(X1,X0)))[flattening 35]
35. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~leq(X1,X0)))[rectify 4]
4. ! [X6] : ! [X7] : (lt(X6,X7) <=> (leq(X6,X7) & ~leq(X7,X6)))[input]
166. ~leq(X0,X1) | lt(X0,X1) | eqdim(X1,X0) (1:9) [resolution 165,104]
104. ~leq(X1,X0) | ~leq(X0,X1) | eqdim(X0,X1) (0:9) [cnf transformation 72]
72. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 71]
71. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[nnf transformation 34]
34. ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[flattening 7]
7. ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 159
Passive clauses: 402
Generated clauses: 793
Final active clauses: 127
Final passive clauses: 172
Input formulas: 17
Initial clauses: 37

Selected by SInE selection: 17
SInE iterations: 1

Duplicate literals: 7
Fw subsumption resolutions: 63

Simple tautologies: 28
Forward subsumptions: 272

Binary resolution: 598
Factoring: 9

Backtracking splits: 39
Backtracking splits refuted: 14

Memory used: 255KB
Time elapsed: 0.050 s
------------------------------
% Success in time 0.102 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Wed Jan 23 15:10:05 2019
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_dim\theorems\conversions\dim_prime_linear_bounded_discrete_theorems_goal6.all.tptp"
Input read from 
============================ end of footer ===========================
