Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_dim\theorems\conversions\dim_prime_linear_bounded_discrete_theorems_goal4.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 508
Passive clauses: 2684
Generated clauses: 6912
Final active clauses: 219
Final passive clauses: 1257
Input formulas: 17
Initial clauses: 37

Selected by SInE selection: 17
SInE iterations: 1

Duplicate literals: 63
Fw subsumption resolutions: 570

Simple tautologies: 116
Forward subsumptions: 3379

Binary resolution: 5439
Factoring: 55

Backtracking splits: 221
Backtracking splits refuted: 105
Backtracking splits refuted at zero level: 2

Memory used: 895KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_dim\theorems\conversions\dim_prime_linear_bounded_discrete_theorems_goal4.all
Refutation found. Thanks to Tanya!
2275. $false (5:0) [subsumption resolution 2258,79]
79. eqdim(sK0,sK1) (0:3) [cnf transformation 60]
60. eqdim(sK0,sK1) & lt(sK2,sK0) & ~lt(sK2,sK1)[skolemisation 46]
46. ? [X0,X1,X2] : (eqdim(X0,X1) & lt(X2,X0) & ~lt(X2,X1))[flattening 45]
45. ? [X0,X1,X2] : ((eqdim(X0,X1) & lt(X2,X0)) & ~lt(X2,X1))[ennf transformation 18]
18. ~! [X0,X1,X2] : ((eqdim(X0,X1) & lt(X2,X0)) => lt(X2,X1))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : ((eqdim(X0,X1) & lt(X2,X0)) => lt(X2,X1))[negated conjecture]
2258. ~eqdim(sK0,sK1) (5:3) [resolution 2239,80]
80. lt(sK2,sK0) (0:3) [cnf transformation 60]
2239. ~lt(sK2,X100) | ~eqdim(X100,sK1) (4:6) [resolution 536,81]
81. ~lt(sK2,sK1) (0:3) [cnf transformation 60]
536. lt(X27,X29) | ~eqdim(X28,X29) | ~lt(X27,X28) (3:9) [resolution 418,156]
156. leq(X1,X0) | lt(X0,X1) (0:6) [subsumption resolution 99,92]
92. leq(X1,X0) | leq(X0,X1) (0:6) [cnf transformation 25]
25. ! [X0,X1] : (leq(X0,X1) | leq(X1,X0))[flattening 24]
24. ! [X0] : ! [X1] : (leq(X0,X1) | leq(X1,X0))[rectify 14]
14. ! [X23] : ! [X24] : (leq(X23,X24) | leq(X24,X23))[input]
99. lt(X0,X1) | leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 72]
72. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 71]
71. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[nnf transformation 56]
56. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~leq(X1,X0)))[ennf transformation 37]
37. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~leq(X1,X0)))[flattening 36]
36. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~leq(X1,X0)))[rectify 4]
4. ! [X7] : ! [X8] : (lt(X7,X8) <=> (leq(X7,X8) & ~leq(X8,X7)))[input]
418. ~leq(X14,X15) | ~lt(X15,X16) | ~eqdim(X16,X14) (2:9) [resolution 168,95]
95. leq(X0,X1) | ~eqdim(X0,X1) (0:6) [cnf transformation 55]
55. ! [X0,X1] : (~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0)))[ennf transformation 44]
44. ! [X0,X1] : (eqdim(X0,X1) => (leq(X0,X1) & leq(X1,X0)))[unused predicate definition removal 35]
35. ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[flattening 34]
34. ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[rectify 7]
7. ! [X11] : ! [X12] : (eqdim(X11,X12) <=> (leq(X11,X12) & leq(X12,X11)))[input]
168. ~leq(X2,X0) | ~leq(X0,X1) | ~lt(X1,X2) (1:9) [resolution 104,98]
98. ~leq(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 72]
104. leq(X0,X2) | ~leq(X1,X2) | ~leq(X0,X1) (0:9) [cnf transformation 59]
59. ! [X0,X1,X2] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2))[flattening 58]
58. ! [X0,X1,X2] : ((~leq(X0,X1) | ~leq(X1,X2)) | leq(X0,X2))[ennf transformation 41]
41. ! [X0,X1,X2] : ((leq(X0,X1) & leq(X1,X2)) => leq(X0,X2))[flattening 40]
40. ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & leq(X1,X2)) => leq(X0,X2))[rectify 15]
15. ! [X25] : ! [X26] : ! [X27] : ((leq(X25,X26) & leq(X26,X27)) => leq(X25,X27))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 258
Passive clauses: 914
Generated clauses: 2196
Final active clauses: 258
Final passive clauses: 590
Input formulas: 17
Initial clauses: 30

Unused predicate definitions: 3
Selected by SInE selection: 17
SInE iterations: 1

Duplicate literals: 79
Fw subsumption resolutions: 168

Simple tautologies: 57
Equational tautologies: 1
Forward subsumptions: 1031

Binary resolution: 1375
Factoring: 12
Forward superposition: 464
Backward superposition: 72

Unique components: 3

Memory used: 383KB
Time elapsed: 0.139 s
------------------------------
% Success in time 0.654 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Wed Jan 23 15:09:59 2019
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_dim\theorems\conversions\dim_prime_linear_bounded_discrete_theorems_goal4.all.tptp"
Input read from 
============================ end of footer ===========================
