Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2333
Passive clauses: 23917
Generated clauses: 39907
Input formulas: 64
Initial clauses: 129

Selected by SInE selection: 64
SInE iterations: 1

Duplicate literals: 93
Fw subsumption resolutions: 2876

Simple tautologies: 83
Forward subsumptions: 11549

Binary resolution: 34399
Factoring: 35

Backtracking splits: 1469
Backtracking splits refuted: 147
Backtracking splits refuted at zero level: 24

Memory used: 7164KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2927
Passive clauses: 36552
Generated clauses: 91190
Input formulas: 64
Initial clauses: 128

Selected by SInE selection: 64
SInE iterations: 1

Duplicate literals: 3378
Fw subsumption resolutions: 6564
Fw demodulations: 38

Simple tautologies: 4084
Equational tautologies: 20
Forward subsumptions: 41749

Binary resolution: 66599
Factoring: 126
Forward superposition: 13783
Backward superposition: 848
Self superposition: 2

Unique components: 7

Memory used: 16375KB
Time elapsed: 1.400 s
------------------------------
% remaining time: 1181 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fsr=off:fde=none:gsp=input_only:lcm=reverse:nwc=10:ptb=off:ssec=off:stl=60:sos=on:sgo=on:sio=off:spl=backtracking:sfv=off:sp=reverse_arity_23 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 8577
Passive clauses: 65715
Generated clauses: 321976
Input formulas: 64
Initial clauses: 126

Duplicate literals: 709
Trivial inequalities: 1067
Bw subsumption resolutions: 49910
Fw demodulations: 97724
Condensations: 460

Simple tautologies: 1148
Equational tautologies: 521
Forward subsumptions: 124376

Binary resolution: 68639
Factoring: 14
Forward superposition: 2689
Backward superposition: 3504

Backtracking splits: 5512
Backtracking splits refuted: 3305
Backtracking splits refuted at zero level: 79

Memory used: 39274KB
Time elapsed: 3.300 s
------------------------------
% remaining time: 1148 next slice time: 95
dis+2_4_bs=off:ep=on:nwc=1.5:nicw=on:ptb=off:ssec=off:sac=on:sio=off:spl=backtracking_70 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f39013,plain,(
  $false),
  inference(subsumption_resolution,[],[f39001,f330])).
fof(f330,plain,(
  ~cont(sK2,sK1)),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  eqdim(sK0,sK1) & cont(sK2,sum(sK0,sK1)) & ~cont(sK2,sK0) & ~cont(sK2,sK1) & ~zex(sK0) & ~zex(sK1) & ~cont(intersection(sK2,sK0),sK0)),
  inference(skolemisation,[status(esa)],[f184])).
fof(f184,plain,(
  ? [X0,X1,X2] : (eqdim(X0,X1) & cont(X2,sum(X0,X1)) & ~cont(X2,X0) & ~cont(X2,X1) & ~zex(X0) & ~zex(X1) & ~cont(intersection(X2,X0),X0))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ? [X0,X1,X2] : ((eqdim(X0,X1) & cont(X2,sum(X0,X1)) & ~cont(X2,X0) & ~cont(X2,X1) & ~zex(X0) & ~zex(X1)) & ~cont(intersection(X2,X0),X0))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ~! [X0,X1,X2] : ((eqdim(X0,X1) & cont(X2,sum(X0,X1)) & ~cont(X2,X0) & ~cont(X2,X1) & ~zex(X0) & ~zex(X1)) => cont(intersection(X2,X0),X0))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ~! [X0] : ! [X1] : ! [X2] : ((eqdim(X0,X1) & cont(X2,sum(X0,X1)) & ~cont(X2,X0) & ~cont(X2,X1) & ~zex(X0) & ~zex(X1)) => cont(intersection(X2,X0),X0))),
  inference(rectify,[],[f7])).
fof(f7,negated_conjecture,(
  ~! [X13] : ! [X14] : ! [X15] : ((eqdim(X13,X14) & cont(X15,sum(X13,X14)) & ~cont(X15,X13) & ~cont(X15,X14) & ~zex(X13) & ~zex(X14)) => cont(intersection(X15,X13),X13))),
  file(unknown,unknown)).
fof(f39001,plain,(
  cont(sK2,sK1)),
  inference(resolution,[],[f37384,f328])).
fof(f328,plain,(
  cont(sK2,sum(sK0,sK1))),
  inference(cnf_transformation,[],[f263])).
fof(f37384,plain,(
  ( ! [X4] : (~cont(sK2,sum(sK0,X4)) | cont(sK2,X4)) )),
  inference(subsumption_resolution,[],[f37278,f329])).
fof(f329,plain,(
  ~cont(sK2,sK0)),
  inference(cnf_transformation,[],[f263])).
fof(f37278,plain,(
  ( ! [X4] : (cont(sK2,X4) | cont(sK2,sK0) | ~cont(sK2,sum(sK0,X4))) )),
  inference(superposition,[],[f446,f37245])).
fof(f37245,plain,(
  difference(sK2,sK0) = sK2),
  inference(subsumption_resolution,[],[f37219,f30669])).
fof(f30669,plain,(
  p(sK2,sK2)),
  inference(resolution,[],[f30663,f3187])).
fof(f3187,plain,(
  ( ! [X2,X3] : (~cont(X2,difference(X2,X3)) | p(X2,X2)) )),
  inference(resolution,[],[f802,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (~p(X2,difference(X0,X1)) | p(X2,X0)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))),
  inference(ennf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : (p(X2,difference(X0,X1)) => p(X2,X0))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0] : ! [X1] : ! [X2] : (p(X2,difference(X0,X1)) => p(X2,X0))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X139] : ! [X140] : ! [X141] : (p(X141,difference(X139,X140)) => p(X141,X139))),
  file(unknown,unknown)).
fof(f802,plain,(
  ( ! [X2,X3] : (p(X2,difference(X2,X3)) | ~cont(X2,difference(X2,X3))) )),
  inference(subsumption_resolution,[],[f799,f352])).
fof(f352,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X154] : ! [X155] : (zex(X154) => (~cont(X155,X154) & ~cont(X154,X155)))),
  file(unknown,unknown)).
fof(f799,plain,(
  ( ! [X2,X3] : (p(X2,difference(X2,X3)) | ~cont(X2,difference(X2,X3)) | zex(difference(X2,X3))) )),
  inference(resolution,[],[f426,f382])).
fof(f382,plain,(
  ( ! [X0,X1] : (eqdim(X0,difference(X0,X1)) | zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) | eqdim(X0,difference(X0,X1)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : (~zex(difference(X0,X1)) => eqdim(X0,difference(X0,X1)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1] : (~zex(difference(X0,X1)) => eqdim(X0,difference(X0,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X26] : ! [X27] : (~zex(difference(X26,X27)) => eqdim(X26,difference(X26,X27)))),
  file(unknown,unknown)).
fof(f426,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X101] : ! [X102] : (p(X101,X102) <=> (cont(X101,X102) & eqdim(X101,X102)))),
  file(unknown,unknown)).
fof(f30663,plain,(
  cont(sK2,difference(sK2,sK0))),
  inference(subsumption_resolution,[],[f30633,f468])).
fof(f468,plain,(
  ~zex(sK2)),
  inference(resolution,[],[f353,f328])).
fof(f353,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f192])).
fof(f30633,plain,(
  zex(sK2) | cont(sK2,difference(sK2,sK0))),
  inference(resolution,[],[f30623,f427])).
fof(f427,plain,(
  ( ! [X0,X1] : (~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X39] : ! [X40] : (zex(difference(X39,X40)) <=> (zex(X39) | cont(X39,X40)))),
  file(unknown,unknown)).
fof(f30623,plain,(
  zex(difference(sK2,difference(sK2,sK0)))),
  inference(resolution,[],[f30567,f384])).
fof(f384,plain,(
  ( ! [X0,X1] : (p(difference(X0,X1),X0) | zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))),
  inference(ennf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0] : ! [X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X129] : ! [X130] : (~zex(difference(X129,X130)) => p(difference(X129,X130),X129))),
  file(unknown,unknown)).
fof(f30567,plain,(
  ( ! [X53] : (~p(difference(X53,difference(sK2,sK0)),sK2)) )),
  inference(resolution,[],[f29590,f26186])).
fof(f26186,plain,(
  ( ! [X24,X23] : (~p(difference(X23,X24),X24)) )),
  inference(resolution,[],[f25589,f990])).
fof(f990,plain,(
  ( ! [X0,X1] : (~po(X1,difference(X0,X1))) )),
  inference(subsumption_resolution,[],[f989,f742])).
fof(f742,plain,(
  ( ! [X6,X5] : (~po(X5,X6) | ~zex(X6)) )),
  inference(resolution,[],[f611,f352])).
fof(f611,plain,(
  ( ! [X4,X5] : (cont(sK22(X5,X4),X5) | ~po(X4,X5)) )),
  inference(resolution,[],[f419,f424])).
fof(f424,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f319])).
fof(f419,plain,(
  ( ! [X0,X1] : (p(sK22(X1,X0),X1) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0,X1] : ((~po(X0,X1) | (p(sK22(X1,X0),X0) & p(sK22(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(skolemisation,[status(esa)],[f314])).
fof(f314,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(rectify,[],[f313])).
fof(f313,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X2] : (~p(X2,X0) | ~p(X2,X1)) | po(X0,X1)))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0] : ! [X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X105] : ! [X106] : (po(X105,X106) <=> ? [X107] : (p(X107,X105) & p(X107,X106)))),
  file(unknown,unknown)).
fof(f989,plain,(
  ( ! [X0,X1] : (~po(X1,difference(X0,X1)) | zex(difference(X0,X1))) )),
  inference(subsumption_resolution,[],[f376,f711])).
fof(f711,plain,(
  ( ! [X2,X3] : (~po(X2,X3) | ~zex(X2)) )),
  inference(resolution,[],[f608,f352])).
fof(f608,plain,(
  ( ! [X4,X5] : (cont(sK22(X5,X4),X4) | ~po(X4,X5)) )),
  inference(resolution,[],[f418,f424])).
fof(f418,plain,(
  ( ! [X0,X1] : (p(sK22(X1,X0),X0) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f315])).
fof(f376,plain,(
  ( ! [X0,X1] : (~po(X1,difference(X0,X1)) | zex(difference(X0,X1)) | zex(X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : (zex(X1) | zex(difference(X0,X1)) | ~po(X1,difference(X0,X1)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : ((zex(X1) | zex(difference(X0,X1))) | ~po(X1,difference(X0,X1)))),
  inference(ennf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : ((~zex(X1) & ~zex(difference(X0,X1))) => ~po(X1,difference(X0,X1)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1] : ((~zex(X1) & ~zex(difference(X0,X1))) => ~po(X1,difference(X0,X1)))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X135] : ! [X136] : ((~zex(X136) & ~zex(difference(X135,X136))) => ~po(X136,difference(X135,X136)))),
  file(unknown,unknown)).
fof(f25589,plain,(
  ( ! [X2,X3] : (po(X2,X3) | ~p(X3,X2)) )),
  inference(resolution,[],[f25555,f19900])).
fof(f19900,plain,(
  ( ! [X10,X9] : (p(X10,difference(X9,sK25)) | ~p(X10,X9)) )),
  inference(subsumption_resolution,[],[f19820,f991])).
fof(f991,plain,(
  ( ! [X0] : (~po(X0,sK25)) )),
  inference(superposition,[],[f990,f545])).
fof(f545,plain,(
  ( ! [X0] : (difference(sK25,X0) = sK25) )),
  inference(resolution,[],[f530,f452])).
fof(f452,plain,(
  zex(sK25)),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  zex(sK25)),
  inference(skolemisation,[status(esa)],[f182])).
fof(f182,plain,(
  ? [X0] : zex(X0)),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ? [X65] : zex(X65)),
  file(unknown,unknown)).
fof(f530,plain,(
  ( ! [X2,X3] : (~zex(X2) | difference(X2,X3) = sK25) )),
  inference(resolution,[],[f524,f428])).
fof(f428,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f321])).
fof(f524,plain,(
  ( ! [X0] : (~zex(X0) | sK25 = X0) )),
  inference(resolution,[],[f508,f452])).
fof(f508,plain,(
  ( ! [X0,X1] : (~zex(X1) | ~zex(X0) | X0 = X1) )),
  inference(resolution,[],[f453,f462])).
fof(f462,plain,(
  ( ! [X0,X1] : (~sQ1_eqProxy(X0,X1) | X0 = X1) )),
  introduced(equality_proxy_axiom,[])).
fof(f453,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X0,X1) | ~zex(X1) | ~zex(X0)) )),
  inference(equality_proxy_replacement,[],[f362])).
fof(f362,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X85] : ! [X86] : ((zex(X85) & zex(X86)) => X85 = X86)),
  file(unknown,unknown)).
fof(f19820,plain,(
  ( ! [X10,X9] : (po(X10,sK25) | p(X10,difference(X9,sK25)) | ~p(X10,X9)) )),
  inference(superposition,[],[f445,f19799])).
fof(f19799,plain,(
  ( ! [X26] : (intersection(X26,sK25) = sK25) )),
  inference(resolution,[],[f19763,f452])).
fof(f19763,plain,(
  ( ! [X43,X42] : (~zex(X42) | intersection(X43,X42) = sK25) )),
  inference(resolution,[],[f5474,f7319])).
fof(f7319,plain,(
  ( ! [X5] : (po(X5,sK14(X5,sK25)) | sK25 = X5) )),
  inference(resolution,[],[f7299,f1487])).
fof(f1487,plain,(
  ( ! [X0,X1] : (po(X1,sK14(X0,X1)) | po(X0,sK14(X0,X1)) | X0 = X1) )),
  inference(resolution,[],[f456,f462])).
fof(f456,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X0,X1) | po(X1,sK14(X1,X0)) | po(X0,sK14(X1,X0))) )),
  inference(equality_proxy_replacement,[],[f385])).
fof(f385,plain,(
  ( ! [X0,X1] : (X0 = X1 | po(X1,sK14(X1,X0)) | po(X0,sK14(X1,X0))) )),
  inference(cnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0,X1] : (((po(X0,sK14(X1,X0)) | po(X1,sK14(X1,X0))) & (~po(X0,sK14(X1,X0)) | ~po(X1,sK14(X1,X0)))) | X0 = X1)),
  inference(skolemisation,[status(esa)],[f287])).
fof(f287,plain,(
  ! [X0,X1] : (? [X2] : ((po(X0,X2) | po(X1,X2)) & (~po(X0,X2) | ~po(X1,X2))) | X0 = X1)),
  inference(nnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : (? [X2] : (po(X0,X2) <~> po(X1,X2)) | X0 = X1)),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : (! [X2] : (po(X0,X2) <=> po(X1,X2)) => X0 = X1)),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0] : ! [X1] : (! [X2] : (po(X0,X2) <=> po(X1,X2)) => X0 = X1)),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X44] : ! [X45] : (! [X46] : (po(X44,X46) <=> po(X45,X46)) => X44 = X45)),
  file(unknown,unknown)).
fof(f7299,plain,(
  ( ! [X1] : (~po(sK25,X1)) )),
  inference(subsumption_resolution,[],[f7297,f991])).
fof(f7297,plain,(
  ( ! [X1] : (po(sK2,sK25) | ~po(sK25,X1)) )),
  inference(duplicate_literal_removal,[],[f7285])).
fof(f7285,plain,(
  ( ! [X1] : (~po(sK25,X1) | po(sK2,sK25) | ~po(sK25,X1)) )),
  inference(resolution,[],[f7157,f738])).
fof(f738,plain,(
  ( ! [X19,X20,X18] : (~p(sK22(X20,X19),X18) | po(X18,X19) | ~po(X19,X20)) )),
  inference(resolution,[],[f420,f418])).
fof(f420,plain,(
  ( ! [X0,X3,X1] : (~p(X3,X1) | po(X0,X1) | ~p(X3,X0)) )),
  inference(cnf_transformation,[],[f315])).
fof(f7157,plain,(
  ( ! [X7] : (p(sK22(X7,sK25),sK2) | ~po(sK25,X7)) )),
  inference(forward_demodulation,[],[f7134,f7081])).
fof(f7081,plain,(
  difference(sK2,sum(sK0,sK1)) = sK25),
  inference(superposition,[],[f2900,f486])).
fof(f486,plain,(
  ( ! [X1] : (sum(sK25,X1) = X1) )),
  inference(superposition,[],[f484,f350])).
fof(f350,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = sum(X1,X0)) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : sum(X0,X1) = sum(X1,X0)),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X16] : ! [X17] : sum(X16,X17) = sum(X17,X16)),
  file(unknown,unknown)).
fof(f484,plain,(
  ( ! [X0] : (sum(X0,sK25) = X0) )),
  inference(resolution,[],[f351,f452])).
fof(f351,plain,(
  ( ! [X0,X1] : (~zex(X1) | sum(X0,X1) = X0) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : (~zex(X1) | sum(X0,X1) = X0)),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : (zex(X1) => sum(X0,X1) = X0)),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1] : (zex(X1) => sum(X0,X1) = X0)),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X9] : ! [X10] : (zex(X10) => sum(X9,X10) = X9)),
  file(unknown,unknown)).
fof(f2900,plain,(
  ( ! [X24] : (sum(X24,difference(sK2,sum(sK0,sK1))) = X24) )),
  inference(resolution,[],[f507,f328])).
fof(f507,plain,(
  ( ! [X2,X0,X1] : (~cont(X0,X1) | sum(X2,difference(X0,X1)) = X2) )),
  inference(resolution,[],[f429,f351])).
fof(f429,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f321])).
fof(f7134,plain,(
  ( ! [X7] : (p(sK22(X7,sK25),sK2) | ~po(difference(sK2,sum(sK0,sK1)),X7)) )),
  inference(superposition,[],[f614,f7081])).
fof(f614,plain,(
  ( ! [X6,X8,X7] : (p(sK22(X6,difference(X7,X8)),X7) | ~po(difference(X7,X8),X6)) )),
  inference(resolution,[],[f435,f418])).
fof(f5474,plain,(
  ( ! [X2,X0,X1] : (~po(intersection(X0,X1),X2) | ~zex(X1)) )),
  inference(resolution,[],[f623,f352])).
fof(f623,plain,(
  ( ! [X6,X8,X7] : (cont(sK22(X6,intersection(X7,X8)),X8) | ~po(intersection(X7,X8),X6)) )),
  inference(resolution,[],[f449,f418])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (~p(X2,intersection(X0,X1)) | cont(X2,X1)) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(flattening,[],[f323])).
fof(f323,plain,(
  ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(nnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X54] : ! [X55] : ! [X56] : ((cont(X56,X54) & cont(X56,X55) & eqdim(X56,intersection(X54,X55))) <=> p(X56,intersection(X54,X55)))),
  file(unknown,unknown)).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (po(X2,intersection(X0,X1)) | p(X2,difference(X0,X1)) | ~p(X2,X0)) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : (~p(X2,X0) | po(X2,intersection(X0,X1)) | p(X2,difference(X0,X1)))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : ((~p(X2,X0) | po(X2,intersection(X0,X1))) | p(X2,difference(X0,X1)))),
  inference(ennf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((p(X2,X0) & ~po(X2,intersection(X0,X1))) => p(X2,difference(X0,X1)))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((p(X2,X0) & ~po(X2,intersection(X0,X1))) => p(X2,difference(X0,X1)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X145] : ! [X146] : ! [X147] : ((p(X147,X145) & ~po(X147,intersection(X145,X146))) => p(X147,difference(X145,X146)))),
  file(unknown,unknown)).
fof(f25555,plain,(
  ( ! [X23,X21,X22] : (~p(X22,difference(X21,X23)) | po(X21,X22)) )),
  inference(subsumption_resolution,[],[f25531,f24898])).
fof(f24898,plain,(
  ( ! [X12,X11] : (~p(X11,X12) | ~zex(X11)) )),
  inference(resolution,[],[f24207,f353])).
fof(f24207,plain,(
  ( ! [X8,X7] : (cont(X7,difference(X8,sK25)) | ~p(X7,X8)) )),
  inference(resolution,[],[f19900,f424])).
fof(f25531,plain,(
  ( ! [X23,X21,X22] : (po(X21,X22) | zex(X22) | ~p(X22,difference(X21,X23))) )),
  inference(resolution,[],[f4914,f930])).
fof(f930,plain,(
  ( ! [X0,X1] : (po(X1,X0) | zex(X0) | ~p(X0,X1)) )),
  inference(resolution,[],[f929,f420])).
fof(f929,plain,(
  ( ! [X0] : (p(X0,X0) | zex(X0)) )),
  inference(duplicate_literal_removal,[],[f928])).
fof(f928,plain,(
  ( ! [X0] : (p(X0,X0) | zex(X0) | p(X0,X0) | zex(X0)) )),
  inference(resolution,[],[f380,f379])).
fof(f379,plain,(
  ( ! [X0,X1] : (po(X0,sK13(X1,X0)) | p(X0,X1) | zex(X0)) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : (zex(X0) | (po(X0,sK13(X1,X0)) & ~po(X1,sK13(X1,X0))) | p(X0,X1))),
  inference(skolemisation,[status(esa)],[f225])).
fof(f225,plain,(
  ! [X0,X1] : (zex(X0) | ? [X2] : (po(X0,X2) & ~po(X1,X2)) | p(X0,X1))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((zex(X0) | ? [X2] : (po(X0,X2) & ~po(X1,X2))) | p(X0,X1))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ((~zex(X0) & ! [X2] : (po(X0,X2) => po(X1,X2))) => p(X0,X1))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1] : ((~zex(X0) & ! [X2] : (po(X0,X2) => po(X1,X2))) => p(X0,X1))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X41] : ! [X42] : ((~zex(X41) & ! [X43] : (po(X41,X43) => po(X42,X43))) => p(X41,X42))),
  file(unknown,unknown)).
fof(f380,plain,(
  ( ! [X0,X1] : (~po(X1,sK13(X1,X0)) | p(X0,X1) | zex(X0)) )),
  inference(cnf_transformation,[],[f286])).
fof(f4914,plain,(
  ( ! [X2,X0,X1] : (~po(difference(X0,X1),X2) | po(X0,X2)) )),
  inference(duplicate_literal_removal,[],[f4889])).
fof(f4889,plain,(
  ( ! [X2,X0,X1] : (~po(difference(X0,X1),X2) | po(X0,X2) | ~po(difference(X0,X1),X2)) )),
  inference(resolution,[],[f614,f739])).
fof(f739,plain,(
  ( ! [X23,X21,X22] : (~p(sK22(X22,X23),X21) | po(X21,X22) | ~po(X23,X22)) )),
  inference(resolution,[],[f420,f419])).
fof(f29590,plain,(
  ( ! [X3] : (p(X3,difference(sK2,sK0)) | ~p(X3,sK2)) )),
  inference(subsumption_resolution,[],[f29504,f991])).
fof(f29504,plain,(
  ( ! [X3] : (po(X3,sK25) | p(X3,difference(sK2,sK0)) | ~p(X3,sK2)) )),
  inference(superposition,[],[f445,f29472])).
fof(f29472,plain,(
  intersection(sK2,sK0) = sK25),
  inference(resolution,[],[f29406,f524])).
fof(f29406,plain,(
  zex(intersection(sK2,sK0))),
  inference(resolution,[],[f934,f333])).
fof(f333,plain,(
  ~cont(intersection(sK2,sK0),sK0)),
  inference(cnf_transformation,[],[f263])).
fof(f934,plain,(
  ( ! [X6,X5] : (cont(intersection(X5,X6),X6) | zex(intersection(X5,X6))) )),
  inference(resolution,[],[f929,f449])).
fof(f37219,plain,(
  difference(sK2,sK0) = sK2 | ~p(sK2,sK2)),
  inference(resolution,[],[f2755,f30530])).
fof(f30530,plain,(
  ( ! [X6] : (cont(X6,difference(sK2,sK0)) | ~p(X6,sK2)) )),
  inference(resolution,[],[f29590,f424])).
fof(f2755,plain,(
  ( ! [X14,X15] : (~cont(X14,difference(X14,X15)) | difference(X14,X15) = X14) )),
  inference(subsumption_resolution,[],[f2739,f352])).
fof(f2739,plain,(
  ( ! [X14,X15] : (~cont(X14,difference(X14,X15)) | difference(X14,X15) = X14 | zex(difference(X14,X15))) )),
  inference(resolution,[],[f896,f668])).
fof(f668,plain,(
  ( ! [X4,X5] : (cont(difference(X4,X5),X4) | zex(difference(X4,X5))) )),
  inference(resolution,[],[f384,f424])).
fof(f896,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~cont(X0,X1) | X0 = X1) )),
  inference(resolution,[],[f454,f462])).
fof(f454,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X0,X1) | ~cont(X1,X0) | ~cont(X0,X1)) )),
  inference(equality_proxy_replacement,[],[f367])).
fof(f367,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~cont(X1,X0) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1)),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((~cont(X0,X1) | ~cont(X1,X0)) | X0 = X1)),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((cont(X0,X1) & cont(X1,X0)) => X0 = X1)),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1] : ((cont(X0,X1) & cont(X1,X0)) => X0 = X1)),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X149] : ! [X150] : ((cont(X149,X150) & cont(X150,X149)) => X149 = X150)),
  file(unknown,unknown)).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2] : (~cont(X2,sum(X0,X1)) | cont(X2,X0) | cont(difference(X2,X0),X1))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : ((~cont(X2,sum(X0,X1)) | cont(X2,X0)) | cont(difference(X2,X0),X1))),
  inference(ennf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X23] : ! [X24] : ! [X25] : ((cont(X25,sum(X23,X24)) & ~cont(X25,X23)) => cont(difference(X25,X23),X24))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1820
Passive clauses: 13899
Generated clauses: 39244
Final active clauses: 1288
Final passive clauses: 10132
Input formulas: 64
Initial clauses: 128

Duplicate literals: 520
Fw subsumption resolutions: 3040
Fw demodulations: 2997
Bw demodulations: 337

Simple tautologies: 1482
Equational tautologies: 51
Forward subsumptions: 18061

Binary resolution: 26104
Factoring: 9
Forward superposition: 2036
Backward superposition: 3462
Self superposition: 2

Backtracking splits: 60
Backtracking splits refuted: 29
Backtracking splits refuted at zero level: 15

Memory used: 6268KB
Time elapsed: 0.432 s
------------------------------
% Success in time 5.635 s
