Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f13372,plain,(
  $false),
  inference(subsumption_resolution,[],[f13370,f12957])).
fof(f12957,plain,(
  cont(sK1,sK1)),
  inference(resolution,[],[f12925,f287])).
fof(f287,plain,(
  ( ! [X0] : (zex(X0) | cont(X0,X0)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X119] : (~zex(X119) <=> cont(X119,X119))),
  file(unknown,unknown)).
fof(f12925,plain,(
  ~zex(sK1)),
  inference(backtracking_split_refutation,[],[f413,f417_D,f12924])).
fof(f12924,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f12857,f11703])).
fof(f11703,plain,(
  ~zex(sK1) | $spl1),
  inference(resolution,[],[f11656,f290])).
fof(f290,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X125] : ! [X126] : (zex(X125) => (~cont(X126,X125) & ~cont(X125,X126)))),
  file(unknown,unknown)).
fof(f11656,plain,(
  cont(sK0,sK1) | $spl1),
  inference(backtracking_split_refutation,[],[f733,f737_D,f11648])).
fof(f11648,plain,(
  $false | ($spl1 | $spl13)),
  inference(resolution,[],[f11384,f743])).
fof(f743,plain,(
  p(difference(sK0,sK1),sK0) | ($spl1 | $spl13)),
  inference(resolution,[],[f739,f316])).
fof(f316,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : ! [X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X52] : ! [X53] : (~zex(difference(X52,X53)) => p(difference(X52,X53),X52))),
  file(unknown,unknown)).
fof(f739,plain,(
  ~zex(difference(sK0,sK1)) | ($spl1 | $spl13)),
  inference(resolution,[],[f737,f654])).
fof(f654,plain,(
  ( ! [X0] : (cont(sK0,X0) | ~zex(difference(sK0,X0))) ) | $spl1),
  inference(resolution,[],[f357,f417])).
fof(f357,plain,(
  ( ! [X0,X1] : (zex(X0) | cont(X0,X1) | ~zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X31] : ! [X32] : (zex(difference(X31,X32)) <=> (zex(X31) | cont(X31,X32)))),
  file(unknown,unknown)).
fof(f11384,plain,(
  ( ! [X9] : (~p(difference(sK0,sK1),X9)) ) | ($spl1 | $spl13)),
  inference(subsumption_resolution,[],[f11380,f743])).
fof(f11380,plain,(
  ( ! [X9] : (~p(difference(sK0,sK1),sK0) | ~p(difference(sK0,sK1),X9)) )),
  inference(resolution,[],[f11332,f5437])).
fof(f5437,plain,(
  ( ! [X33,X32] : (po(X32,intersection(X33,sK1)) | ~p(X32,sK0) | ~p(X32,X33)) )),
  inference(resolution,[],[f1378,f895])).
fof(f895,plain,(
  ( ! [X0] : (~po(sK0,difference(X0,sK1))) )),
  inference(resolution,[],[f879,f272])).
fof(f272,plain,(
  ( ! [X2] : (po(sK1,X2) | ~po(sK0,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X2] : ((~po(sK0,X2) | po(sK1,X2)) & (~po(sK1,X2) | po(sK0,X2))) & sK0 != sK1),
  inference(skolemisation,[status(esa)],[f211])).
fof(f211,plain,(
  ? [X0,X1] : (! [X2] : ((~po(X0,X2) | po(X1,X2)) & (~po(X1,X2) | po(X0,X2))) & X0 != X1)),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ? [X0,X1] : (! [X2] : (po(X0,X2) <=> po(X1,X2)) & X0 != X1)),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ~! [X0,X1] : (! [X2] : (po(X0,X2) <=> po(X1,X2)) => X0 = X1)),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : (! [X2] : (po(X0,X2) <=> po(X1,X2)) => X0 = X1)),
  file(unknown,unknown)).
fof(f879,plain,(
  ( ! [X0,X1] : (~po(X1,difference(X0,X1))) )),
  inference(subsumption_resolution,[],[f878,f561])).
fof(f561,plain,(
  ( ! [X2,X3] : (~po(X2,X3) | ~zex(X3)) )),
  inference(resolution,[],[f505,f290])).
fof(f505,plain,(
  ( ! [X0,X1] : (cont(sK18(X1,X0),X1) | ~po(X0,X1)) )),
  inference(resolution,[],[f346,f354])).
fof(f354,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X74] : ! [X75] : (p(X74,X75) <=> (cont(X74,X75) & eqdim(X74,X75)))),
  file(unknown,unknown)).
fof(f346,plain,(
  ( ! [X0,X1] : (p(sK18(X1,X0),X1) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1] : ((~po(X0,X1) | (p(sK18(X1,X0),X0) & p(sK18(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(skolemisation,[status(esa)],[f257])).
fof(f257,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(rectify,[],[f256])).
fof(f256,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X2] : (~p(X2,X0) | ~p(X2,X1)) | po(X0,X1)))),
  inference(nnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X98] : ! [X99] : (po(X98,X99) <=> ? [X100] : (p(X100,X98) & p(X100,X99)))),
  file(unknown,unknown)).
fof(f878,plain,(
  ( ! [X0,X1] : (~po(X1,difference(X0,X1)) | zex(difference(X0,X1))) )),
  inference(subsumption_resolution,[],[f310,f555])).
fof(f555,plain,(
  ( ! [X2,X3] : (~po(X2,X3) | ~zex(X2)) )),
  inference(resolution,[],[f497,f290])).
fof(f497,plain,(
  ( ! [X0,X1] : (cont(sK18(X1,X0),X0) | ~po(X0,X1)) )),
  inference(resolution,[],[f345,f354])).
fof(f345,plain,(
  ( ! [X0,X1] : (p(sK18(X1,X0),X0) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f258])).
fof(f310,plain,(
  ( ! [X0,X1] : (~po(X1,difference(X0,X1)) | zex(difference(X0,X1)) | zex(X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : (zex(X1) | zex(difference(X0,X1)) | ~po(X1,difference(X0,X1)))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : ((zex(X1) | zex(difference(X0,X1))) | ~po(X1,difference(X0,X1)))),
  inference(ennf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : ((~zex(X1) & ~zex(difference(X0,X1))) => ~po(X1,difference(X0,X1)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1] : ((~zex(X1) & ~zex(difference(X0,X1))) => ~po(X1,difference(X0,X1)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X58] : ! [X59] : ((~zex(X59) & ~zex(difference(X58,X59))) => ~po(X59,difference(X58,X59)))),
  file(unknown,unknown)).
fof(f1378,plain,(
  ( ! [X14,X15,X13,X16] : (po(X16,difference(X14,X15)) | ~p(X13,X14) | ~p(X13,X16) | po(X13,intersection(X14,X15))) )),
  inference(resolution,[],[f370,f347])).
fof(f347,plain,(
  ( ! [X0,X3,X1] : (~p(X3,X1) | ~p(X3,X0) | po(X0,X1)) )),
  inference(cnf_transformation,[],[f258])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (p(X2,difference(X0,X1)) | po(X2,intersection(X0,X1)) | ~p(X2,X0)) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : (~p(X2,X0) | po(X2,intersection(X0,X1)) | p(X2,difference(X0,X1)))),
  inference(flattening,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : ((~p(X2,X0) | po(X2,intersection(X0,X1))) | p(X2,difference(X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((p(X2,X0) & ~po(X2,intersection(X0,X1))) => p(X2,difference(X0,X1)))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((p(X2,X0) & ~po(X2,intersection(X0,X1))) => p(X2,difference(X0,X1)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X68] : ! [X69] : ! [X70] : ((p(X70,X68) & ~po(X70,intersection(X68,X69))) => p(X70,difference(X68,X69)))),
  file(unknown,unknown)).
fof(f11332,plain,(
  ( ! [X46] : (~po(difference(sK0,sK1),X46)) )),
  inference(resolution,[],[f3668,f895])).
fof(f3668,plain,(
  ( ! [X10,X11,X9] : (po(X9,difference(X9,X10)) | ~po(difference(X9,X10),X11)) )),
  inference(duplicate_literal_removal,[],[f3647])).
fof(f3647,plain,(
  ( ! [X10,X11,X9] : (~po(difference(X9,X10),X11) | po(X9,difference(X9,X10)) | ~po(difference(X9,X10),X11)) )),
  inference(resolution,[],[f521,f591])).
fof(f591,plain,(
  ( ! [X10,X11,X9] : (~p(sK18(X9,X10),X11) | po(X11,X10) | ~po(X10,X9)) )),
  inference(resolution,[],[f347,f345])).
fof(f521,plain,(
  ( ! [X6,X8,X7] : (p(sK18(X6,difference(X7,X8)),X7) | ~po(difference(X7,X8),X6)) )),
  inference(resolution,[],[f361,f345])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~p(X2,difference(X0,X1)) | p(X2,X0)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : (p(X2,difference(X0,X1)) => p(X2,X0))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1] : ! [X2] : (p(X2,difference(X0,X1)) => p(X2,X0))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X62] : ! [X63] : ! [X64] : (p(X64,difference(X62,X63)) => p(X64,X62))),
  file(unknown,unknown)).
fof(f737,plain,(
  ~cont(sK0,sK1) | $spl13),
  inference(cnf_transformation,[],[f737_D])).
fof(f737_D,plain,(
  ~cont(sK0,sK1) <=> ~$spl13),
  introduced(backtracking_splitting_component,[])).
fof(f733,plain,(
  ~cont(sK1,sK0) | ~cont(sK0,sK1)),
  inference(resolution,[],[f380,f377])).
fof(f377,plain,(
  ~sQ1_eqProxy(sK0,sK1)),
  inference(equality_proxy_replacement,[],[f274])).
fof(f274,plain,(
  sK0 != sK1),
  inference(cnf_transformation,[],[f212])).
fof(f380,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X0,X1) | ~cont(X1,X0) | ~cont(X0,X1)) )),
  inference(equality_proxy_replacement,[],[f301])).
fof(f301,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~cont(X1,X0) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1)),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : ((~cont(X0,X1) | ~cont(X1,X0)) | X0 = X1)),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : ((cont(X0,X1) & cont(X1,X0)) => X0 = X1)),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1] : ((cont(X0,X1) & cont(X1,X0)) => X0 = X1)),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X120] : ! [X121] : ((cont(X120,X121) & cont(X121,X120)) => X120 = X121)),
  file(unknown,unknown)).
fof(f12857,plain,(
  zex(sK1) | $spl1),
  inference(backtracking_split_refutation,[],[f416,f418_D,f12855])).
fof(f12855,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f12842,f11656])).
fof(f12842,plain,(
  ~cont(sK0,sK1) | $spl2),
  inference(backtracking_split_refutation,[],[f736,f738_D,f12834])).
fof(f12834,plain,(
  $false | ($spl2 | $spl14)),
  inference(resolution,[],[f11429,f751])).
fof(f751,plain,(
  p(difference(sK1,sK0),sK1) | ($spl2 | $spl14)),
  inference(resolution,[],[f741,f316])).
fof(f741,plain,(
  ~zex(difference(sK1,sK0)) | ($spl2 | $spl14)),
  inference(resolution,[],[f738,f655])).
fof(f655,plain,(
  ( ! [X1] : (cont(sK1,X1) | ~zex(difference(sK1,X1))) ) | $spl2),
  inference(resolution,[],[f357,f418])).
fof(f11429,plain,(
  ( ! [X11] : (~p(difference(sK1,sK0),X11)) ) | ($spl2 | $spl14)),
  inference(subsumption_resolution,[],[f11425,f751])).
fof(f11425,plain,(
  ( ! [X11] : (~p(difference(sK1,sK0),sK1) | ~p(difference(sK1,sK0),X11)) )),
  inference(resolution,[],[f11341,f5440])).
fof(f5440,plain,(
  ( ! [X41,X40] : (po(X40,intersection(X41,sK0)) | ~p(X40,sK1) | ~p(X40,X41)) )),
  inference(resolution,[],[f1378,f896])).
fof(f896,plain,(
  ( ! [X1] : (~po(sK1,difference(X1,sK0))) )),
  inference(resolution,[],[f879,f273])).
fof(f273,plain,(
  ( ! [X2] : (po(sK0,X2) | ~po(sK1,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f11341,plain,(
  ( ! [X61] : (~po(difference(sK1,sK0),X61)) )),
  inference(resolution,[],[f3668,f896])).
fof(f738,plain,(
  ~cont(sK1,sK0) | $spl14),
  inference(cnf_transformation,[],[f738_D])).
fof(f738_D,plain,(
  ~cont(sK1,sK0) <=> ~$spl14),
  introduced(backtracking_splitting_component,[])).
fof(f736,plain,(
  ~cont(sK0,sK1) | ~cont(sK1,sK0)),
  inference(resolution,[],[f380,f397])).
fof(f397,plain,(
  ~sQ1_eqProxy(sK1,sK0)),
  inference(resolution,[],[f386,f377])).
fof(f386,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X1,X0) | ~sQ1_eqProxy(X0,X1)) )),
  introduced(equality_proxy_axiom,[])).
fof(f418,plain,(
  ~zex(sK1) | $spl2),
  inference(cnf_transformation,[],[f418_D])).
fof(f418_D,plain,(
  ~zex(sK1) <=> ~$spl2),
  introduced(backtracking_splitting_component,[])).
fof(f416,plain,(
  ~zex(sK0) | ~zex(sK1)),
  inference(resolution,[],[f379,f397])).
fof(f379,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X0,X1) | ~zex(X1) | ~zex(X0)) )),
  inference(equality_proxy_replacement,[],[f298])).
fof(f298,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X82] : ! [X83] : ((zex(X82) & zex(X83)) => X82 = X83)),
  file(unknown,unknown)).
fof(f417,plain,(
  ~zex(sK0) | $spl1),
  inference(cnf_transformation,[],[f417_D])).
fof(f417_D,plain,(
  ~zex(sK0) <=> ~$spl1),
  introduced(backtracking_splitting_component,[])).
fof(f413,plain,(
  ~zex(sK1) | ~zex(sK0)),
  inference(resolution,[],[f379,f377])).
fof(f13370,plain,(
  ~cont(sK1,sK1)),
  inference(resolution,[],[f13366,f4163])).
fof(f4163,plain,(
  ( ! [X0] : (p(X0,X0) | ~cont(X0,X0)) )),
  inference(subsumption_resolution,[],[f4097,f289])).
fof(f289,plain,(
  ( ! [X0] : (~less(X0,X0)) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X76] : ~less(X76,X76)),
  file(unknown,unknown)).
fof(f4097,plain,(
  ( ! [X0] : (~cont(X0,X0) | p(X0,X0) | less(X0,X0)) )),
  inference(resolution,[],[f627,f289])).
fof(f627,plain,(
  ( ! [X2,X3] : (less(X3,X2) | ~cont(X2,X3) | p(X2,X3) | less(X2,X3)) )),
  inference(resolution,[],[f356,f341])).
fof(f341,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | less(X1,X0) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X87] : ! [X88] : (eqdim(X87,X88) <=> (~less(X87,X88) & ~less(X88,X87)))),
  file(unknown,unknown)).
fof(f356,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f13366,plain,(
  ~p(sK1,sK1)),
  inference(backtracking_split_refutation,[],[f13253,f13258_D,f13365])).
fof(f13365,plain,(
  $false | $spl398),
  inference(subsumption_resolution,[],[f13364,f12957])).
fof(f13364,plain,(
  ~cont(sK1,sK1) | $spl398),
  inference(resolution,[],[f13258,f4163])).
fof(f13258,plain,(
  ( ! [X9] : (~p(sK1,X9)) ) | $spl398),
  inference(cnf_transformation,[],[f13258_D])).
fof(f13258_D,plain,(
  ( ! [X9] : (~p(sK1,X9)) ) <=> ~$spl398),
  introduced(backtracking_splitting_component,[])).
fof(f13253,plain,(
  ( ! [X9] : (~p(sK1,sK1) | ~p(sK1,X9)) )),
  inference(resolution,[],[f13164,f5440])).
fof(f13164,plain,(
  ( ! [X0] : (~po(sK1,X0)) )),
  inference(resolution,[],[f13057,f273])).
fof(f13057,plain,(
  ( ! [X7] : (~po(sK0,X7)) )),
  inference(resolution,[],[f13037,f345])).
fof(f13037,plain,(
  ( ! [X0] : (~p(X0,sK0)) )),
  inference(backtracking_split_refutation,[],[f11183,f12941_D,f13036])).
fof(f13036,plain,(
  $false | $spl376),
  inference(subsumption_resolution,[],[f12941,f12926])).
fof(f12926,plain,(
  zex(sK0)),
  inference(backtracking_split_refutation,[],[f413,f417_D,f12924])).
fof(f12941,plain,(
  ~zex(sK0) | $spl376),
  inference(cnf_transformation,[],[f12941_D])).
fof(f12941_D,plain,(
  ~zex(sK0) <=> ~$spl376),
  introduced(backtracking_splitting_component,[])).
fof(f11183,plain,(
  ( ! [X0] : (~p(X0,sK0) | ~zex(sK0)) )),
  inference(factoring,[],[f7971])).
fof(f7971,plain,(
  ( ! [X0,X1] : (~p(X1,sK0) | ~p(X1,X0) | ~zex(X0)) )),
  inference(resolution,[],[f3885,f5437])).
fof(f3885,plain,(
  ( ! [X4,X5,X3] : (~po(X3,intersection(X4,X5)) | ~zex(X4)) )),
  inference(resolution,[],[f526,f290])).
fof(f526,plain,(
  ( ! [X10,X11,X9] : (cont(sK18(intersection(X9,X10),X11),X9) | ~po(X11,intersection(X9,X10))) )),
  inference(resolution,[],[f372,f346])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (~p(X2,intersection(X0,X1)) | cont(X2,X0)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X47] : ! [X48] : ! [X49] : ((cont(X49,X47) & cont(X49,X48) & eqdim(X49,intersection(X47,X48))) <=> p(X49,intersection(X47,X48)))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1275
Passive clauses: 7483
Generated clauses: 13356
Final active clauses: 438
Final passive clauses: 1932
Input formulas: 52
Initial clauses: 108

Selected by SInE selection: 52
SInE iterations: 1

Duplicate literals: 79
Fw subsumption resolutions: 428

Simple tautologies: 60
Forward subsumptions: 5158

Binary resolution: 11883
Factoring: 46

Backtracking splits: 399
Backtracking splits refuted: 86
Backtracking splits refuted at zero level: 16

Memory used: 2686KB
Time elapsed: 0.164 s
------------------------------
% Success in time 0.173 s
