Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1200 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2380
Passive clauses: 21657
Generated clauses: 38947
Input formulas: 38
Initial clauses: 90

Selected by SInE selection: 38
SInE iterations: 1

Duplicate literals: 181
Fw subsumption resolutions: 1742

Simple tautologies: 161
Forward subsumptions: 14194

Binary resolution: 35263
Factoring: 98

Backtracking splits: 1280
Backtracking splits refuted: 227
Backtracking splits refuted at zero level: 92

Memory used: 6780KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 3017
Passive clauses: 35587
Generated clauses: 98030
Input formulas: 38
Initial clauses: 89

Selected by SInE selection: 38
SInE iterations: 1

Duplicate literals: 6593
Fw subsumption resolutions: 11102
Fw demodulations: 229
Bw demodulations: 3

Simple tautologies: 4915
Equational tautologies: 29
Forward subsumptions: 40815

Binary resolution: 69974
Factoring: 213
Forward superposition: 9560
Backward superposition: 700

Unique components: 4

Memory used: 14839KB
Time elapsed: 1.400 s
------------------------------
% remaining time: 1181 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fsr=off:fde=none:gsp=input_only:lcm=reverse:nwc=10:ptb=off:ssec=off:stl=60:sos=on:sgo=on:sio=off:spl=backtracking:sfv=off:sp=reverse_arity_23 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 9460
Passive clauses: 47254
Generated clauses: 231690
Input formulas: 38
Initial clauses: 87

Duplicate literals: 802
Trivial inequalities: 9
Bw subsumption resolutions: 31812
Fw demodulations: 39747
Condensations: 636

Simple tautologies: 957
Equational tautologies: 710
Forward subsumptions: 110716

Binary resolution: 100968
Factoring: 27
Forward superposition: 6973
Backward superposition: 9979
Self superposition: 3
Equality factoring: 2

Backtracking splits: 3667
Backtracking splits refuted: 1836
Backtracking splits refuted at zero level: 59

Memory used: 46054KB
Time elapsed: 3.300 s
------------------------------
% remaining time: 1148 next slice time: 95
dis+2_4_bs=off:ep=on:nwc=1.5:nicw=on:ptb=off:ssec=off:sac=on:sio=off:spl=backtracking_70 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f93687,plain,(
  $false),
  inference(subsumption_resolution,[],[f93686,f203])).
fof(f203,plain,(
  pp(sK1,sK0)),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  pp(sK1,sK0) & ~zex(difference(sK0,sK1)) & p(difference(sK0,sK1),sK0) & ~pp(difference(sK0,sK1),sK0)),
  inference(skolemisation,[status(esa)],[f107])).
fof(f107,plain,(
  ? [X0,X1] : (pp(X1,X0) & ~zex(difference(X0,X1)) & p(difference(X0,X1),X0) & ~pp(difference(X0,X1),X0))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ? [X0,X1] : ((pp(X1,X0) & ~zex(difference(X0,X1)) & p(difference(X0,X1),X0)) & ~pp(difference(X0,X1),X0))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ((pp(X1,X0) & ~zex(difference(X0,X1)) & p(difference(X0,X1),X0)) => pp(difference(X0,X1),X0))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : ((pp(X1,X0) & ~zex(difference(X0,X1)) & p(difference(X0,X1),X0)) => pp(difference(X0,X1),X0))),
  file(unknown,unknown)).
fof(f93686,plain,(
  ~pp(sK1,sK0)),
  inference(resolution,[],[f93684,f263])).
fof(f263,plain,(
  ( ! [X0,X1] : (p(X0,X1) | ~pp(X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((~pp(X0,X1) | (p(X0,X1) & X0 != X1)) & (~p(X0,X1) | X0 = X1 | pp(X0,X1)))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : ((~pp(X0,X1) | (p(X0,X1) & X0 != X1)) & ((~p(X0,X1) | X0 = X1) | pp(X0,X1)))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : (pp(X0,X1) <=> (p(X0,X1) & X0 != X1))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : (pp(X0,X1) <=> (p(X0,X1) & ~X0 = X1))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1] : (pp(X0,X1) <=> (p(X0,X1) & ~X0 = X1))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X54] : ! [X55] : (pp(X54,X55) <=> (p(X54,X55) & ~X54 = X55))),
  file(unknown,unknown)).
fof(f93684,plain,(
  ~p(sK1,sK0)),
  inference(backtracking_split_refutation,[],[f1690,f1734_D,f93679])).
fof(f93679,plain,(
  $false | $spl9),
  inference(subsumption_resolution,[],[f93665,f1734])).
fof(f93665,plain,(
  ~p(sK1,sK0)),
  inference(resolution,[],[f93511,f267])).
fof(f267,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X52] : ! [X53] : (p(X52,X53) <=> (cont(X52,X53) & eqdim(X52,X53)))),
  file(unknown,unknown)).
fof(f93511,plain,(
  ~eqdim(sK1,sK0)),
  inference(backtracking_split_refutation,[],[f1701,f1802_D,f93485])).
fof(f93485,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f93484,f316])).
fof(f316,plain,(
  ~zex(sK1)),
  inference(resolution,[],[f312,f203])).
fof(f312,plain,(
  ( ! [X0,X1] : (~pp(X0,X1) | ~zex(X0)) )),
  inference(resolution,[],[f308,f263])).
fof(f308,plain,(
  ( ! [X4,X3] : (~p(X3,X4) | ~zex(X3)) )),
  inference(resolution,[],[f266,f223])).
fof(f223,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X87] : ! [X88] : (zex(X87) => (~cont(X88,X87) & ~cont(X87,X88)))),
  file(unknown,unknown)).
fof(f266,plain,(
  ( ! [X0,X1] : (cont(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f93484,plain,(
  zex(sK1) | $spl10),
  inference(subsumption_resolution,[],[f93278,f1814])).
fof(f1814,plain,(
  eqdim(sK0,sK1) | $spl10),
  inference(resolution,[],[f1802,f473])).
fof(f473,plain,(
  ( ! [X10,X9] : (~eqdim(X10,X9) | eqdim(X9,X10)) )),
  inference(subsumption_resolution,[],[f460,f261])).
fof(f261,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X41] : ! [X42] : (eqdim(X41,X42) <=> (~less(X41,X42) & ~less(X42,X41)))),
  file(unknown,unknown)).
fof(f460,plain,(
  ( ! [X10,X9] : (less(X9,X10) | eqdim(X9,X10) | ~eqdim(X10,X9)) )),
  inference(resolution,[],[f262,f260])).
fof(f260,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f187])).
fof(f262,plain,(
  ( ! [X0,X1] : (less(X1,X0) | less(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f187])).
fof(f93278,plain,(
  ~eqdim(sK0,sK1) | zex(sK1) | $spl10),
  inference(resolution,[],[f93272,f219])).
fof(f219,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X81] : (~zex(X81) <=> cont(X81,X81))),
  file(unknown,unknown)).
fof(f93272,plain,(
  ( ! [X3] : (~cont(X3,sK1) | ~eqdim(sK0,X3)) ) | $spl10),
  inference(subsumption_resolution,[],[f93206,f1692])).
fof(f1692,plain,(
  ( ! [X5] : (~cont(X5,sK1) | cont(X5,sK0)) )),
  inference(resolution,[],[f1682,f282])).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (~cont(X1,X2) | cont(X0,X2) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X84] : ! [X85] : ! [X86] : ((cont(X84,X85) & cont(X85,X86)) => cont(X84,X86))),
  file(unknown,unknown)).
fof(f1682,plain,(
  cont(sK1,sK0)),
  inference(subsumption_resolution,[],[f1681,f316])).
fof(f1681,plain,(
  zex(sK1) | cont(sK1,sK0)),
  inference(subsumption_resolution,[],[f1670,f289])).
fof(f289,plain,(
  zex(sK16)),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  zex(sK16)),
  inference(skolemisation,[status(esa)],[f105])).
fof(f105,plain,(
  ? [X0] : zex(X0)),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ? [X29] : zex(X29)),
  file(unknown,unknown)).
fof(f1670,plain,(
  ~zex(sK16) | zex(sK1) | cont(sK1,sK0)),
  inference(superposition,[],[f275,f1665])).
fof(f1665,plain,(
  difference(sK1,sK0) = sK16),
  inference(resolution,[],[f1615,f203])).
fof(f1615,plain,(
  ( ! [X16] : (~pp(X16,sK0) | difference(X16,sK0) = sK16) )),
  inference(resolution,[],[f1520,f353])).
fof(f353,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | difference(X0,X1) = sK16) )),
  inference(resolution,[],[f344,f277])).
fof(f277,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X15] : ! [X16] : (zex(difference(X15,X16)) <=> (zex(X15) | cont(X15,X16)))),
  file(unknown,unknown)).
fof(f344,plain,(
  ( ! [X0] : (~zex(X0) | sK16 = X0) )),
  inference(resolution,[],[f334,f289])).
fof(f334,plain,(
  ( ! [X0,X1] : (~zex(X1) | ~zex(X0) | X0 = X1) )),
  inference(resolution,[],[f290,f297])).
fof(f297,plain,(
  ( ! [X0,X1] : (~sQ1_eqProxy(X0,X1) | X0 = X1) )),
  introduced(equality_proxy_axiom,[])).
fof(f290,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X0,X1) | ~zex(X1) | ~zex(X0)) )),
  inference(equality_proxy_replacement,[],[f227])).
fof(f227,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X36] : ! [X37] : ((zex(X36) & zex(X37)) => X36 = X37)),
  file(unknown,unknown)).
fof(f1520,plain,(
  ( ! [X0] : (cont(X0,sK0) | ~pp(X0,sK0)) )),
  inference(backward_demodulation,[],[f1496,f1306])).
fof(f1306,plain,(
  ( ! [X0] : (~pp(X0,difference(sK0,sK1)) | cont(X0,sK0)) )),
  inference(resolution,[],[f1061,f263])).
fof(f1061,plain,(
  ( ! [X1] : (~p(X1,difference(sK0,sK1)) | cont(X1,sK0)) )),
  inference(resolution,[],[f1033,f266])).
fof(f1033,plain,(
  ( ! [X5] : (~cont(X5,difference(sK0,sK1)) | cont(X5,sK0)) )),
  inference(resolution,[],[f1023,f282])).
fof(f1023,plain,(
  cont(difference(sK0,sK1),sK0)),
  inference(subsumption_resolution,[],[f1022,f204])).
fof(f204,plain,(
  ~zex(difference(sK0,sK1))),
  inference(cnf_transformation,[],[f148])).
fof(f1022,plain,(
  zex(difference(sK0,sK1)) | cont(difference(sK0,sK1),sK0)),
  inference(subsumption_resolution,[],[f1014,f289])).
fof(f1014,plain,(
  ~zex(sK16) | zex(difference(sK0,sK1)) | cont(difference(sK0,sK1),sK0)),
  inference(superposition,[],[f275,f1007])).
fof(f1007,plain,(
  difference(difference(sK0,sK1),sK0) = sK16),
  inference(resolution,[],[f832,f205])).
fof(f205,plain,(
  p(difference(sK0,sK1),sK0)),
  inference(cnf_transformation,[],[f148])).
fof(f832,plain,(
  ( ! [X4,X3] : (~p(X3,X4) | difference(X3,X4) = sK16) )),
  inference(resolution,[],[f353,f266])).
fof(f1496,plain,(
  difference(sK0,sK1) = sK0),
  inference(subsumption_resolution,[],[f1491,f206])).
fof(f206,plain,(
  ~pp(difference(sK0,sK1),sK0)),
  inference(cnf_transformation,[],[f148])).
fof(f1491,plain,(
  pp(difference(sK0,sK1),sK0) | difference(sK0,sK1) = sK0),
  inference(resolution,[],[f588,f205])).
fof(f588,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | pp(X0,X1) | X0 = X1) )),
  inference(resolution,[],[f294,f297])).
fof(f294,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X0,X1) | pp(X0,X1) | ~p(X0,X1)) )),
  inference(equality_proxy_replacement,[],[f265])).
fof(f265,plain,(
  ( ! [X0,X1] : (pp(X0,X1) | X0 = X1 | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f275,plain,(
  ( ! [X0,X1] : (~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f198])).
fof(f93206,plain,(
  ( ! [X3] : (~eqdim(sK0,X3) | ~cont(X3,sK1) | ~cont(X3,sK0)) ) | $spl10),
  inference(resolution,[],[f55531,f6097])).
fof(f6097,plain,(
  ( ! [X10,X8,X9] : (~less(intersection(X9,X10),X8) | ~cont(X8,X10) | ~cont(X8,X9)) )),
  inference(subsumption_resolution,[],[f6071,f261])).
fof(f6071,plain,(
  ( ! [X10,X8,X9] : (~cont(X8,X9) | ~cont(X8,X10) | eqdim(X8,intersection(X9,X10)) | ~less(intersection(X9,X10),X8)) )),
  inference(resolution,[],[f770,f225])).
fof(f225,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : (~less(X0,X1) | ~less(X1,X0))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X31] : ! [X32] : (less(X31,X32) => ~less(X32,X31))),
  file(unknown,unknown)).
fof(f770,plain,(
  ( ! [X6,X8,X7] : (less(X6,intersection(X8,X7)) | ~cont(X6,X8) | ~cont(X6,X7) | eqdim(X6,intersection(X8,X7))) )),
  inference(resolution,[],[f283,f269])).
fof(f269,plain,(
  ( ! [X0,X1] : (~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X43] : ! [X44] : (leq(X43,X44) <=> (less(X43,X44) | eqdim(X43,X44)))),
  file(unknown,unknown)).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0)) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X21] : ! [X22] : ! [X23] : ((cont(X23,X21) & cont(X23,X22)) => leq(X23,intersection(X21,X22)))),
  file(unknown,unknown)).
fof(f55531,plain,(
  ( ! [X4] : (less(intersection(sK0,sK1),X4) | ~eqdim(sK0,X4)) ) | $spl10),
  inference(resolution,[],[f55516,f556])).
fof(f556,plain,(
  ( ! [X4,X5,X3] : (~less(X3,X5) | less(X3,X4) | ~eqdim(X5,X4)) )),
  inference(resolution,[],[f281,f271])).
fof(f271,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f193])).
fof(f281,plain,(
  ( ! [X2,X0,X1] : (~leq(X1,X2) | less(X0,X2) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : (~less(X0,X1) | ~leq(X1,X2) | less(X0,X2))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : ((~less(X0,X1) | ~leq(X1,X2)) | less(X0,X2))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X33] : ! [X34] : ! [X35] : ((less(X33,X34) & leq(X34,X35)) => less(X33,X35))),
  file(unknown,unknown)).
fof(f55516,plain,(
  less(intersection(sK0,sK1),sK0) | $spl10),
  inference(backtracking_split_refutation,[],[f50606,f55417_D,f55515])).
fof(f55515,plain,(
  $false | ($spl10 | $spl31)),
  inference(subsumption_resolution,[],[f55508,f1814])).
fof(f55508,plain,(
  ~eqdim(sK0,sK1) | $spl31),
  inference(resolution,[],[f55417,f271])).
fof(f55417,plain,(
  ~leq(sK0,sK1) | $spl31),
  inference(cnf_transformation,[],[f55417_D])).
fof(f55417_D,plain,(
  ~leq(sK0,sK1) <=> ~$spl31),
  introduced(backtracking_splitting_component,[])).
fof(f50606,plain,(
  less(intersection(sK0,sK1),sK0) | ~leq(sK0,sK1)),
  inference(subsumption_resolution,[],[f50605,f313])).
fof(f313,plain,(
  ~zex(sK0)),
  inference(resolution,[],[f309,f205])).
fof(f309,plain,(
  ( ! [X6,X5] : (~p(X5,X6) | ~zex(X6)) )),
  inference(resolution,[],[f266,f222])).
fof(f222,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f113])).
fof(f50605,plain,(
  zex(sK0) | less(intersection(sK0,sK1),sK0) | ~leq(sK0,sK1)),
  inference(forward_demodulation,[],[f50540,f1496])).
fof(f50540,plain,(
  less(intersection(sK0,sK1),sK0) | ~leq(sK0,sK1) | zex(difference(sK0,sK1))),
  inference(superposition,[],[f855,f1496])).
fof(f855,plain,(
  ( ! [X4,X3] : (less(intersection(difference(X3,X4),X4),difference(X3,X4)) | ~leq(X3,X4) | zex(difference(X3,X4))) )),
  inference(resolution,[],[f279,f489])).
fof(f489,plain,(
  ( ! [X0] : (p(X0,X0) | zex(X0)) )),
  inference(subsumption_resolution,[],[f480,f472])).
fof(f472,plain,(
  ( ! [X0] : (eqdim(X0,X0)) )),
  inference(subsumption_resolution,[],[f455,f221])).
fof(f221,plain,(
  ( ! [X0] : (~less(X0,X0)) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X30] : ~less(X30,X30)),
  file(unknown,unknown)).
fof(f455,plain,(
  ( ! [X0] : (less(X0,X0) | eqdim(X0,X0)) )),
  inference(resolution,[],[f262,f221])).
fof(f480,plain,(
  ( ! [X0] : (~eqdim(X0,X0) | p(X0,X0) | zex(X0)) )),
  inference(resolution,[],[f268,f219])).
fof(f268,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f279,plain,(
  ( ! [X2,X0,X1] : (~p(X2,difference(X0,X1)) | less(intersection(X2,X1),X2) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~p(X2,difference(X0,X1)) | less(intersection(X2,X1),X2))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~p(X2,difference(X0,X1))) | less(intersection(X2,X1),X2))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & p(X2,difference(X0,X1))) => less(intersection(X2,X1),X2))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & p(X2,difference(X0,X1))) => less(intersection(X2,X1),X2))),
  inference(rectify,[],[f6])).
fof(f6,axiom,(
  ! [X12] : ! [X13] : ! [X14] : ((leq(X12,X13) & p(X14,difference(X12,X13))) => less(intersection(X14,X13),X14))),
  file(unknown,unknown)).
fof(f1802,plain,(
  eqdim(sK1,sK0) | $spl10),
  inference(cnf_transformation,[],[f1802_D])).
fof(f1802_D,plain,(
  eqdim(sK1,sK0) <=> ~$spl10),
  introduced(backtracking_splitting_component,[])).
fof(f1701,plain,(
  less(sK1,sK0) | eqdim(sK1,sK0)),
  inference(resolution,[],[f1685,f269])).
fof(f1685,plain,(
  leq(sK1,sK0)),
  inference(resolution,[],[f1682,f226])).
fof(f226,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1)) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X27] : ! [X28] : (cont(X27,X28) => leq(X27,X28))),
  file(unknown,unknown)).
fof(f1734,plain,(
  p(sK1,sK0) | $spl9),
  inference(cnf_transformation,[],[f1734_D])).
fof(f1734_D,plain,(
  p(sK1,sK0) <=> ~$spl9),
  introduced(backtracking_splitting_component,[])).
fof(f1690,plain,(
  ~eqdim(sK1,sK0) | p(sK1,sK0)),
  inference(resolution,[],[f1682,f268])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 2846
Passive clauses: 38173
Generated clauses: 95636
Final active clauses: 1669
Final passive clauses: 25055
Input formulas: 38
Initial clauses: 89

Duplicate literals: 2058
Fw subsumption resolutions: 6261
Fw demodulations: 3746
Bw demodulations: 347

Simple tautologies: 1940
Equational tautologies: 151
Forward subsumptions: 44964

Binary resolution: 72899
Factoring: 23
Forward superposition: 3551
Backward superposition: 4553

Backtracking splits: 52
Backtracking splits refuted: 21
Backtracking splits refuted at zero level: 9

Memory used: 16630KB
Time elapsed: 0.997 s
------------------------------
% Success in time 6.214 s
