Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2092
Passive clauses: 18675
Generated clauses: 36478
Input formulas: 38
Initial clauses: 89

Selected by SInE selection: 38
SInE iterations: 1

Duplicate literals: 362
Fw subsumption resolutions: 2083

Simple tautologies: 90
Forward subsumptions: 14690

Binary resolution: 32650
Factoring: 108

Backtracking splits: 967
Backtracking splits refuted: 160
Backtracking splits refuted at zero level: 53

Memory used: 6908KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2598
Passive clauses: 29029
Generated clauses: 73523
Input formulas: 38
Initial clauses: 88

Selected by SInE selection: 38
SInE iterations: 1

Duplicate literals: 4467
Fw subsumption resolutions: 10620

Simple tautologies: 1786
Equational tautologies: 11
Forward subsumptions: 28629

Binary resolution: 50951
Factoring: 203
Forward superposition: 7089
Backward superposition: 616

Unique components: 3

Memory used: 13432KB
Time elapsed: 1.400 s
------------------------------
% remaining time: 1181 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fsr=off:fde=none:gsp=input_only:lcm=reverse:nwc=10:ptb=off:ssec=off:stl=60:sos=on:sgo=on:sio=off:spl=backtracking:sfv=off:sp=reverse_arity_23 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f7254,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f7253])).
fof(f7253,plain,(
  sK0 != sK0),
  inference(forward_demodulation,[],[f7134,f5967])).
fof(f5967,plain,(
  difference(sK0,sK1) = sK0),
  inference(backtracking_split_refutation,[],[f5228,f5828_D,f5962])).
fof(f5962,plain,(
  $false | $spl205),
  inference(subsumption_resolution,[],[f5825,f5946])).
fof(f5946,plain,(
  ~cont(difference(sK0,sK1),sK0) | $spl205),
  inference(backtracking_split_refutation,[],[f5263,f5836_D,f5945])).
fof(f5945,plain,(
  $false | ($spl205 | $spl207)),
  inference(trivial_inequality_removal,[],[f5944])).
fof(f5944,plain,(
  sK0 != sK0 | ($spl205 | $spl207)),
  inference(forward_demodulation,[],[f5941,f5836])).
fof(f5941,plain,(
  difference(sK0,sK1) != sK0 | $spl205),
  inference(resolution,[],[f5828,f257])).
fof(f257,plain,(
  ( ! [X0,X1] : (~pp(X0,X1) | X0 != X1) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((~pp(X0,X1) | (p(X0,X1) & X0 != X1)) & (~p(X0,X1) | X0 = X1 | pp(X0,X1)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : ((~pp(X0,X1) | (p(X0,X1) & X0 != X1)) & ((~p(X0,X1) | X0 = X1) | pp(X0,X1)))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : (pp(X0,X1) <=> (p(X0,X1) & X0 != X1))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : (pp(X0,X1) <=> (p(X0,X1) & ~X0 = X1))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1] : (pp(X0,X1) <=> (p(X0,X1) & ~X0 = X1))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X54] : ! [X55] : (pp(X54,X55) <=> (p(X54,X55) & ~X54 = X55))),
  file(unknown,unknown)).
fof(f5836,plain,(
  difference(sK0,sK1) = sK0 | $spl207),
  inference(cnf_transformation,[],[f5836_D])).
fof(f5836_D,plain,(
  difference(sK0,sK1) = sK0 <=> ~$spl207),
  introduced(backtracking_splitting_component,[])).
fof(f5263,plain,(
  ~cont(difference(sK0,sK1),sK0) | difference(sK0,sK1) = sK0),
  inference(resolution,[],[f5053,f228])).
fof(f228,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1)),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : ((~cont(X0,X1) | ~cont(X1,X0)) | X0 = X1)),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ((cont(X0,X1) & cont(X1,X0)) => X0 = X1)),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1] : ((cont(X0,X1) & cont(X1,X0)) => X0 = X1)),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X82] : ! [X83] : ((cont(X82,X83) & cont(X83,X82)) => X82 = X83)),
  file(unknown,unknown)).
fof(f5053,plain,(
  cont(sK0,difference(sK0,sK1))),
  inference(backtracking_split_refutation,[],[f331,f332_D,f5050])).
fof(f5050,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f4738,f5049])).
fof(f5049,plain,(
  ~cont(sK13(difference(sK0,sK1),sK1),sK16)),
  inference(backtracking_split_refutation,[],[f4740,f4748_D,f5042])).
fof(f5042,plain,(
  $false | $spl165),
  inference(subsumption_resolution,[],[f843,f5041])).
fof(f5041,plain,(
  zex(sK13(difference(sK0,sK1),sK1)) | $spl165),
  inference(forward_demodulation,[],[f5035,f1685])).
fof(f1685,plain,(
  difference(sK13(difference(sK0,sK1),sK1),sK16) = sK13(difference(sK0,sK1),sK1)),
  inference(resolution,[],[f1375,f223])).
fof(f223,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | difference(X0,X1) = X0) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : (~less(X1,X0) | difference(X0,X1) = X0)),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X4] : ! [X5] : (less(X5,X4) => difference(X4,X5) = X4)),
  file(unknown,unknown)).
fof(f1375,plain,(
  less(sK16,sK13(difference(sK0,sK1),sK1))),
  inference(forward_demodulation,[],[f1374,f1250])).
fof(f1250,plain,(
  difference(sK0,sK0) = sK16),
  inference(forward_demodulation,[],[f1243,f1235])).
fof(f1235,plain,(
  difference(sK1,sK1) = difference(sK0,sK0)),
  inference(resolution,[],[f321,f770])).
fof(f770,plain,(
  zex(difference(sK0,sK0))),
  inference(resolution,[],[f753,f276])).
fof(f276,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X15] : ! [X16] : (zex(difference(X15,X16)) <=> (zex(X15) | cont(X15,X16)))),
  file(unknown,unknown)).
fof(f753,plain,(
  cont(sK0,sK0)),
  inference(resolution,[],[f381,f218])).
fof(f218,plain,(
  ( ! [X0] : (zex(X0) | cont(X0,X0)) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X81] : (~zex(X81) <=> cont(X81,X81))),
  file(unknown,unknown)).
fof(f381,plain,(
  ~zex(sK0)),
  inference(backtracking_split_refutation,[],[f370,f371_D,f379])).
fof(f379,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f204,f378])).
fof(f378,plain,(
  ( ! [X2] : (zex(difference(sK0,X2))) ) | $spl10),
  inference(resolution,[],[f371,f275])).
fof(f275,plain,(
  ( ! [X0,X1] : (~zex(X0) | zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f198])).
fof(f204,plain,(
  ~zex(difference(sK0,sK1))),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ~zex(sK1) & ~zex(difference(sK0,sK1)) & po(sK1,difference(sK0,sK1))),
  inference(skolemisation,[status(esa)],[f107])).
fof(f107,plain,(
  ? [X0,X1] : (~zex(X1) & ~zex(difference(X0,X1)) & po(X1,difference(X0,X1)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ? [X0,X1] : ((~zex(X1) & ~zex(difference(X0,X1))) & po(X1,difference(X0,X1)))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ((~zex(X1) & ~zex(difference(X0,X1))) => ~po(X1,difference(X0,X1)))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : ((~zex(X1) & ~zex(difference(X0,X1))) => ~po(X1,difference(X0,X1)))),
  file(unknown,unknown)).
fof(f371,plain,(
  zex(sK0) | $spl10),
  inference(cnf_transformation,[],[f371_D])).
fof(f371_D,plain,(
  zex(sK0) <=> ~$spl10),
  introduced(backtracking_splitting_component,[])).
fof(f370,plain,(
  less(sK5(sK0),sK0) | zex(sK0) | $spl9),
  inference(resolution,[],[f369,f216])).
fof(f216,plain,(
  ( ! [X0] : (mindim(X0) | less(sK5(X0),X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1)))) & (zex(X0) | (less(sK5(X0),X0) & ~zex(sK5(X0))) | mindim(X0)))),
  inference(skolemisation,[status(esa)],[f162])).
fof(f162,plain,(
  ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1)))) & (zex(X0) | ? [X2] : (less(X2,X0) & ~zex(X2)) | mindim(X0)))),
  inference(rectify,[],[f161])).
fof(f161,plain,(
  ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1)))) & (zex(X0) | ? [X1] : (less(X1,X0) & ~zex(X1)) | mindim(X0)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1)))) & ((zex(X0) | ? [X1] : (less(X1,X0) & ~zex(X1))) | mindim(X0)))),
  inference(nnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : (mindim(X0) <=> (~zex(X0) & ! [X1] : (~less(X1,X0) | zex(X1))))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (mindim(X0) <=> (~zex(X0) & ! [X1] : (less(X1,X0) => zex(X1))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X47] : (mindim(X47) <=> (~zex(X47) & ! [X48] : (less(X48,X47) => zex(X48))))),
  file(unknown,unknown)).
fof(f369,plain,(
  ~mindim(sK0) | $spl9),
  inference(cnf_transformation,[],[f369_D])).
fof(f369_D,plain,(
  ~mindim(sK0) <=> ~$spl9),
  introduced(backtracking_splitting_component,[])).
fof(f321,plain,(
  ( ! [X0] : (~zex(X0) | difference(sK1,sK1) = X0) )),
  inference(resolution,[],[f293,f226])).
fof(f226,plain,(
  ( ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X36] : ! [X37] : ((zex(X36) & zex(X37)) => X36 = X37)),
  file(unknown,unknown)).
fof(f293,plain,(
  zex(difference(sK1,sK1))),
  inference(resolution,[],[f289,f276])).
fof(f289,plain,(
  cont(sK1,sK1)),
  inference(resolution,[],[f203,f218])).
fof(f203,plain,(
  ~zex(sK1)),
  inference(cnf_transformation,[],[f148])).
fof(f1243,plain,(
  difference(sK1,sK1) = sK16),
  inference(resolution,[],[f321,f288])).
fof(f288,plain,(
  zex(sK16)),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  zex(sK16)),
  inference(skolemisation,[status(esa)],[f105])).
fof(f105,plain,(
  ? [X0] : zex(X0)),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ? [X29] : zex(X29)),
  file(unknown,unknown)).
fof(f1374,plain,(
  less(difference(sK0,sK0),sK13(difference(sK0,sK1),sK1))),
  inference(forward_demodulation,[],[f1337,f1235])).
fof(f1337,plain,(
  less(difference(sK1,sK1),sK13(difference(sK0,sK1),sK1))),
  inference(resolution,[],[f322,f843])).
fof(f322,plain,(
  ( ! [X1] : (zex(X1) | less(difference(sK1,sK1),X1)) )),
  inference(resolution,[],[f293,f227])).
fof(f227,plain,(
  ( ! [X0,X1] : (~zex(X0) | zex(X1) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : (~zex(X0) | zex(X1) | less(X0,X1))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ((~zex(X0) | zex(X1)) | less(X0,X1))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X38] : ! [X39] : ((zex(X38) & ~zex(X39)) => less(X38,X39))),
  file(unknown,unknown)).
fof(f5035,plain,(
  zex(difference(sK13(difference(sK0,sK1),sK1),sK16)) | $spl165),
  inference(resolution,[],[f4748,f276])).
fof(f843,plain,(
  ~zex(sK13(difference(sK0,sK1),sK1))),
  inference(resolution,[],[f405,f222])).
fof(f222,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X87] : ! [X88] : (zex(X87) => (~cont(X88,X87) & ~cont(X87,X88)))),
  file(unknown,unknown)).
fof(f405,plain,(
  cont(sK13(difference(sK0,sK1),sK1),sK1)),
  inference(resolution,[],[f327,f259])).
fof(f259,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X52] : ! [X53] : (p(X52,X53) <=> (cont(X52,X53) & eqdim(X52,X53)))),
  file(unknown,unknown)).
fof(f327,plain,(
  p(sK13(difference(sK0,sK1),sK1),sK1)),
  inference(resolution,[],[f205,f262])).
fof(f262,plain,(
  ( ! [X0,X1] : (~po(X0,X1) | p(sK13(X1,X0),X0)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((~po(X0,X1) | (p(sK13(X1,X0),X0) & p(sK13(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(skolemisation,[status(esa)],[f188])).
fof(f188,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(rectify,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X2] : (~p(X2,X0) | ~p(X2,X1)) | po(X0,X1)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X56] : ! [X57] : (po(X56,X57) <=> ? [X58] : (p(X58,X56) & p(X58,X57)))),
  file(unknown,unknown)).
fof(f205,plain,(
  po(sK1,difference(sK0,sK1))),
  inference(cnf_transformation,[],[f148])).
fof(f4748,plain,(
  cont(sK13(difference(sK0,sK1),sK1),sK16) | $spl165),
  inference(cnf_transformation,[],[f4748_D])).
fof(f4748_D,plain,(
  cont(sK13(difference(sK0,sK1),sK1),sK16) <=> ~$spl165),
  introduced(backtracking_splitting_component,[])).
fof(f4740,plain,(
  sK13(difference(sK0,sK1),sK1) = sK16 | cont(sK13(difference(sK0,sK1),sK1),sK16) | $spl3),
  inference(backtracking_split_refutation,[],[f4735,f4736_D,f4737])).
fof(f4737,plain,(
  $false | $spl163),
  inference(subsumption_resolution,[],[f841,f4736])).
fof(f841,plain,(
  leq(sK13(difference(sK0,sK1),sK1),sK1)),
  inference(resolution,[],[f405,f225])).
fof(f225,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1)) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X27] : ! [X28] : (cont(X27,X28) => leq(X27,X28))),
  file(unknown,unknown)).
fof(f4736,plain,(
  ~leq(sK13(difference(sK0,sK1),sK1),sK1) | $spl163),
  inference(cnf_transformation,[],[f4736_D])).
fof(f4736_D,plain,(
  ~leq(sK13(difference(sK0,sK1),sK1),sK1) <=> ~$spl163),
  introduced(backtracking_splitting_component,[])).
fof(f4735,plain,(
  sK13(difference(sK0,sK1),sK1) = sK16 | cont(sK13(difference(sK0,sK1),sK1),sK16) | ~leq(sK13(difference(sK0,sK1),sK1),sK1) | $spl3),
  inference(forward_demodulation,[],[f4734,f1250])).
fof(f4734,plain,(
  difference(sK0,sK0) = sK13(difference(sK0,sK1),sK1) | cont(sK13(difference(sK0,sK1),sK1),sK16) | ~leq(sK13(difference(sK0,sK1),sK1),sK1) | $spl3),
  inference(forward_demodulation,[],[f4733,f1235])).
fof(f4733,plain,(
  cont(sK13(difference(sK0,sK1),sK1),sK16) | ~leq(sK13(difference(sK0,sK1),sK1),sK1) | difference(sK1,sK1) = sK13(difference(sK0,sK1),sK1) | $spl3),
  inference(forward_demodulation,[],[f4726,f2350])).
fof(f2350,plain,(
  difference(sK13(difference(sK0,sK1),sK1),sK1) = sK16),
  inference(forward_demodulation,[],[f2349,f1250])).
fof(f2349,plain,(
  difference(sK0,sK0) = difference(sK13(difference(sK0,sK1),sK1),sK1)),
  inference(forward_demodulation,[],[f2344,f1235])).
fof(f2344,plain,(
  difference(sK1,sK1) = difference(sK13(difference(sK0,sK1),sK1),sK1)),
  inference(resolution,[],[f839,f321])).
fof(f839,plain,(
  zex(difference(sK13(difference(sK0,sK1),sK1),sK1))),
  inference(resolution,[],[f405,f276])).
fof(f4726,plain,(
  cont(sK13(difference(sK0,sK1),sK1),difference(sK13(difference(sK0,sK1),sK1),sK1)) | ~leq(sK13(difference(sK0,sK1),sK1),sK1) | difference(sK1,sK1) = sK13(difference(sK0,sK1),sK1) | $spl3),
  inference(resolution,[],[f1159,f1229])).
fof(f1229,plain,(
  ( ! [X0] : (cont(X0,X0) | difference(sK1,sK1) = X0) )),
  inference(resolution,[],[f321,f218])).
fof(f1159,plain,(
  ( ! [X0] : (~cont(sK13(difference(sK0,sK1),sK1),X0) | cont(sK13(difference(sK0,sK1),sK1),difference(X0,sK1)) | ~leq(X0,sK1)) ) | $spl3),
  inference(resolution,[],[f749,f277])).
fof(f277,plain,(
  ( ! [X2,X0,X1] : (~less(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)) | ~cont(X2,X0) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X0) | ~less(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X0) | ~less(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & cont(X2,X0) & less(intersection(X2,X1),X2)) => cont(X2,difference(X0,X1)))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & cont(X2,X0) & less(intersection(X2,X1),X2)) => cont(X2,difference(X0,X1)))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X6] : ! [X7] : ! [X8] : ((leq(X6,X7) & cont(X8,X6) & less(intersection(X8,X7),X8)) => cont(X8,difference(X6,X7)))),
  file(unknown,unknown)).
fof(f749,plain,(
  less(intersection(sK13(difference(sK0,sK1),sK1),sK1),sK13(difference(sK0,sK1),sK1)) | $spl3),
  inference(backtracking_split_refutation,[],[f465,f743_D,f747])).
fof(f747,plain,(
  $false | ($spl3 | $spl39)),
  inference(subsumption_resolution,[],[f743,f741])).
fof(f741,plain,(
  leq(sK0,sK1) | $spl3),
  inference(backtracking_split_refutation,[],[f338,f345_D,f739])).
fof(f739,plain,(
  $false | ($spl3 | $spl5)),
  inference(subsumption_resolution,[],[f383,f731])).
fof(f731,plain,(
  ~cont(sK0,sK0) | ($spl3 | $spl5)),
  inference(superposition,[],[f332,f367])).
fof(f367,plain,(
  difference(sK0,sK1) = sK0 | $spl5),
  inference(resolution,[],[f355,f223])).
fof(f355,plain,(
  less(sK1,sK0) | $spl5),
  inference(backtracking_split_refutation,[],[f350,f352_D,f353])).
fof(f353,plain,(
  $false | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f347,f352])).
fof(f347,plain,(
  ~eqdim(sK0,sK1) | $spl5),
  inference(resolution,[],[f345,f270])).
fof(f270,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X43] : ! [X44] : (leq(X43,X44) <=> (less(X43,X44) | eqdim(X43,X44)))),
  file(unknown,unknown)).
fof(f352,plain,(
  eqdim(sK0,sK1) | $spl6),
  inference(cnf_transformation,[],[f352_D])).
fof(f352_D,plain,(
  eqdim(sK0,sK1) <=> ~$spl6),
  introduced(backtracking_splitting_component,[])).
fof(f350,plain,(
  less(sK1,sK0) | eqdim(sK0,sK1) | $spl5),
  inference(resolution,[],[f346,f267])).
fof(f267,plain,(
  ( ! [X0,X1] : (less(X1,X0) | less(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X41] : ! [X42] : (eqdim(X41,X42) <=> (~less(X41,X42) & ~less(X42,X41)))),
  file(unknown,unknown)).
fof(f346,plain,(
  ~less(sK0,sK1) | $spl5),
  inference(resolution,[],[f345,f269])).
fof(f269,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f193])).
fof(f383,plain,(
  cont(sK0,sK0) | $spl5),
  inference(resolution,[],[f382,f218])).
fof(f382,plain,(
  ~zex(sK0) | $spl5),
  inference(forward_demodulation,[],[f204,f367])).
fof(f345,plain,(
  ~leq(sK0,sK1) | $spl5),
  inference(cnf_transformation,[],[f345_D])).
fof(f345_D,plain,(
  ~leq(sK0,sK1) <=> ~$spl5),
  introduced(backtracking_splitting_component,[])).
fof(f338,plain,(
  cont(difference(sK0,sK1),sK0) | ~leq(sK0,sK1)),
  inference(resolution,[],[f291,f279])).
fof(f279,plain,(
  ( ! [X2,X0,X1] : (~cont(X2,difference(X0,X1)) | cont(X2,X0) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,difference(X0,X1)) | cont(X2,X0))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,difference(X0,X1))) | cont(X2,X0))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & cont(X2,difference(X0,X1))) => cont(X2,X0))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & cont(X2,difference(X0,X1))) => cont(X2,X0))),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X9] : ! [X10] : ! [X11] : ((leq(X9,X10) & cont(X11,difference(X9,X10))) => cont(X11,X9))),
  file(unknown,unknown)).
fof(f291,plain,(
  cont(difference(sK0,sK1),difference(sK0,sK1))),
  inference(resolution,[],[f204,f218])).
fof(f743,plain,(
  ~leq(sK0,sK1) | $spl39),
  inference(cnf_transformation,[],[f743_D])).
fof(f743_D,plain,(
  ~leq(sK0,sK1) <=> ~$spl39),
  introduced(backtracking_splitting_component,[])).
fof(f465,plain,(
  less(intersection(sK13(difference(sK0,sK1),sK1),sK1),sK13(difference(sK0,sK1),sK1)) | ~leq(sK0,sK1)),
  inference(resolution,[],[f328,f278])).
fof(f278,plain,(
  ( ! [X2,X0,X1] : (~p(X2,difference(X0,X1)) | less(intersection(X2,X1),X2) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~p(X2,difference(X0,X1)) | less(intersection(X2,X1),X2))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~p(X2,difference(X0,X1))) | less(intersection(X2,X1),X2))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & p(X2,difference(X0,X1))) => less(intersection(X2,X1),X2))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & p(X2,difference(X0,X1))) => less(intersection(X2,X1),X2))),
  inference(rectify,[],[f6])).
fof(f6,axiom,(
  ! [X12] : ! [X13] : ! [X14] : ((leq(X12,X13) & p(X14,difference(X12,X13))) => less(intersection(X14,X13),X14))),
  file(unknown,unknown)).
fof(f328,plain,(
  p(sK13(difference(sK0,sK1),sK1),difference(sK0,sK1))),
  inference(resolution,[],[f205,f263])).
fof(f263,plain,(
  ( ! [X0,X1] : (~po(X0,X1) | p(sK13(X1,X0),X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f4738,plain,(
  cont(sK13(difference(sK0,sK1),sK1),sK16) | $spl3),
  inference(backtracking_split_refutation,[],[f4729,f4730_D,f4731])).
fof(f4731,plain,(
  $false | $spl161),
  inference(subsumption_resolution,[],[f2003,f4730])).
fof(f2003,plain,(
  leq(sK1,sK1)),
  inference(backtracking_split_refutation,[],[f1999,f2000_D,f2001])).
fof(f2001,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f295,f2000])).
fof(f295,plain,(
  leq(sK1,sK1)),
  inference(resolution,[],[f289,f225])).
fof(f2000,plain,(
  ~leq(sK1,sK1) | $spl96),
  inference(cnf_transformation,[],[f2000_D])).
fof(f2000_D,plain,(
  ~leq(sK1,sK1) <=> ~$spl96),
  introduced(backtracking_splitting_component,[])).
fof(f1999,plain,(
  ( ! [X4] : (~p(X4,sK16) | less(intersection(X4,sK1),X4) | ~leq(sK1,sK1)) )),
  inference(forward_demodulation,[],[f1993,f1250])).
fof(f1993,plain,(
  ( ! [X4] : (~p(X4,difference(sK0,sK0)) | less(intersection(X4,sK1),X4) | ~leq(sK1,sK1)) )),
  inference(superposition,[],[f278,f1235])).
fof(f4730,plain,(
  ~leq(sK1,sK1) | $spl161),
  inference(cnf_transformation,[],[f4730_D])).
fof(f4730_D,plain,(
  ~leq(sK1,sK1) <=> ~$spl161),
  introduced(backtracking_splitting_component,[])).
fof(f4729,plain,(
  cont(sK13(difference(sK0,sK1),sK1),sK16) | ~leq(sK1,sK1) | $spl3),
  inference(forward_demodulation,[],[f4728,f1250])).
fof(f4728,plain,(
  cont(sK13(difference(sK0,sK1),sK1),difference(sK0,sK0)) | ~leq(sK1,sK1) | $spl3),
  inference(forward_demodulation,[],[f4723,f1235])).
fof(f4723,plain,(
  cont(sK13(difference(sK0,sK1),sK1),difference(sK1,sK1)) | ~leq(sK1,sK1) | $spl3),
  inference(resolution,[],[f1159,f405])).
fof(f332,plain,(
  ~cont(sK0,difference(sK0,sK1)) | $spl3),
  inference(cnf_transformation,[],[f332_D])).
fof(f332_D,plain,(
  ~cont(sK0,difference(sK0,sK1)) <=> ~$spl3),
  introduced(backtracking_splitting_component,[])).
fof(f331,plain,(
  p(sK0,difference(sK0,sK1)) | ~cont(sK0,difference(sK0,sK1))),
  inference(resolution,[],[f290,f261])).
fof(f261,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f290,plain,(
  eqdim(sK0,difference(sK0,sK1))),
  inference(resolution,[],[f204,f232])).
fof(f232,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | eqdim(X0,difference(X0,X1))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) | eqdim(X0,difference(X0,X1)))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : (~zex(difference(X0,X1)) => eqdim(X0,difference(X0,X1)))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1] : (~zex(difference(X0,X1)) => eqdim(X0,difference(X0,X1)))),
  inference(rectify,[],[f2])).
fof(f2,axiom,(
  ! [X2] : ! [X3] : (~zex(difference(X2,X3)) => eqdim(X2,difference(X2,X3)))),
  file(unknown,unknown)).
fof(f5825,plain,(
  cont(difference(sK0,sK1),sK0)),
  inference(backtracking_split_refutation,[],[f338,f5054_D,f5822])).
fof(f5822,plain,(
  $false | $spl180),
  inference(subsumption_resolution,[],[f5304,f5818])).
fof(f5818,plain,(
  less(sK0,sK13(sK0,sK1)) | $spl180),
  inference(backtracking_split_refutation,[],[f5488,f5489_D,f5817])).
fof(f5817,plain,(
  $false | ($spl180 | $spl198)),
  inference(subsumption_resolution,[],[f5296,f5815])).
fof(f5815,plain,(
  less(sK1,sK13(sK0,sK1)) | ($spl180 | $spl198)),
  inference(resolution,[],[f5602,f5489])).
fof(f5602,plain,(
  ( ! [X1] : (~eqdim(sK0,X1) | less(sK1,X1)) ) | $spl180),
  inference(resolution,[],[f5079,f270])).
fof(f5079,plain,(
  ( ! [X0] : (~leq(sK0,X0) | less(sK1,X0)) ) | $spl180),
  inference(resolution,[],[f5066,f280])).
fof(f280,plain,(
  ( ! [X2,X0,X1] : (~less(X0,X1) | ~leq(X1,X2) | less(X0,X2)) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : (~less(X0,X1) | ~leq(X1,X2) | less(X0,X2))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : ((~less(X0,X1) | ~leq(X1,X2)) | less(X0,X2))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X33] : ! [X34] : ! [X35] : ((less(X33,X34) & leq(X34,X35)) => less(X33,X35))),
  file(unknown,unknown)).
fof(f5066,plain,(
  less(sK1,sK0) | $spl180),
  inference(backtracking_split_refutation,[],[f5061,f5063_D,f5064])).
fof(f5064,plain,(
  $false | ($spl180 | $spl182)),
  inference(subsumption_resolution,[],[f5057,f5063])).
fof(f5057,plain,(
  ~eqdim(sK0,sK1) | $spl180),
  inference(resolution,[],[f5054,f270])).
fof(f5063,plain,(
  eqdim(sK0,sK1) | $spl182),
  inference(cnf_transformation,[],[f5063_D])).
fof(f5063_D,plain,(
  eqdim(sK0,sK1) <=> ~$spl182),
  introduced(backtracking_splitting_component,[])).
fof(f5061,plain,(
  less(sK1,sK0) | eqdim(sK0,sK1) | $spl180),
  inference(resolution,[],[f5056,f267])).
fof(f5056,plain,(
  ~less(sK0,sK1) | $spl180),
  inference(resolution,[],[f5054,f269])).
fof(f5296,plain,(
  ~less(sK1,sK13(sK0,sK1)) | $spl180),
  inference(superposition,[],[f887,f5082])).
fof(f5082,plain,(
  difference(sK0,sK1) = sK0 | $spl180),
  inference(resolution,[],[f5066,f223])).
fof(f887,plain,(
  ~less(sK1,sK13(difference(sK0,sK1),sK1))),
  inference(resolution,[],[f406,f266])).
fof(f266,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | ~less(X1,X0)) )),
  inference(cnf_transformation,[],[f191])).
fof(f406,plain,(
  eqdim(sK13(difference(sK0,sK1),sK1),sK1)),
  inference(resolution,[],[f327,f260])).
fof(f260,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f5489,plain,(
  eqdim(sK0,sK13(sK0,sK1)) | $spl198),
  inference(cnf_transformation,[],[f5489_D])).
fof(f5489_D,plain,(
  eqdim(sK0,sK13(sK0,sK1)) <=> ~$spl198),
  introduced(backtracking_splitting_component,[])).
fof(f5488,plain,(
  less(sK0,sK13(sK0,sK1)) | eqdim(sK0,sK13(sK0,sK1)) | $spl180),
  inference(resolution,[],[f5303,f267])).
fof(f5303,plain,(
  ~less(sK13(sK0,sK1),sK0) | $spl180),
  inference(superposition,[],[f1057,f5082])).
fof(f1057,plain,(
  ~less(sK13(difference(sK0,sK1),sK1),difference(sK0,sK1))),
  inference(resolution,[],[f467,f265])).
fof(f265,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f467,plain,(
  eqdim(sK13(difference(sK0,sK1),sK1),difference(sK0,sK1))),
  inference(resolution,[],[f328,f260])).
fof(f5304,plain,(
  ~less(sK0,sK13(sK0,sK1)) | $spl180),
  inference(superposition,[],[f1058,f5082])).
fof(f1058,plain,(
  ~less(difference(sK0,sK1),sK13(difference(sK0,sK1),sK1))),
  inference(resolution,[],[f467,f266])).
fof(f5054,plain,(
  ~leq(sK0,sK1) | $spl180),
  inference(cnf_transformation,[],[f5054_D])).
fof(f5054_D,plain,(
  ~leq(sK0,sK1) <=> ~$spl180),
  introduced(backtracking_splitting_component,[])).
fof(f5828,plain,(
  pp(sK0,difference(sK0,sK1)) | $spl205),
  inference(cnf_transformation,[],[f5828_D])).
fof(f5828_D,plain,(
  pp(sK0,difference(sK0,sK1)) <=> ~$spl205),
  introduced(backtracking_splitting_component,[])).
fof(f5228,plain,(
  difference(sK0,sK1) = sK0 | pp(sK0,difference(sK0,sK1))),
  inference(resolution,[],[f5052,f258])).
fof(f258,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | X0 = X1 | pp(X0,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f5052,plain,(
  p(sK0,difference(sK0,sK1))),
  inference(backtracking_split_refutation,[],[f331,f332_D,f5050])).
fof(f7134,plain,(
  difference(sK0,sK1) != sK0),
  inference(backtracking_split_refutation,[],[f5263,f5971_D,f7130])).
fof(f7130,plain,(
  $false | $spl221),
  inference(subsumption_resolution,[],[f6920,f7122])).
fof(f7122,plain,(
  cont(sK13(sK0,sK1),sK16) | $spl221),
  inference(backtracking_split_refutation,[],[f7117,f7118_D,f7119])).
fof(f7119,plain,(
  $false | ($spl221 | $spl254)),
  inference(subsumption_resolution,[],[f6031,f7118])).
fof(f6031,plain,(
  leq(sK13(sK0,sK1),sK1) | $spl221),
  inference(superposition,[],[f841,f5971])).
fof(f7118,plain,(
  ~leq(sK13(sK0,sK1),sK1) | $spl254),
  inference(cnf_transformation,[],[f7118_D])).
fof(f7118_D,plain,(
  ~leq(sK13(sK0,sK1),sK1) <=> ~$spl254),
  introduced(backtracking_splitting_component,[])).
fof(f7117,plain,(
  cont(sK13(sK0,sK1),sK16) | ~leq(sK13(sK0,sK1),sK1) | $spl221),
  inference(forward_demodulation,[],[f7109,f6426])).
fof(f6426,plain,(
  difference(sK13(sK0,sK1),sK1) = sK16 | $spl221),
  inference(forward_demodulation,[],[f6425,f1250])).
fof(f6425,plain,(
  difference(sK0,sK0) = difference(sK13(sK0,sK1),sK1) | $spl221),
  inference(forward_demodulation,[],[f6418,f1235])).
fof(f6418,plain,(
  difference(sK1,sK1) = difference(sK13(sK0,sK1),sK1) | $spl221),
  inference(resolution,[],[f6030,f321])).
fof(f6030,plain,(
  zex(difference(sK13(sK0,sK1),sK1)) | $spl221),
  inference(superposition,[],[f839,f5971])).
fof(f7109,plain,(
  cont(sK13(sK0,sK1),difference(sK13(sK0,sK1),sK1)) | ~leq(sK13(sK0,sK1),sK1) | $spl221),
  inference(resolution,[],[f5990,f6103])).
fof(f6103,plain,(
  cont(sK13(sK0,sK1),sK13(sK0,sK1)) | $spl221),
  inference(resolution,[],[f6032,f218])).
fof(f6032,plain,(
  ~zex(sK13(sK0,sK1)) | $spl221),
  inference(superposition,[],[f843,f5971])).
fof(f5990,plain,(
  ( ! [X0] : (~cont(sK13(sK0,sK1),X0) | cont(sK13(sK0,sK1),difference(X0,sK1)) | ~leq(X0,sK1)) ) | $spl221),
  inference(forward_demodulation,[],[f5989,f5971])).
fof(f5989,plain,(
  ( ! [X0] : (cont(sK13(sK0,sK1),difference(X0,sK1)) | ~cont(sK13(difference(sK0,sK1),sK1),X0) | ~leq(X0,sK1)) ) | $spl221),
  inference(forward_demodulation,[],[f5983,f5971])).
fof(f5983,plain,(
  ( ! [X0] : (cont(sK13(difference(sK0,sK1),sK1),difference(X0,sK1)) | ~cont(sK13(difference(sK0,sK1),sK1),X0) | ~leq(X0,sK1)) )),
  inference(resolution,[],[f5832,f277])).
fof(f5832,plain,(
  less(intersection(sK13(difference(sK0,sK1),sK1),sK1),sK13(difference(sK0,sK1),sK1))),
  inference(backtracking_split_refutation,[],[f465,f5827_D,f5830])).
fof(f5830,plain,(
  $false | $spl204),
  inference(subsumption_resolution,[],[f5827,f5826])).
fof(f5826,plain,(
  leq(sK0,sK1)),
  inference(backtracking_split_refutation,[],[f338,f5054_D,f5822])).
fof(f5827,plain,(
  ~leq(sK0,sK1) | $spl204),
  inference(cnf_transformation,[],[f5827_D])).
fof(f5827_D,plain,(
  ~leq(sK0,sK1) <=> ~$spl204),
  introduced(backtracking_splitting_component,[])).
fof(f6920,plain,(
  ~cont(sK13(sK0,sK1),sK16) | $spl221),
  inference(superposition,[],[f5049,f5971])).
fof(f5971,plain,(
  difference(sK0,sK1) = sK0 | $spl221),
  inference(cnf_transformation,[],[f5971_D])).
fof(f5971_D,plain,(
  difference(sK0,sK1) = sK0 <=> ~$spl221),
  introduced(backtracking_splitting_component,[])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1135
Passive clauses: 2033
Generated clauses: 7349
Final active clauses: 450
Final passive clauses: 269
Input formulas: 38
Initial clauses: 86

Duplicate literals: 13
Trivial inequalities: 7
Bw subsumption resolutions: 184
Fw demodulations: 1622

Simple tautologies: 90
Equational tautologies: 90
Forward subsumptions: 3032
Fw demodulations to eq. taut.: 6

Binary resolution: 3940
Forward superposition: 54
Backward superposition: 556

Backtracking splits: 260
Backtracking splits refuted: 163
Backtracking splits refuted at zero level: 16

Memory used: 1279KB
Time elapsed: 0.086 s
------------------------------
% Success in time 1.955 s
