Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2257
Passive clauses: 21435
Generated clauses: 37355
Input formulas: 38
Initial clauses: 91

Selected by SInE selection: 38
SInE iterations: 1

Duplicate literals: 129
Fw subsumption resolutions: 2229

Simple tautologies: 61
Forward subsumptions: 11315

Binary resolution: 33381
Factoring: 61

Backtracking splits: 1100
Backtracking splits refuted: 261
Backtracking splits refuted at zero level: 49

Memory used: 6652KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2969
Passive clauses: 32021
Generated clauses: 97989
Input formulas: 38
Initial clauses: 90

Selected by SInE selection: 38
SInE iterations: 1

Duplicate literals: 6102
Fw subsumption resolutions: 10057

Simple tautologies: 4524
Equational tautologies: 17
Forward subsumptions: 46779

Binary resolution: 68758
Factoring: 194
Forward superposition: 12141
Backward superposition: 1056

Unique components: 5

Memory used: 13688KB
Time elapsed: 1.400 s
------------------------------
% remaining time: 1181 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fsr=off:fde=none:gsp=input_only:lcm=reverse:nwc=10:ptb=off:ssec=off:stl=60:sos=on:sgo=on:sio=off:spl=backtracking:sfv=off:sp=reverse_arity_23 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 7323
Passive clauses: 35868
Generated clauses: 255647
Input formulas: 38
Initial clauses: 88

Duplicate literals: 392
Trivial inequalities: 170
Bw subsumption resolutions: 17347
Fw demodulations: 39210
Condensations: 537

Simple tautologies: 1340
Equational tautologies: 2660
Forward subsumptions: 167037

Binary resolution: 117194
Factoring: 143
Forward superposition: 5939
Backward superposition: 10770
Equality factoring: 3

Backtracking splits: 2482
Backtracking splits refuted: 1229
Backtracking splits refuted at zero level: 25

Memory used: 38378KB
Time elapsed: 3.300 s
------------------------------
% remaining time: 1148 next slice time: 95
dis+2_4_bs=off:ep=on:nwc=1.5:nicw=on:ptb=off:ssec=off:sac=on:sio=off:spl=backtracking_70 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f337192,plain,(
  $false),
  inference(subsumption_resolution,[],[f337181,f337072])).
fof(f337072,plain,(
  leq(sK0,sK1)),
  inference(backtracking_split_refutation,[],[f336616,f337045_D,f337069])).
fof(f337069,plain,(
  $false | $spl86),
  inference(subsumption_resolution,[],[f337052,f203])).
fof(f203,plain,(
  p(sK1,sK0)),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  p(sK1,sK0) & min(sK2) & cont(sK2,sK0) & ~cont(sK2,sK1) & ~cont(sK2,difference(sK0,sK1))),
  inference(skolemisation,[status(esa)],[f107])).
fof(f107,plain,(
  ? [X0,X1,X2] : (p(X1,X0) & min(X2) & cont(X2,X0) & ~cont(X2,X1) & ~cont(X2,difference(X0,X1)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ? [X0,X1,X2] : ((p(X1,X0) & min(X2) & cont(X2,X0)) & (~cont(X2,X1) & ~cont(X2,difference(X0,X1))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1,X2] : ((p(X1,X0) & min(X2) & cont(X2,X0)) => (cont(X2,X1) | cont(X2,difference(X0,X1))))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : ! [X2] : ((p(X1,X0) & min(X2) & cont(X2,X0)) => (cont(X2,X1) | cont(X2,difference(X0,X1))))),
  file(unknown,unknown)).
fof(f337052,plain,(
  ~p(sK1,sK0) | $spl86),
  inference(resolution,[],[f337046,f482])).
fof(f482,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | ~p(X1,X0)) )),
  inference(resolution,[],[f468,f268])).
fof(f268,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X37] : ! [X38] : (p(X37,X38) <=> (cont(X37,X38) & eqdim(X37,X38)))),
  file(unknown,unknown)).
fof(f468,plain,(
  ( ! [X10,X9] : (~eqdim(X10,X9) | eqdim(X9,X10)) )),
  inference(subsumption_resolution,[],[f453,f265])).
fof(f265,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X50] : ! [X51] : (eqdim(X50,X51) <=> (~less(X50,X51) & ~less(X51,X50)))),
  file(unknown,unknown)).
fof(f453,plain,(
  ( ! [X10,X9] : (less(X9,X10) | eqdim(X9,X10) | ~eqdim(X10,X9)) )),
  inference(resolution,[],[f266,f264])).
fof(f264,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f266,plain,(
  ( ! [X0,X1] : (less(X1,X0) | less(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f337046,plain,(
  ~eqdim(sK0,sK1) | $spl86),
  inference(resolution,[],[f337045,f272])).
fof(f272,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X52] : ! [X53] : (leq(X52,X53) <=> (less(X52,X53) | eqdim(X52,X53)))),
  file(unknown,unknown)).
fof(f337045,plain,(
  ~leq(sK0,sK1) | $spl86),
  inference(cnf_transformation,[],[f337045_D])).
fof(f337045_D,plain,(
  ~leq(sK0,sK1) <=> ~$spl86),
  introduced(backtracking_splitting_component,[])).
fof(f336616,plain,(
  less(sK2,intersection(sK2,sK1)) | ~leq(sK0,sK1) | $spl70),
  inference(backtracking_split_refutation,[],[f6705,f306765_D,f336612])).
fof(f336612,plain,(
  $false | ($spl70 | $spl71)),
  inference(subsumption_resolution,[],[f336446,f741])).
fof(f741,plain,(
  p(sK2,sK2)),
  inference(backtracking_split_refutation,[],[f722,f737_D,f739])).
fof(f739,plain,(
  $false | $spl5),
  inference(subsumption_resolution,[],[f738,f304])).
fof(f304,plain,(
  ~zex(sK2)),
  inference(resolution,[],[f224,f205])).
fof(f205,plain,(
  cont(sK2,sK0)),
  inference(cnf_transformation,[],[f148])).
fof(f224,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X88] : ! [X89] : (zex(X88) => (~cont(X89,X88) & ~cont(X88,X89)))),
  file(unknown,unknown)).
fof(f738,plain,(
  zex(sK2) | $spl5),
  inference(resolution,[],[f737,f495])).
fof(f495,plain,(
  ( ! [X0] : (p(X0,X0) | zex(X0)) )),
  inference(resolution,[],[f477,f220])).
fof(f220,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X82] : (~zex(X82) <=> cont(X82,X82))),
  file(unknown,unknown)).
fof(f477,plain,(
  ( ! [X4] : (~cont(X4,X4) | p(X4,X4)) )),
  inference(resolution,[],[f269,f467])).
fof(f467,plain,(
  ( ! [X0] : (eqdim(X0,X0)) )),
  inference(subsumption_resolution,[],[f448,f222])).
fof(f222,plain,(
  ( ! [X0] : (~less(X0,X0)) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X39] : ~less(X39,X39)),
  file(unknown,unknown)).
fof(f448,plain,(
  ( ! [X0] : (less(X0,X0) | eqdim(X0,X0)) )),
  inference(resolution,[],[f266,f222])).
fof(f269,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f737,plain,(
  ~p(sK2,sK2) | $spl5),
  inference(cnf_transformation,[],[f737_D])).
fof(f737_D,plain,(
  ~p(sK2,sK2) <=> ~$spl5),
  introduced(backtracking_splitting_component,[])).
fof(f722,plain,(
  c(sK0,sK0) | ~p(sK2,sK2)),
  inference(resolution,[],[f579,f205])).
fof(f579,plain,(
  ( ! [X2,X3] : (~cont(X2,X3) | c(X3,sK0) | ~p(X2,sK2)) )),
  inference(resolution,[],[f569,f275])).
fof(f275,plain,(
  ( ! [X0,X3,X1] : (~cont(X3,X1) | c(X0,X1) | ~cont(X3,X0)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : ((~c(X0,X1) | (cont(sK15(X1,X0),X0) & cont(sK15(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))),
  inference(skolemisation,[status(esa)],[f195])).
fof(f195,plain,(
  ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))),
  inference(rectify,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | c(X0,X1)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X72] : ! [X73] : (c(X72,X73) <=> ? [X74] : (cont(X74,X72) & cont(X74,X73)))),
  file(unknown,unknown)).
fof(f569,plain,(
  ( ! [X1] : (cont(X1,sK0) | ~p(X1,sK2)) )),
  inference(resolution,[],[f560,f267])).
fof(f267,plain,(
  ( ! [X0,X1] : (cont(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f560,plain,(
  ( ! [X11] : (~cont(X11,sK2) | cont(X11,sK0)) )),
  inference(resolution,[],[f283,f205])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (~cont(X1,X2) | cont(X0,X2) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X85] : ! [X86] : ! [X87] : ((cont(X85,X86) & cont(X86,X87)) => cont(X85,X87))),
  file(unknown,unknown)).
fof(f336446,plain,(
  ~p(sK2,sK2) | ($spl70 | $spl71)),
  inference(resolution,[],[f327232,f206])).
fof(f206,plain,(
  ~cont(sK2,sK1)),
  inference(cnf_transformation,[],[f148])).
fof(f327232,plain,(
  ( ! [X30] : (cont(X30,sK1) | ~p(X30,sK2)) ) | ($spl70 | $spl71)),
  inference(superposition,[],[f287,f327178])).
fof(f327178,plain,(
  intersection(sK2,sK1) = sK2 | ($spl70 | $spl71)),
  inference(subsumption_resolution,[],[f327167,f204])).
fof(f204,plain,(
  min(sK2)),
  inference(cnf_transformation,[],[f148])).
fof(f327167,plain,(
  intersection(sK2,sK1) = sK2 | ~min(sK2) | ($spl70 | $spl71)),
  inference(resolution,[],[f326980,f1956])).
fof(f1956,plain,(
  ( ! [X6,X7] : (~p(X6,X7) | X6 = X7 | ~min(X7)) )),
  inference(resolution,[],[f599,f218])).
fof(f218,plain,(
  ( ! [X0,X1] : (~pp(X1,X0) | ~min(X0)) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((~min(X0) | (~zex(X0) & ! [X1] : ~pp(X1,X0))) & (zex(X0) | pp(sK6(X0),X0) | min(X0)))),
  inference(skolemisation,[status(esa)],[f162])).
fof(f162,plain,(
  ! [X0] : ((~min(X0) | (~zex(X0) & ! [X1] : ~pp(X1,X0))) & (zex(X0) | ? [X2] : pp(X2,X0) | min(X0)))),
  inference(rectify,[],[f161])).
fof(f161,plain,(
  ! [X0] : ((~min(X0) | (~zex(X0) & ! [X1] : ~pp(X1,X0))) & (zex(X0) | ? [X1] : pp(X1,X0) | min(X0)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0] : ((~min(X0) | (~zex(X0) & ! [X1] : ~pp(X1,X0))) & ((zex(X0) | ? [X1] : pp(X1,X0)) | min(X0)))),
  inference(nnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : (min(X0) <=> (~zex(X0) & ! [X1] : ~pp(X1,X0)))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (min(X0) <=> (~zex(X0) & ! [X1] : ~pp(X1,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X20] : (min(X20) <=> (~zex(X20) & ! [X21] : ~pp(X21,X20)))),
  file(unknown,unknown)).
fof(f599,plain,(
  ( ! [X0,X1] : (pp(X0,X1) | ~p(X0,X1) | X0 = X1) )),
  inference(resolution,[],[f295,f298])).
fof(f298,plain,(
  ( ! [X0,X1] : (~sQ1_eqProxy(X0,X1) | X0 = X1) )),
  introduced(equality_proxy_axiom,[])).
fof(f295,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X0,X1) | pp(X0,X1) | ~p(X0,X1)) )),
  inference(equality_proxy_replacement,[],[f263])).
fof(f263,plain,(
  ( ! [X0,X1] : (pp(X0,X1) | X0 = X1 | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : ((~pp(X0,X1) | (p(X0,X1) & X0 != X1)) & (~p(X0,X1) | X0 = X1 | pp(X0,X1)))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : ((~pp(X0,X1) | (p(X0,X1) & X0 != X1)) & ((~p(X0,X1) | X0 = X1) | pp(X0,X1)))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : (pp(X0,X1) <=> (p(X0,X1) & X0 != X1))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (pp(X0,X1) <=> (p(X0,X1) & ~X0 = X1))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1] : (pp(X0,X1) <=> (p(X0,X1) & ~X0 = X1))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X32] : ! [X33] : (pp(X32,X33) <=> (p(X32,X33) & ~X32 = X33))),
  file(unknown,unknown)).
fof(f326980,plain,(
  p(intersection(sK2,sK1),sK2) | ($spl70 | $spl71)),
  inference(backtracking_split_refutation,[],[f306781,f326971_D,f326977])).
fof(f326977,plain,(
  $false | ($spl70 | $spl85)),
  inference(subsumption_resolution,[],[f326972,f306745])).
fof(f306745,plain,(
  ~zex(intersection(sK2,sK1)) | $spl70),
  inference(cnf_transformation,[],[f306745_D])).
fof(f306745_D,plain,(
  ~zex(intersection(sK2,sK1)) <=> ~$spl70),
  introduced(backtracking_splitting_component,[])).
fof(f326972,plain,(
  zex(intersection(sK2,sK1)) | $spl85),
  inference(resolution,[],[f326971,f235])).
fof(f235,plain,(
  ( ! [X0,X1] : (cont(intersection(X0,X1),X0) | zex(intersection(X0,X1))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))),
  inference(ennf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X24] : ! [X25] : (~zex(intersection(X24,X25)) => cont(intersection(X24,X25),X24))),
  file(unknown,unknown)).
fof(f326971,plain,(
  ~cont(intersection(sK2,sK1),sK2) | $spl85),
  inference(cnf_transformation,[],[f326971_D])).
fof(f326971_D,plain,(
  ~cont(intersection(sK2,sK1),sK2) <=> ~$spl85),
  introduced(backtracking_splitting_component,[])).
fof(f306781,plain,(
  p(intersection(sK2,sK1),sK2) | ~cont(intersection(sK2,sK1),sK2) | $spl71),
  inference(resolution,[],[f306765,f269])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (~p(X2,intersection(X0,X1)) | cont(X2,X1)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))),
  inference(nnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X29] : ! [X30] : ! [X31] : ((cont(X31,X29) & cont(X31,X30) & eqdim(X31,intersection(X29,X30))) <=> p(X31,intersection(X29,X30)))),
  file(unknown,unknown)).
fof(f306765,plain,(
  eqdim(intersection(sK2,sK1),sK2) | $spl71),
  inference(cnf_transformation,[],[f306765_D])).
fof(f306765_D,plain,(
  eqdim(intersection(sK2,sK1),sK2) <=> ~$spl71),
  introduced(backtracking_splitting_component,[])).
fof(f6705,plain,(
  less(sK2,intersection(sK2,sK1)) | ~leq(sK0,sK1) | eqdim(intersection(sK2,sK1),sK2)),
  inference(subsumption_resolution,[],[f6669,f205])).
fof(f6669,plain,(
  less(sK2,intersection(sK2,sK1)) | ~leq(sK0,sK1) | ~cont(sK2,sK0) | eqdim(intersection(sK2,sK1),sK2)),
  inference(resolution,[],[f1071,f207])).
fof(f207,plain,(
  ~cont(sK2,difference(sK0,sK1))),
  inference(cnf_transformation,[],[f148])).
fof(f1071,plain,(
  ( ! [X2,X0,X1] : (less(X0,intersection(X0,X2)) | cont(X0,difference(X1,X2)) | ~leq(X1,X2) | ~cont(X0,X1) | eqdim(intersection(X0,X2),X0)) )),
  inference(resolution,[],[f279,f266])).
fof(f279,plain,(
  ( ! [X2,X0,X1] : (~less(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)) | ~cont(X2,X0) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X0) | ~less(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X0) | ~less(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & cont(X2,X0) & less(intersection(X2,X1),X2)) => cont(X2,difference(X0,X1)))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & cont(X2,X0) & less(intersection(X2,X1),X2)) => cont(X2,difference(X0,X1)))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X7] : ! [X8] : ! [X9] : ((leq(X7,X8) & cont(X9,X7) & less(intersection(X9,X8),X9)) => cont(X9,difference(X7,X8)))),
  file(unknown,unknown)).
fof(f337181,plain,(
  ~leq(sK0,sK1)),
  inference(backtracking_split_refutation,[],[f4784,f306745_D,f337173])).
fof(f337173,plain,(
  $false | $spl70),
  inference(subsumption_resolution,[],[f337125,f326981])).
fof(f326981,plain,(
  cont(intersection(sK2,sK1),sK2) | $spl70),
  inference(backtracking_split_refutation,[],[f306781,f326971_D,f326977])).
fof(f337125,plain,(
  ~cont(intersection(sK2,sK1),sK2) | $spl70),
  inference(resolution,[],[f337071,f1420])).
fof(f1420,plain,(
  ( ! [X6,X5] : (~less(X6,X5) | ~cont(X5,X6)) )),
  inference(subsumption_resolution,[],[f1403,f265])).
fof(f1403,plain,(
  ( ! [X6,X5] : (eqdim(X5,X6) | ~cont(X5,X6) | ~less(X6,X5)) )),
  inference(resolution,[],[f493,f226])).
fof(f226,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : (~less(X0,X1) | ~less(X1,X0))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X40] : ! [X41] : (less(X40,X41) => ~less(X41,X40))),
  file(unknown,unknown)).
fof(f493,plain,(
  ( ! [X4,X5] : (less(X4,X5) | eqdim(X4,X5) | ~cont(X4,X5)) )),
  inference(resolution,[],[f270,f227])).
fof(f227,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X34] : ! [X35] : (cont(X34,X35) => leq(X34,X35))),
  file(unknown,unknown)).
fof(f270,plain,(
  ( ! [X0,X1] : (~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f193])).
fof(f337071,plain,(
  less(sK2,intersection(sK2,sK1)) | $spl70),
  inference(backtracking_split_refutation,[],[f336616,f337045_D,f337069])).
fof(f4784,plain,(
  ~leq(sK0,sK1) | ~zex(intersection(sK2,sK1))),
  inference(subsumption_resolution,[],[f4781,f205])).
fof(f4781,plain,(
  ~cont(sK2,sK0) | ~leq(sK0,sK1) | ~zex(intersection(sK2,sK1))),
  inference(resolution,[],[f207,f1075])).
fof(f1075,plain,(
  ( ! [X6,X8,X7] : (cont(X6,difference(X7,X8)) | ~cont(X6,X7) | ~leq(X7,X8) | ~zex(intersection(X6,X8))) )),
  inference(subsumption_resolution,[],[f1073,f224])).
fof(f1073,plain,(
  ( ! [X6,X8,X7] : (cont(X6,difference(X7,X8)) | ~cont(X6,X7) | ~leq(X7,X8) | zex(X6) | ~zex(intersection(X6,X8))) )),
  inference(resolution,[],[f279,f229])).
fof(f229,plain,(
  ( ! [X0,X1] : (less(X0,X1) | zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : (~zex(X0) | zex(X1) | less(X0,X1))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ((~zex(X0) | zex(X1)) | less(X0,X1))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X47] : ! [X48] : ((zex(X47) & ~zex(X48)) => less(X47,X48))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 5726
Passive clauses: 134731
Generated clauses: 355323
Final active clauses: 2898
Final passive clauses: 101388
Input formulas: 38
Initial clauses: 90

Duplicate literals: 10567
Fw subsumption resolutions: 15951
Fw demodulations: 23859
Bw demodulations: 1813

Simple tautologies: 2877
Equational tautologies: 641
Forward subsumptions: 173577

Binary resolution: 264423
Factoring: 111
Forward superposition: 8307
Backward superposition: 13505

Backtracking splits: 86
Backtracking splits refuted: 43
Backtracking splits refuted at zero level: 14

Memory used: 56161KB
Time elapsed: 4.443 s
------------------------------
% Success in time 9.633 s
