Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2390
Passive clauses: 24331
Generated clauses: 40010
Input formulas: 63
Initial clauses: 125

Selected by SInE selection: 63
SInE iterations: 1

Duplicate literals: 157
Fw subsumption resolutions: 1458

Simple tautologies: 76
Forward subsumptions: 12967

Binary resolution: 36884
Factoring: 32

Backtracking splits: 1132
Backtracking splits refuted: 148
Backtracking splits refuted at zero level: 12

Memory used: 7675KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2979
Passive clauses: 38994
Generated clauses: 76595
Input formulas: 63
Initial clauses: 124

Selected by SInE selection: 63
SInE iterations: 1

Duplicate literals: 3583
Fw subsumption resolutions: 5198
Fw demodulations: 17

Simple tautologies: 1442
Equational tautologies: 39
Forward subsumptions: 28638

Binary resolution: 49486
Factoring: 117
Forward superposition: 17684
Backward superposition: 812
Self superposition: 2

Unique components: 4

Memory used: 18038KB
Time elapsed: 1.400 s
------------------------------
% remaining time: 1181 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fsr=off:fde=none:gsp=input_only:lcm=reverse:nwc=10:ptb=off:ssec=off:stl=60:sos=on:sgo=on:sio=off:spl=backtracking:sfv=off:sp=reverse_arity_23 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 13710
Passive clauses: 64616
Generated clauses: 248823
Input formulas: 63
Initial clauses: 122

Duplicate literals: 317
Trivial inequalities: 305
Bw subsumption resolutions: 40720
Fw demodulations: 36076
Condensations: 1069

Simple tautologies: 899
Equational tautologies: 344
Forward subsumptions: 125183

Binary resolution: 94585
Factoring: 34
Forward superposition: 7001
Backward superposition: 4362

Backtracking splits: 4185
Backtracking splits refuted: 2599
Backtracking splits refuted at zero level: 14

Memory used: 38762KB
Time elapsed: 3.300 s
------------------------------
% remaining time: 1148 next slice time: 95
dis+2_4_bs=off:ep=on:nwc=1.5:nicw=on:ptb=off:ssec=off:sac=on:sio=off:spl=backtracking_70 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 6329
Passive clauses: 237056
Generated clauses: 552144
Input formulas: 63
Initial clauses: 124

Duplicate literals: 11379
Fw subsumption resolutions: 30027
Fw demodulations: 44995
Bw demodulations: 183

Simple tautologies: 5269
Equational tautologies: 247
Forward subsumptions: 227623

Binary resolution: 402780
Factoring: 74
Forward superposition: 36538
Backward superposition: 27132
Self superposition: 1

Backtracking splits: 97
Backtracking splits refuted: 32
Backtracking splits refuted at zero level: 10

Memory used: 149293KB
Time elapsed: 9.600 s
------------------------------
% remaining time: 1052 next slice time: 4
lrs-1003_12_bs=unit_only:bsr=on:bms=on:cond=fast:drc=off:ep=on:fde=none:lcm=predicate:nwc=10:nicw=on:ssec=off:stl=60:sgo=on:sio=off:spo=on:sfv=off:sp=occurrence:updr=off_2 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 1690
Passive clauses: 6283
Generated clauses: 27544
Input formulas: 63
Initial clauses: 124
Reactivated clauses: 115

Duplicate literals: 134
Fw subsumption resolutions: 925
Bw subsumption resolutions: 468
Fw demodulations: 1506
Bw demodulations: 7007
Condensations: 144

Simple tautologies: 498
Equational tautologies: 54
Forward subsumptions: 16726
Backward subsumptions: 3232
Bw demodulations to eq. taut.: 764
Subsumed by BDD marking: 880

Binary resolution: 13193
Factoring: 30
Forward superposition: 4478
Backward superposition: 6538

Split clauses: 81
Split components: 163
Unique components: 5453
Introduced splitting names: 32
BDD propositional clauses: 11

Memory used: 8315KB
Time elapsed: 0.500 s
------------------------------
% remaining time: 1047 next slice time: 22
lrs+1010_50_bs=off:drc=off:flr=on:lcm=predicate:nwc=1.2:stl=60:sgo=on:sfv=off_16 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 4387
Passive clauses: 12529
Generated clauses: 242960
Input formulas: 63
Initial clauses: 122

Duplicate literals: 41245
Fw subsumption resolutions: 16058
Fw literal rewrites: 4670

Simple tautologies: 9084
Equational tautologies: 92
Forward subsumptions: 162130

Binary resolution: 174768
Factoring: 36
Forward superposition: 6612
Backward superposition: 2384
Self superposition: 1
Equality resolution: 1

Unique components: 4

Memory used: 12025KB
Time elapsed: 2.300 s
------------------------------
% remaining time: 1024 next slice time: 41
dis-1010_5_bd=off:bs=off:cond=fast:ep=on:fde=none:lcm=predicate:nwc=1.3:nicw=on:ptb=off:ssec=off:sos=on:sac=on:sgo=on:sio=off:spl=backtracking:updr=off_30 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 7021
Passive clauses: 67213
Generated clauses: 155590
Input formulas: 63
Initial clauses: 124

Duplicate literals: 9152
Fw subsumption resolutions: 6872
Fw demodulations: 45823
Condensations: 45

Simple tautologies: 140
Forward subsumptions: 27410

Binary resolution: 84343
Factoring: 577
Forward superposition: 9335

Backtracking splits: 89
Backtracking splits refuted: 42

Memory used: 50020KB
Time elapsed: 4.200 s
------------------------------
% remaining time: 982 next slice time: 237
lrs+1_2:1_bs=off:ep=on:gsp=input_only:lcm=predicate:nwc=1.1:nicw=on:ptb=off:ssec=off:stl=60:sos=on:sac=on:sgo=on:sio=off:spo=on:spl=backtracking:sp=occurrence:updr=off_175 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 5497
Passive clauses: 167717
Generated clauses: 384102
Input formulas: 63
Initial clauses: 124
Discarded non-redundant clauses: 62555

Duplicate literals: 1164
Fw subsumption resolutions: 7205
Fw demodulations: 44331
Bw demodulations: 81296

Simple tautologies: 1224
Equational tautologies: 47
Forward subsumptions: 160457
Bw demodulations to eq. taut.: 8610

Binary resolution: 148675
Factoring: 343
Forward superposition: 8681
Backward superposition: 5095

Backtracking splits: 88
Backtracking splits refuted: 21
Backtracking splits refuted at zero level: 6

Memory used: 53602KB
Time elapsed: 23.800 s
------------------------------
% remaining time: 744 next slice time: 311
lrs+11_40_bd=off:bs=unit_only:bsr=unit_only:drc=off:flr=on:lcm=predicate:nwc=1:ptb=off:ssec=off:stl=60:sac=on:spl=backtracking:sfv=off:updr=off_230 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f132418,plain,(
  $false),
  inference(subsumption_resolution,[],[f131766,f322])).
fof(f322,plain,(
  eqdim(sK0,sK1)),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  eqdim(sK0,sK1) & ~zex(sK0) & ~zex(sK1) & ~eqdim(sK0,sum(sK0,sK1))),
  inference(skolemisation,[status(esa)],[f181])).
fof(f181,plain,(
  ? [X0,X1] : (eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & ~eqdim(X0,sum(X0,X1)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ? [X0,X1] : ((eqdim(X0,X1) & ~zex(X0) & ~zex(X1)) & ~eqdim(X0,sum(X0,X1)))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ~! [X0,X1] : ((eqdim(X0,X1) & ~zex(X0) & ~zex(X1)) => eqdim(X0,sum(X0,X1)))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ~! [X0] : ! [X1] : ((eqdim(X0,X1) & ~zex(X0) & ~zex(X1)) => eqdim(X0,sum(X0,X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X11] : ! [X12] : ((eqdim(X11,X12) & ~zex(X11) & ~zex(X12)) => eqdim(X11,sum(X11,X12)))),
  file(unknown,unknown)).
fof(f131766,plain,(
  ~eqdim(sK0,sK1)),
  inference(superposition,[],[f325,f131761])).
fof(f131761,plain,(
  sum(sK0,sK1) = sK1),
  inference(subsumption_resolution,[],[f13709,f131759])).
fof(f131759,plain,(
  cont(sum(sK0,sK1),sK1)),
  inference(subsumption_resolution,[],[f131758,f17600])).
fof(f17600,plain,(
  ~zex(sum(sK0,sK1))),
  inference(forward_literal_rewriting,[],[f17598,f340])).
fof(f340,plain,(
  ( ! [X0] : (~zex(X0) | ~cont(X0,X0)) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X145] : (~zex(X145) <=> cont(X145,X145))),
  file(unknown,unknown)).
fof(f17598,plain,(
  cont(sum(sK0,sK1),sum(sK0,sK1))),
  inference(resolution,[],[f17574,f415])).
fof(f415,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f313])).
fof(f313,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X98] : ! [X99] : (p(X98,X99) <=> (cont(X98,X99) & eqdim(X98,X99)))),
  file(unknown,unknown)).
fof(f17574,plain,(
  p(sum(sK0,sK1),sum(sK0,sK1))),
  inference(forward_demodulation,[],[f17573,f16535])).
fof(f16535,plain,(
  sum(sK0,sK1) = sum(sK0,sum(sK0,sK1))),
  inference(resolution,[],[f16527,f349])).
fof(f349,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | sum(X0,X1) = X1) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : (~less(X0,X1) | sum(X0,X1) = X1)),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : (less(X0,X1) => sum(X0,X1) = X1)),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => sum(X0,X1) = X1)),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X15] : ! [X16] : (less(X15,X16) => sum(X15,X16) = X16)),
  file(unknown,unknown)).
fof(f16527,plain,(
  less(sK0,sum(sK0,sK1))),
  inference(subsumption_resolution,[],[f16526,f325])).
fof(f16526,plain,(
  eqdim(sK0,sum(sK0,sK1)) | less(sK0,sum(sK0,sK1))),
  inference(forward_literal_rewriting,[],[f16516,f1273])).
fof(f1273,plain,(
  ( ! [X12,X11] : (~eqdim(X12,X11) | eqdim(X11,X12)) )),
  inference(subsumption_resolution,[],[f1257,f410])).
fof(f410,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X87] : ! [X88] : (eqdim(X87,X88) <=> (~less(X87,X88) & ~less(X88,X87)))),
  file(unknown,unknown)).
fof(f1257,plain,(
  ( ! [X12,X11] : (less(X11,X12) | eqdim(X11,X12) | ~eqdim(X12,X11)) )),
  inference(resolution,[],[f411,f409])).
fof(f409,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f309])).
fof(f411,plain,(
  ( ! [X0,X1] : (less(X1,X0) | less(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f309])).
fof(f16516,plain,(
  less(sK0,sum(sK0,sK1)) | eqdim(sum(sK0,sK1),sK0)),
  inference(resolution,[],[f16483,f2509])).
fof(f2509,plain,(
  ( ! [X2,X3] : (~cont(X2,X3) | less(X2,X3) | eqdim(X3,X2)) )),
  inference(resolution,[],[f2505,f411])).
fof(f2505,plain,(
  ( ! [X8,X7] : (~less(X8,X7) | ~cont(X7,X8)) )),
  inference(subsumption_resolution,[],[f2496,f410])).
fof(f2496,plain,(
  ( ! [X8,X7] : (eqdim(X7,X8) | ~cont(X7,X8) | ~less(X8,X7)) )),
  inference(resolution,[],[f1241,f351])).
fof(f351,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~less(X1,X0)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : (~less(X0,X1) | ~less(X1,X0))),
  inference(ennf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X77] : ! [X78] : (less(X77,X78) => ~less(X78,X77))),
  file(unknown,unknown)).
fof(f1241,plain,(
  ( ! [X2,X3] : (less(X2,X3) | eqdim(X2,X3) | ~cont(X2,X3)) )),
  inference(resolution,[],[f400,f352])).
fof(f352,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1))),
  inference(ennf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X60] : ! [X61] : (cont(X60,X61) => leq(X60,X61))),
  file(unknown,unknown)).
fof(f400,plain,(
  ( ! [X0,X1] : (~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f301])).
fof(f301,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X89] : ! [X90] : (leq(X89,X90) <=> (less(X89,X90) | eqdim(X89,X90)))),
  file(unknown,unknown)).
fof(f16483,plain,(
  cont(sK0,sum(sK0,sK1))),
  inference(forward_demodulation,[],[f16482,f342])).
fof(f342,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = sum(X1,X0)) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : sum(X0,X1) = sum(X1,X0)),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X13] : ! [X14] : sum(X13,X14) = sum(X14,X13)),
  file(unknown,unknown)).
fof(f16482,plain,(
  cont(sK0,sum(sK1,sK0))),
  inference(subsumption_resolution,[],[f16445,f323])).
fof(f323,plain,(
  ~zex(sK0)),
  inference(cnf_transformation,[],[f258])).
fof(f16445,plain,(
  cont(sK0,sum(sK1,sK0)) | zex(sK0)),
  inference(resolution,[],[f14763,f322])).
fof(f14763,plain,(
  ( ! [X12,X11] : (~eqdim(X11,X12) | cont(X11,sum(X12,X11)) | zex(X11)) )),
  inference(resolution,[],[f10972,f409])).
fof(f10972,plain,(
  ( ! [X6,X7] : (less(X6,X7) | zex(X6) | cont(X6,sum(X7,X6))) )),
  inference(resolution,[],[f6809,f1274])).
fof(f1274,plain,(
  ( ! [X10,X9] : (leq(X10,X9) | less(X9,X10)) )),
  inference(subsumption_resolution,[],[f1266,f402])).
fof(f402,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f302])).
fof(f1266,plain,(
  ( ! [X10,X9] : (less(X9,X10) | eqdim(X10,X9) | leq(X10,X9)) )),
  inference(resolution,[],[f411,f401])).
fof(f401,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | leq(X0,X1)) )),
  inference(cnf_transformation,[],[f302])).
fof(f6809,plain,(
  ( ! [X0,X1] : (~leq(X1,X0) | cont(X0,sum(X1,X0)) | zex(X0)) )),
  inference(resolution,[],[f429,f339])).
fof(f339,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f274])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (~cont(X2,X1) | cont(X2,sum(X0,X1)) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X1) | cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X1)) | cont(X2,sum(X0,X1)))),
  inference(ennf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & cont(X2,X1)) => cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & cont(X2,X1)) => cont(X2,sum(X0,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X17] : ! [X18] : ! [X19] : ((leq(X17,X18) & cont(X19,X18)) => cont(X19,sum(X17,X18)))),
  file(unknown,unknown)).
fof(f17573,plain,(
  p(sum(sK0,sK1),sum(sK0,sum(sK0,sK1)))),
  inference(subsumption_resolution,[],[f17572,f1272])).
fof(f1272,plain,(
  ( ! [X0] : (eqdim(X0,X0)) )),
  inference(subsumption_resolution,[],[f1251,f341])).
fof(f341,plain,(
  ( ! [X0] : (~less(X0,X0)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X76] : ~less(X76,X76)),
  file(unknown,unknown)).
fof(f1251,plain,(
  ( ! [X0] : (less(X0,X0) | eqdim(X0,X0)) )),
  inference(resolution,[],[f411,f341])).
fof(f17572,plain,(
  ~eqdim(sum(sK0,sK1),sum(sK0,sK1)) | p(sum(sK0,sK1),sum(sK0,sum(sK0,sK1)))),
  inference(forward_demodulation,[],[f17548,f16535])).
fof(f17548,plain,(
  ~eqdim(sum(sK0,sK1),sum(sK0,sum(sK0,sK1))) | p(sum(sK0,sK1),sum(sK0,sum(sK0,sK1)))),
  inference(resolution,[],[f16523,f417])).
fof(f417,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)) )),
  inference(cnf_transformation,[],[f314])).
fof(f16523,plain,(
  cont(sum(sK0,sK1),sum(sK0,sum(sK0,sK1)))),
  inference(resolution,[],[f16483,f11073])).
fof(f11073,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | cont(X0,sum(X1,X0))) )),
  inference(subsumption_resolution,[],[f10969,f344])).
fof(f344,plain,(
  ( ! [X0,X1] : (~zex(X0) | ~cont(X1,X0)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X151] : ! [X152] : (zex(X151) => (~cont(X152,X151) & ~cont(X151,X152)))),
  file(unknown,unknown)).
fof(f10969,plain,(
  ( ! [X0,X1] : (cont(X0,sum(X1,X0)) | zex(X0) | ~cont(X1,X0)) )),
  inference(resolution,[],[f6809,f352])).
fof(f131758,plain,(
  cont(sum(sK0,sK1),sK1) | zex(sum(sK0,sK1))),
  inference(subsumption_resolution,[],[f131488,f16525])).
fof(f16525,plain,(
  ~cont(sum(sK0,sK1),sK0)),
  inference(subsumption_resolution,[],[f16515,f325])).
fof(f16515,plain,(
  eqdim(sK0,sum(sK0,sK1)) | ~cont(sum(sK0,sK1),sK0)),
  inference(resolution,[],[f16483,f2508])).
fof(f2508,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | eqdim(X1,X0) | ~cont(X0,X1)) )),
  inference(resolution,[],[f2505,f1241])).
fof(f131488,plain,(
  cont(sum(sK0,sK1),sK1) | cont(sum(sK0,sK1),sK0) | zex(sum(sK0,sK1))),
  inference(superposition,[],[f21812,f16536])).
fof(f16536,plain,(
  sum(sK0,sK1) = difference(sum(sK0,sK1),sK0)),
  inference(resolution,[],[f16527,f350])).
fof(f350,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | difference(X0,X1) = X0) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : (~less(X1,X0) | difference(X0,X1) = X0)),
  inference(ennf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : ! [X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X25] : ! [X26] : (less(X26,X25) => difference(X25,X26) = X25)),
  file(unknown,unknown)).
fof(f21812,plain,(
  ( ! [X8,X9] : (cont(difference(sum(X8,X9),X8),X9) | cont(sum(X8,X9),X8) | zex(sum(X8,X9))) )),
  inference(resolution,[],[f437,f339])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (~cont(X2,sum(X0,X1)) | cont(X2,X0) | cont(difference(X2,X0),X1)) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : (~cont(X2,sum(X0,X1)) | cont(X2,X0) | cont(difference(X2,X0),X1))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : ((~cont(X2,sum(X0,X1)) | cont(X2,X0)) | cont(difference(X2,X0),X1))),
  inference(ennf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X20] : ! [X21] : ! [X22] : ((cont(X22,sum(X20,X21)) & ~cont(X22,X20)) => cont(difference(X22,X20),X21))),
  file(unknown,unknown)).
fof(f13709,plain,(
  ~cont(sum(sK0,sK1),sK1) | sum(sK0,sK1) = sK1),
  inference(resolution,[],[f13707,f358])).
fof(f358,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1)),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X0,X1] : ((~cont(X0,X1) | ~cont(X1,X0)) | X0 = X1)),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ((cont(X0,X1) & cont(X1,X0)) => X0 = X1)),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1] : ((cont(X0,X1) & cont(X1,X0)) => X0 = X1)),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X146] : ! [X147] : ((cont(X146,X147) & cont(X147,X146)) => X146 = X147)),
  file(unknown,unknown)).
fof(f13707,plain,(
  cont(sK1,sum(sK0,sK1))),
  inference(subsumption_resolution,[],[f13688,f324])).
fof(f324,plain,(
  ~zex(sK1)),
  inference(cnf_transformation,[],[f258])).
fof(f13688,plain,(
  zex(sK1) | cont(sK1,sum(sK0,sK1))),
  inference(resolution,[],[f10970,f322])).
fof(f10970,plain,(
  ( ! [X2,X3] : (~eqdim(X3,X2) | zex(X2) | cont(X2,sum(X3,X2))) )),
  inference(resolution,[],[f6809,f402])).
fof(f325,plain,(
  ~eqdim(sK0,sum(sK0,sK1))),
  inference(cnf_transformation,[],[f258])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 4185
Passive clauses: 27477
Generated clauses: 132022
Final active clauses: 4032
Final passive clauses: 22417
Input formulas: 63
Initial clauses: 122

Duplicate literals: 396
Fw subsumption resolutions: 3869
Bw subsumption resolutions: 322
Fw demodulations: 22287
Fw literal rewrites: 2458

Simple tautologies: 346
Equational tautologies: 861
Forward subsumptions: 73696
Backward subsumptions: 706

Binary resolution: 88435
Factoring: 18
Forward superposition: 8267
Backward superposition: 5857
Self superposition: 5
Equality factoring: 2

Memory used: 24562KB
Time elapsed: 1.790 s
------------------------------
% Success in time 47.395 s
