Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 30000 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f11014,plain,(
  $false),
  inference(subsumption_resolution,[],[f11003,f116])).
fof(f116,plain,(
  eqdim(sK0,sK1)),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  eqdim(sK0,sK1) & inc(sK0,sK1)),
  inference(skolemisation,[status(esa)],[f62])).
fof(f62,plain,(
  ? [X0,X1] : (eqdim(X0,X1) & inc(X0,X1))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1] : (eqdim(X0,X1) => ~inc(X0,X1))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : (eqdim(X0,X1) => ~inc(X0,X1))),
  file(unknown,unknown)).
fof(f11003,plain,(
  ~eqdim(sK0,sK1)),
  inference(resolution,[],[f10992,f232])).
fof(f232,plain,(
  ( ! [X8,X7] : (eqdim(X7,X8) | ~eqdim(X8,X7)) )),
  inference(subsumption_resolution,[],[f220,f163])).
fof(f163,plain,(
  ( ! [X0,X1] : (leq(X1,X0) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X39] : ! [X40] : (eqdim(X39,X40) <=> (leq(X39,X40) & leq(X40,X39)))),
  file(unknown,unknown)).
fof(f220,plain,(
  ( ! [X8,X7] : (~leq(X7,X8) | eqdim(X7,X8) | ~eqdim(X8,X7)) )),
  inference(resolution,[],[f164,f162])).
fof(f162,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f114])).
fof(f164,plain,(
  ( ! [X0,X1] : (~leq(X1,X0) | ~leq(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f114])).
fof(f10992,plain,(
  ~eqdim(sK1,sK0)),
  inference(backtracking_split_refutation,[],[f10949,f10950_D,f10987])).
fof(f10987,plain,(
  $false | $spl323),
  inference(subsumption_resolution,[],[f10967,f10130])).
fof(f10130,plain,(
  eqdim(sK4(sK0,sK1),sK0)),
  inference(backtracking_split_refutation,[],[f1604,f6132_D,f10128])).
fof(f10128,plain,(
  $false | $spl224),
  inference(subsumption_resolution,[],[f10117,f116])).
fof(f10117,plain,(
  ~eqdim(sK0,sK1) | $spl224),
  inference(resolution,[],[f10107,f232])).
fof(f10107,plain,(
  ~eqdim(sK1,sK0) | $spl224),
  inference(subsumption_resolution,[],[f10106,f9771])).
fof(f9771,plain,(
  ~sP0(sK1,sK0) | $spl224),
  inference(subsumption_resolution,[],[f9753,f116])).
fof(f9753,plain,(
  ~eqdim(sK0,sK1) | ~sP0(sK1,sK0) | $spl224),
  inference(resolution,[],[f6144,f187])).
fof(f187,plain,(
  ( ! [X0,X1] : (eqdim(sK4(X1,X0),X0) | ~sP0(X0,X1)) )),
  inference(resolution,[],[f152,f148])).
fof(f148,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X37] : ! [X38] : (p(X37,X38) <=> (cont(X37,X38) & eqdim(X37,X38)))),
  file(unknown,unknown)).
fof(f152,plain,(
  ( ! [X0,X1] : (p(sK4(X1,X0),X0) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (less(sK4(X1,X0),X1) & cont(sK4(X1,X0),X1) & p(sK4(X1,X0),X0))) & (! [X3] : (~less(X3,X1) | ~cont(X3,X1) | ~p(X3,X0)) | sP0(X0,X1)))),
  inference(skolemisation,[status(esa)],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ? [X2] : (less(X2,X1) & cont(X2,X1) & p(X2,X0))) & (! [X3] : (~less(X3,X1) | ~cont(X3,X1) | ~p(X3,X0)) | sP0(X0,X1)))),
  inference(rectify,[],[f104])).
fof(f104,plain,(
  ! [X1,X0] : ((~sP0(X1,X0) | ? [X2] : (less(X2,X0) & cont(X2,X0) & p(X2,X1))) & (! [X2] : (~less(X2,X0) | ~cont(X2,X0) | ~p(X2,X1)) | sP0(X1,X0)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X1,X0] : (sP0(X1,X0) <=> ? [X2] : (less(X2,X0) & cont(X2,X0) & p(X2,X1)))),
  introduced(predicate_definition_introduction,[])).
fof(f6144,plain,(
  ( ! [X2] : (~eqdim(sK4(sK0,sK1),X2) | ~eqdim(sK0,X2)) ) | $spl224),
  inference(resolution,[],[f6132,f1995])).
fof(f1995,plain,(
  ( ! [X14,X12,X13] : (eqdim(X12,X13) | ~eqdim(X12,X14) | ~eqdim(X13,X14)) )),
  inference(subsumption_resolution,[],[f1992,f590])).
fof(f590,plain,(
  ( ! [X14,X12,X13] : (~less(X12,X13) | ~eqdim(X13,X14) | ~eqdim(X12,X14)) )),
  inference(resolution,[],[f327,f163])).
fof(f327,plain,(
  ( ! [X17,X15,X16] : (~leq(X15,X16) | ~less(X16,X17) | ~eqdim(X17,X15)) )),
  inference(resolution,[],[f260,f162])).
fof(f260,plain,(
  ( ! [X2,X0,X1] : (~leq(X2,X0) | ~leq(X0,X1) | ~less(X1,X2)) )),
  inference(resolution,[],[f166,f237])).
fof(f237,plain,(
  ( ! [X10,X9] : (~leq(X9,X10) | ~less(X10,X9)) )),
  inference(subsumption_resolution,[],[f227,f160])).
fof(f160,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ((~less(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | less(X0,X1)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ((~less(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | less(X0,X1)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : (less(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : (less(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X35] : ! [X36] : (less(X35,X36) <=> (leq(X35,X36) & ~eqdim(X35,X36)))),
  file(unknown,unknown)).
fof(f227,plain,(
  ( ! [X10,X9] : (~leq(X9,X10) | eqdim(X10,X9) | ~less(X10,X9)) )),
  inference(resolution,[],[f164,f159])).
fof(f159,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f112])).
fof(f166,plain,(
  ( ! [X2,X0,X1] : (leq(X0,X2) | ~leq(X1,X2) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~leq(X1,X2)) | leq(X0,X2))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & leq(X1,X2)) => leq(X0,X2))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & leq(X1,X2)) => leq(X0,X2))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X24] : ! [X25] : ! [X26] : ((leq(X24,X25) & leq(X25,X26)) => leq(X24,X26))),
  file(unknown,unknown)).
fof(f1992,plain,(
  ( ! [X14,X12,X13] : (eqdim(X12,X13) | less(X12,X13) | ~eqdim(X12,X14) | ~eqdim(X13,X14)) )),
  inference(resolution,[],[f812,f163])).
fof(f812,plain,(
  ( ! [X17,X15,X16] : (~leq(X15,X16) | eqdim(X17,X16) | less(X17,X16) | ~eqdim(X17,X15)) )),
  inference(resolution,[],[f263,f162])).
fof(f263,plain,(
  ( ! [X10,X11,X9] : (~leq(X11,X9) | ~leq(X9,X10) | eqdim(X11,X10) | less(X11,X10)) )),
  inference(resolution,[],[f166,f161])).
fof(f161,plain,(
  ( ! [X0,X1] : (~leq(X0,X1) | eqdim(X0,X1) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f112])).
fof(f10106,plain,(
  ~eqdim(sK1,sK0) | sP0(sK1,sK0) | $spl224),
  inference(subsumption_resolution,[],[f10088,f117])).
fof(f117,plain,(
  inc(sK0,sK1)),
  inference(cnf_transformation,[],[f89])).
fof(f10088,plain,(
  ~eqdim(sK1,sK0) | ~inc(sK0,sK1) | sP0(sK1,sK0) | $spl224),
  inference(resolution,[],[f9803,f400])).
fof(f400,plain,(
  ( ! [X0,X1] : (eqdim(sK5(X0,X1),X1) | ~inc(X1,X0) | sP0(X0,X1)) )),
  inference(resolution,[],[f156,f148])).
fof(f156,plain,(
  ( ! [X0,X1] : (p(sK5(X1,X0),X0) | sP0(X1,X0) | ~inc(X0,X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | (less(sK5(X1,X0),X1) & cont(sK5(X1,X0),X1) & p(sK5(X1,X0),X0))) & ((~sP0(X1,X0) & ! [X3] : (~less(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))),
  inference(skolemisation,[status(esa)],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | ? [X2] : (less(X2,X1) & cont(X2,X1) & p(X2,X0))) & ((~sP0(X1,X0) & ! [X3] : (~less(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))),
  inference(rectify,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | ? [X3] : (less(X3,X1) & cont(X3,X1) & p(X3,X0))) & ((~sP0(X1,X0) & ! [X3] : (~less(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : ((~inc(X0,X1) | (sP0(X1,X0) | ? [X3] : (less(X3,X1) & cont(X3,X1) & p(X3,X0)))) & ((~sP0(X1,X0) & ! [X3] : (~less(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : (inc(X0,X1) <=> (sP0(X1,X0) | ? [X3] : (less(X3,X1) & cont(X3,X1) & p(X3,X0))))),
  inference(definition_folding,[],[f87])).
fof(f9803,plain,(
  ( ! [X2] : (~eqdim(sK5(sK1,sK0),X2) | ~eqdim(sK1,X2)) ) | $spl224),
  inference(resolution,[],[f9785,f1995])).
fof(f9785,plain,(
  ~eqdim(sK5(sK1,sK0),sK1) | $spl224),
  inference(subsumption_resolution,[],[f9773,f117])).
fof(f9773,plain,(
  ~inc(sK0,sK1) | ~eqdim(sK5(sK1,sK0),sK1) | $spl224),
  inference(resolution,[],[f9771,f358])).
fof(f358,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~inc(X1,X0) | ~eqdim(sK5(X0,X1),X0)) )),
  inference(resolution,[],[f154,f160])).
fof(f154,plain,(
  ( ! [X0,X1] : (less(sK5(X1,X0),X1) | sP0(X1,X0) | ~inc(X0,X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f6132,plain,(
  ~eqdim(sK4(sK0,sK1),sK0) | $spl224),
  inference(cnf_transformation,[],[f6132_D])).
fof(f6132_D,plain,(
  ~eqdim(sK4(sK0,sK1),sK0) <=> ~$spl224),
  introduced(backtracking_splitting_component,[])).
fof(f1604,plain,(
  ~zex(sK5(sK1,sK0)) | ~eqdim(sK4(sK0,sK1),sK0)),
  inference(resolution,[],[f502,f117])).
fof(f502,plain,(
  ( ! [X2,X3] : (~inc(X2,X3) | ~zex(sK5(X3,X2)) | ~eqdim(sK4(X2,X3),X2)) )),
  inference(resolution,[],[f381,f184])).
fof(f184,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | ~eqdim(sK4(X1,X0),X1)) )),
  inference(resolution,[],[f150,f160])).
fof(f150,plain,(
  ( ! [X0,X1] : (less(sK4(X1,X0),X1) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f106])).
fof(f381,plain,(
  ( ! [X2,X3] : (sP0(X2,X3) | ~inc(X3,X2) | ~zex(sK5(X2,X3))) )),
  inference(resolution,[],[f155,f135])).
fof(f135,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (~zex(X0) | (s(X0) & ~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : (zex(X0) => (s(X0) & ~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (s(X0) & ~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X14] : ! [X15] : (zex(X14) => (s(X14) & ~cont(X15,X14) & ~cont(X14,X15)))),
  file(unknown,unknown)).
fof(f155,plain,(
  ( ! [X0,X1] : (cont(sK5(X1,X0),X1) | sP0(X1,X0) | ~inc(X0,X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f10967,plain,(
  ~eqdim(sK4(sK0,sK1),sK0) | $spl323),
  inference(resolution,[],[f10950,f184])).
fof(f10950,plain,(
  sP0(sK1,sK0) | $spl323),
  inference(cnf_transformation,[],[f10950_D])).
fof(f10950_D,plain,(
  sP0(sK1,sK0) <=> ~$spl323),
  introduced(backtracking_splitting_component,[])).
fof(f10949,plain,(
  ~eqdim(sK1,sK0) | sP0(sK1,sK0)),
  inference(subsumption_resolution,[],[f10931,f117])).
fof(f10931,plain,(
  ~eqdim(sK1,sK0) | ~inc(sK0,sK1) | sP0(sK1,sK0)),
  inference(resolution,[],[f10332,f400])).
fof(f10332,plain,(
  ( ! [X2] : (~eqdim(sK5(sK1,sK0),X2) | ~eqdim(sK1,X2)) )),
  inference(resolution,[],[f10251,f1995])).
fof(f10251,plain,(
  ~eqdim(sK5(sK1,sK0),sK1)),
  inference(backtracking_split_refutation,[],[f2102,f6133_D,f10250])).
fof(f10250,plain,(
  $false | $spl225),
  inference(subsumption_resolution,[],[f6133,f10130])).
fof(f6133,plain,(
  ~eqdim(sK4(sK0,sK1),sK0) | $spl225),
  inference(cnf_transformation,[],[f6133_D])).
fof(f6133_D,plain,(
  ~eqdim(sK4(sK0,sK1),sK0) <=> ~$spl225),
  introduced(backtracking_splitting_component,[])).
fof(f2102,plain,(
  ~eqdim(sK5(sK1,sK0),sK1) | ~eqdim(sK4(sK0,sK1),sK0)),
  inference(resolution,[],[f545,f117])).
fof(f545,plain,(
  ( ! [X2,X3] : (~inc(X2,X3) | ~eqdim(sK5(X3,X2),X3) | ~eqdim(sK4(X2,X3),X2)) )),
  inference(resolution,[],[f358,f184])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 877
Passive clauses: 3947
Generated clauses: 11169
Final active clauses: 437
Final passive clauses: 1264
Input formulas: 23
Initial clauses: 55

Selected by SInE selection: 23
SInE iterations: 1

Duplicate literals: 337
Fw subsumption resolutions: 885

Simple tautologies: 99
Forward subsumptions: 5807

Binary resolution: 8967
Factoring: 76

Backtracking splits: 324
Backtracking splits refuted: 141
Backtracking splits refuted at zero level: 11

Memory used: 1535KB
Time elapsed: 0.131 s
------------------------------
% Success in time 0.135 s
============================= vampire_lin32 ================================
execution finished: Tue Oct 08 20:36:24 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_space_codi/theorems/conversions/inc_theorems_goal3.all.tptp
============================ end of footer ===========================
