Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 30000 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2222
Passive clauses: 11999
Generated clauses: 23262
Input formulas: 24
Initial clauses: 58

Selected by SInE selection: 24
SInE iterations: 1

Duplicate literals: 417
Fw subsumption resolutions: 1137

Simple tautologies: 169
Forward subsumptions: 9265

Binary resolution: 19715
Factoring: 263

Backtracking splits: 959
Backtracking splits refuted: 317
Backtracking splits refuted at zero level: 97

Memory used: 3709KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 29996 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f2118,plain,(
  $false),
  inference(subsumption_resolution,[],[f2117,f119])).
fof(f119,plain,(
  c(sK0,sK1)),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  c(sK0,sK1) & ! [X2] : (~cont(X2,sK0) | ~cont(X2,sK1) | (~p(X2,sK0) & ~p(X2,sK1))) & ~sc(sK0,sK1)),
  inference(skolemisation,[status(esa)],[f67])).
fof(f67,plain,(
  ? [X0,X1] : (c(X0,X1) & ! [X2] : (~cont(X2,X0) | ~cont(X2,X1) | (~p(X2,X0) & ~p(X2,X1))) & ~sc(X0,X1))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ? [X0,X1] : ((c(X0,X1) & ! [X2] : (~cont(X2,X0) | ~cont(X2,X1) | (~p(X2,X0) & ~p(X2,X1)))) & ~sc(X0,X1))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1] : ((c(X0,X1) & ! [X2] : (~cont(X2,X0) | ~cont(X2,X1) | (~p(X2,X0) & ~p(X2,X1)))) => sc(X0,X1))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : ((c(X0,X1) & ! [X2] : (~cont(X2,X0) | ~cont(X2,X1) | (~p(X2,X0) & ~p(X2,X1)))) => sc(X0,X1))),
  file(unknown,unknown)).
fof(f2117,plain,(
  ~c(sK0,sK1)),
  inference(duplicate_literal_removal,[],[f2109])).
fof(f2109,plain,(
  ~c(sK0,sK1) | ~c(sK0,sK1)),
  inference(resolution,[],[f2046,f156])).
fof(f156,plain,(
  ( ! [X0,X1] : (cont(sK6(X1,X0),X0) | ~c(X0,X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((~c(X0,X1) | (cont(sK6(X1,X0),X0) & cont(sK6(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))),
  inference(skolemisation,[status(esa)],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))),
  inference(rectify,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | c(X0,X1)))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16] : ! [X17] : (c(X16,X17) <=> ? [X18] : (cont(X18,X16) & cont(X18,X17)))),
  file(unknown,unknown)).
fof(f2046,plain,(
  ( ! [X9] : (~cont(sK6(sK1,X9),sK0) | ~c(X9,sK1)) )),
  inference(resolution,[],[f2038,f157])).
fof(f157,plain,(
  ( ! [X0,X1] : (cont(sK6(X1,X0),X1) | ~c(X0,X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f2038,plain,(
  ( ! [X10] : (~cont(X10,sK1) | ~cont(X10,sK0)) )),
  inference(subsumption_resolution,[],[f2037,f122])).
fof(f122,plain,(
  ~sc(sK0,sK1)),
  inference(cnf_transformation,[],[f96])).
fof(f2037,plain,(
  ( ! [X10] : (sc(sK0,sK1) | ~cont(X10,sK0) | ~cont(X10,sK1)) )),
  inference(subsumption_resolution,[],[f2036,f152])).
fof(f152,plain,(
  ( ! [X4,X0,X1] : (cont(sK5(X1,X0),X0) | sc(X0,X1) | ~cont(X4,X1) | ~cont(X4,X0)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | ((cont(sK4(X1,X0),X0) & cont(sK4(X1,X0),X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | (cont(sK5(X1,X0),X0) & cont(sK5(X1,X0),X1) & (~less(sK5(X1,X0),X0) | ~less(sK5(X1,X0),X1))) | sc(X0,X1)))),
  inference(skolemisation,[status(esa)],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | ? [X5] : (cont(X5,X0) & cont(X5,X1) & (~less(X5,X0) | ~less(X5,X1))) | sc(X0,X1)))),
  inference(rectify,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | ? [X3] : (cont(X3,X0) & cont(X3,X1) & (~less(X3,X0) | ~less(X3,X1))) | sc(X0,X1)))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & ((! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | ? [X3] : (cont(X3,X0) & cont(X3,X1) & (~less(X3,X0) | ~less(X3,X1)))) | sc(X0,X1)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1)))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((~cont(X3,X0) | ~cont(X3,X1)) | (less(X3,X0) & less(X3,X1)))))),
  inference(ennf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((cont(X3,X0) & cont(X3,X1)) => (less(X3,X0) & less(X3,X1)))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((cont(X3,X0) & cont(X3,X1)) => (less(X3,X0) & less(X3,X1)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X43] : ! [X44] : (sc(X43,X44) <=> (? [X45] : (cont(X45,X43) & cont(X45,X44)) & ! [X45] : ((cont(X45,X43) & cont(X45,X44)) => (less(X45,X43) & less(X45,X44)))))),
  file(unknown,unknown)).
fof(f2036,plain,(
  ( ! [X10] : (~cont(sK5(sK1,sK0),sK0) | sc(sK0,sK1) | ~cont(X10,sK0) | ~cont(X10,sK1)) )),
  inference(subsumption_resolution,[],[f2028,f153])).
fof(f153,plain,(
  ( ! [X4,X0,X1] : (cont(sK5(X1,X0),X1) | sc(X0,X1) | ~cont(X4,X1) | ~cont(X4,X0)) )),
  inference(cnf_transformation,[],[f110])).
fof(f2028,plain,(
  ( ! [X10] : (~cont(sK5(sK1,sK0),sK1) | ~cont(sK5(sK1,sK0),sK0) | sc(sK0,sK1) | ~cont(X10,sK0) | ~cont(X10,sK1)) )),
  inference(duplicate_literal_removal,[],[f2021])).
fof(f2021,plain,(
  ( ! [X10] : (sc(sK0,sK1) | ~cont(X10,sK1) | ~cont(X10,sK0) | ~cont(sK5(sK1,sK0),sK0) | ~cont(sK5(sK1,sK0),sK1) | ~cont(sK5(sK1,sK0),sK0)) )),
  inference(resolution,[],[f496,f339])).
fof(f339,plain,(
  ( ! [X2] : (less(X2,sK0) | ~cont(X2,sK1) | ~cont(X2,sK0)) )),
  inference(duplicate_literal_removal,[],[f336])).
fof(f336,plain,(
  ( ! [X2] : (~cont(X2,sK0) | ~cont(X2,sK1) | less(X2,sK0) | ~cont(X2,sK0)) )),
  inference(resolution,[],[f306,f141])).
fof(f141,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | less(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1] : (~cont(X0,X1) | less(X0,X1) | eqdim(X0,X1))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : (~cont(X0,X1) | (less(X0,X1) | eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (cont(X0,X1) => (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f2])).
fof(f2,axiom,(
  ! [X3] : ! [X4] : (cont(X3,X4) => (less(X3,X4) | eqdim(X3,X4)))),
  file(unknown,unknown)).
fof(f306,plain,(
  ( ! [X0] : (~eqdim(X0,sK0) | ~cont(X0,sK0) | ~cont(X0,sK1)) )),
  inference(duplicate_literal_removal,[],[f303])).
fof(f303,plain,(
  ( ! [X0] : (~eqdim(X0,sK0) | ~cont(X0,sK0) | ~cont(X0,sK1) | ~cont(X0,sK0)) )),
  inference(resolution,[],[f155,f120])).
fof(f120,plain,(
  ( ! [X2] : (~p(X2,sK0) | ~cont(X2,sK1) | ~cont(X2,sK0)) )),
  inference(cnf_transformation,[],[f96])).
fof(f155,plain,(
  ( ! [X0,X1] : (p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ((cont(X0,X1) & eqdim(X0,X1)) => p(X0,X1))),
  inference(unused_predicate_definition_removal,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X19] : ! [X20] : (p(X19,X20) <=> (cont(X19,X20) & eqdim(X19,X20)))),
  file(unknown,unknown)).
fof(f496,plain,(
  ( ! [X0,X1] : (~less(sK5(sK1,X0),X0) | sc(X0,sK1) | ~cont(X1,sK1) | ~cont(X1,X0) | ~cont(sK5(sK1,X0),sK0)) )),
  inference(subsumption_resolution,[],[f488,f153])).
fof(f488,plain,(
  ( ! [X0,X1] : (sc(X0,sK1) | ~less(sK5(sK1,X0),X0) | ~cont(X1,sK1) | ~cont(X1,X0) | ~cont(sK5(sK1,X0),sK1) | ~cont(sK5(sK1,X0),sK0)) )),
  inference(resolution,[],[f154,f333])).
fof(f333,plain,(
  ( ! [X2] : (less(X2,sK1) | ~cont(X2,sK1) | ~cont(X2,sK0)) )),
  inference(duplicate_literal_removal,[],[f330])).
fof(f330,plain,(
  ( ! [X2] : (~cont(X2,sK0) | ~cont(X2,sK1) | less(X2,sK1) | ~cont(X2,sK1)) )),
  inference(resolution,[],[f305,f141])).
fof(f305,plain,(
  ( ! [X1] : (~eqdim(X1,sK1) | ~cont(X1,sK0) | ~cont(X1,sK1)) )),
  inference(duplicate_literal_removal,[],[f304])).
fof(f304,plain,(
  ( ! [X1] : (~eqdim(X1,sK1) | ~cont(X1,sK1) | ~cont(X1,sK1) | ~cont(X1,sK0)) )),
  inference(resolution,[],[f155,f121])).
fof(f121,plain,(
  ( ! [X2] : (~p(X2,sK1) | ~cont(X2,sK1) | ~cont(X2,sK0)) )),
  inference(cnf_transformation,[],[f96])).
fof(f154,plain,(
  ( ! [X4,X0,X1] : (~less(sK5(X1,X0),X1) | sc(X0,X1) | ~less(sK5(X1,X0),X0) | ~cont(X4,X1) | ~cont(X4,X0)) )),
  inference(cnf_transformation,[],[f110])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 313
Passive clauses: 700
Generated clauses: 1997
Final active clauses: 313
Final passive clauses: 316
Input formulas: 24
Initial clauses: 51

Unused predicate definitions: 3
Selected by SInE selection: 24
SInE iterations: 1

Duplicate literals: 105
Fw subsumption resolutions: 223

Simple tautologies: 37
Equational tautologies: 5
Forward subsumptions: 1004

Binary resolution: 1018
Factoring: 72
Forward superposition: 504
Backward superposition: 38

Unique components: 4

Memory used: 383KB
Time elapsed: 0.030 s
------------------------------
% Success in time 0.444 s
============================= vampire_lin32 ================================
execution finished: Tue Oct 08 20:43:16 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_space_codi/theorems/conversions/sc_theorems_goal4.all.tptp
============================ end of footer ===========================
