Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 30000 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f1853,plain,(
  $false),
  inference(subsumption_resolution,[],[f1850,f119])).
fof(f119,plain,(
  sc(sK0,sK1)),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  sc(sK0,sK1) & ~c(sK0,sK1)),
  inference(skolemisation,[status(esa)],[f65])).
fof(f65,plain,(
  ? [X0,X1] : (sc(X0,X1) & ~c(X0,X1))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1] : (sc(X0,X1) => c(X0,X1))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : (sc(X0,X1) => c(X0,X1))),
  file(unknown,unknown)).
fof(f1850,plain,(
  ~sc(sK0,sK1)),
  inference(resolution,[],[f934,f1790])).
fof(f1790,plain,(
  ~c(sK1,sK0)),
  inference(resolution,[],[f888,f120])).
fof(f120,plain,(
  ~c(sK0,sK1)),
  inference(cnf_transformation,[],[f92])).
fof(f888,plain,(
  ( ! [X2,X3] : (c(X2,X3) | ~c(X3,X2)) )),
  inference(duplicate_literal_removal,[],[f884])).
fof(f884,plain,(
  ( ! [X2,X3] : (c(X2,X3) | ~c(X3,X2) | ~c(X3,X2)) )),
  inference(resolution,[],[f270,f154])).
fof(f154,plain,(
  ( ! [X0,X1] : (cont(sK4(X1,X0),X1) | ~c(X0,X1)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((~c(X0,X1) | (cont(sK4(X1,X0),X0) & cont(sK4(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))),
  inference(skolemisation,[status(esa)],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))),
  inference(rectify,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | c(X0,X1)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X18] : ! [X19] : (c(X18,X19) <=> ? [X20] : (cont(X20,X18) & cont(X20,X19)))),
  file(unknown,unknown)).
fof(f270,plain,(
  ( ! [X2,X3,X1] : (~cont(sK4(X1,X2),X3) | c(X3,X2) | ~c(X2,X1)) )),
  inference(resolution,[],[f155,f153])).
fof(f153,plain,(
  ( ! [X0,X1] : (cont(sK4(X1,X0),X0) | ~c(X0,X1)) )),
  inference(cnf_transformation,[],[f109])).
fof(f155,plain,(
  ( ! [X0,X3,X1] : (~cont(X3,X1) | ~cont(X3,X0) | c(X0,X1)) )),
  inference(cnf_transformation,[],[f109])).
fof(f934,plain,(
  ( ! [X2,X3] : (c(X2,X3) | ~sc(X3,X2)) )),
  inference(duplicate_literal_removal,[],[f931])).
fof(f931,plain,(
  ( ! [X2,X3] : (c(X2,X3) | ~sc(X3,X2) | ~sc(X3,X2)) )),
  inference(resolution,[],[f272,f157])).
fof(f157,plain,(
  ( ! [X0,X1] : (cont(sK5(X1,X0),X1) | ~sc(X0,X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | ((cont(sK5(X1,X0),X0) & cont(sK5(X1,X0),X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | (cont(sK6(X1,X0),X0) & cont(sK6(X1,X0),X1) & (~less(sK6(X1,X0),X0) | ~less(sK6(X1,X0),X1))) | sc(X0,X1)))),
  inference(skolemisation,[status(esa)],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | ? [X5] : (cont(X5,X0) & cont(X5,X1) & (~less(X5,X0) | ~less(X5,X1))) | sc(X0,X1)))),
  inference(rectify,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | ? [X3] : (cont(X3,X0) & cont(X3,X1) & (~less(X3,X0) | ~less(X3,X1))) | sc(X0,X1)))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1))))) & ((! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | ? [X3] : (cont(X3,X0) & cont(X3,X1) & (~less(X3,X0) | ~less(X3,X1)))) | sc(X0,X1)))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (less(X3,X0) & less(X3,X1)))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((~cont(X3,X0) | ~cont(X3,X1)) | (less(X3,X0) & less(X3,X1)))))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((cont(X3,X0) & cont(X3,X1)) => (less(X3,X0) & less(X3,X1)))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : (sc(X0,X1) <=> (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((cont(X3,X0) & cont(X3,X1)) => (less(X3,X0) & less(X3,X1)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X43] : ! [X44] : (sc(X43,X44) <=> (? [X45] : (cont(X45,X43) & cont(X45,X44)) & ! [X45] : ((cont(X45,X43) & cont(X45,X44)) => (less(X45,X43) & less(X45,X44)))))),
  file(unknown,unknown)).
fof(f272,plain,(
  ( ! [X8,X7,X9] : (~cont(sK5(X7,X8),X9) | c(X9,X8) | ~sc(X8,X7)) )),
  inference(resolution,[],[f155,f156])).
fof(f156,plain,(
  ( ! [X0,X1] : (cont(sK5(X1,X0),X0) | ~sc(X0,X1)) )),
  inference(cnf_transformation,[],[f113])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 386
Passive clauses: 1006
Generated clauses: 1746
Final active clauses: 247
Final passive clauses: 441
Input formulas: 24
Initial clauses: 56

Selected by SInE selection: 24
SInE iterations: 1

Duplicate literals: 33
Fw subsumption resolutions: 109

Simple tautologies: 27
Forward subsumptions: 563

Binary resolution: 1375
Factoring: 37

Backtracking splits: 66
Backtracking splits refuted: 29
Backtracking splits refuted at zero level: 22

Memory used: 511KB
Time elapsed: 0.021 s
------------------------------
% Success in time 0.038 s
============================= vampire_lin32 ================================
execution finished: Tue Oct 08 20:43:07 2013
total CPU time used: 0
The command was "vampire_lin32 --mode casc --proof tptp -t 3000"
Input read from  /stl/torsten/git/macleod/qs/multidim_space_codi/theorems/conversions/sc_theorems_goal1.all.tptp
============================ end of footer ===========================
