Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal12.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 659
Passive clauses: 2554
Generated clauses: 5153
Final active clauses: 659
Final passive clauses: 1837
Input formulas: 43
Initial clauses: 91

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 55
Fw subsumption resolutions: 87

Simple tautologies: 33
Forward subsumptions: 2484

Binary resolution: 4913
Factoring: 13

Memory used: 895KB
Time elapsed: 0.405 s
------------------------------
% remaining time: 5993 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal12.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 791
Passive clauses: 4110
Generated clauses: 9136
Final active clauses: 791
Final passive clauses: 3028
Input formulas: 43
Initial clauses: 86

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 175
Fw subsumption resolutions: 662

Simple tautologies: 99
Equational tautologies: 6
Forward subsumptions: 4391

Binary resolution: 6414
Factoring: 127
Forward superposition: 1642
Backward superposition: 46

Unique components: 4

Memory used: 1663KB
Time elapsed: 1.402 s
------------------------------
% remaining time: 5978 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal12.all
Refutation found. Thanks to Tanya!
4566. $false (0:0) [subsumption resolution 216,4562]
4562. ~p(sK2,sK1) (9:3) [resolution 4541,273]
273. eqdim(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 203]
203. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 202]
202. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[nnf transformation 94]
94. ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[flattening 93]
93. ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[rectify 7]
7. ! [X15] : ! [X16] : (p(X15,X16) <=> (cont(X15,X16) & eqdim(X15,X16)))[input]
4541. ~eqdim(sK2,sK1) (8:3) [resolution 4527,236]
236. ~lt(X0,X1) | ~eqdim(X0,X1) (0:6) [cnf transformation 133]
133. ! [X0,X1] : (~lt(X0,X1) | ~eqdim(X0,X1))[ennf transformation 62]
62. ! [X0,X1] : (lt(X0,X1) => ~eqdim(X0,X1))[flattening 61]
61. ! [X0] : ! [X1] : (lt(X0,X1) => ~eqdim(X0,X1))[rectify 31]
31. ! [X66] : ! [X67] : (lt(X66,X67) => ~eqdim(X66,X67))[input]
4527. lt(sK2,sK1) (7:3) [resolution 4419,2314]
2314. ~lt(intersection(sK0,sK1),X3) | lt(sK2,X3) (3:8) [resolution 2296,289]
289. ~lt(X0,X1) | ~lt(X1,X2) | lt(X0,X2) (0:9) [cnf transformation 165]
165. ! [X0,X1,X2] : (~lt(X0,X1) | ~lt(X1,X2) | lt(X0,X2))[flattening 164]
164. ! [X0,X1,X2] : ((~lt(X0,X1) | ~lt(X1,X2)) | lt(X0,X2))[ennf transformation 114]
114. ! [X0,X1,X2] : ((lt(X0,X1) & lt(X1,X2)) => lt(X0,X2))[flattening 113]
113. ! [X0] : ! [X1] : ! [X2] : ((lt(X0,X1) & lt(X1,X2)) => lt(X0,X2))[rectify 27]
27. ! [X57] : ! [X58] : ! [X59] : ((lt(X57,X58) & lt(X58,X59)) => lt(X57,X59))[input]
2296. lt(sK2,intersection(sK0,sK1)) (2:5) [backtracking split refutation 584,2294,2295]
2295. $false {108} (2:0) [subsumption resolution 2294,2293]
2293. leq(sK2,intersection(sK0,sK1)) (2:5) [backtracking split refutation 583,2102,2291]
2291. $false {93} (1:0) [subsumption resolution 326,2289]
2289. ~cont(sK2,sK1) {93} (3:3) [backtracking split refutation 2280,2287,2288]
2288. $false {107} (1:0) [subsumption resolution 305,2287]
305. cont(sK2,sK0) (1:3) [backtracking split refutation 298,300,303]
303. $false {1} (0:0) [subsumption resolution 215,302]
302. ~p(sK2,sK0) {1} (2:3) [resolution 300,272]
272. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 203]
215. p(sK2,sK0) (0:3) [cnf transformation 169]
169. po(sK0,sK1) & p(sK2,sK0) & p(sK2,sK1) & ~p(sK2,intersection(sK0,sK1))[skolemisation 121]
121. ? [X0,X1,X2] : (po(X0,X1) & p(X2,X0) & p(X2,X1) & ~p(X2,intersection(X0,X1)))[flattening 120]
120. ? [X0,X1,X2] : ((po(X0,X1) & p(X2,X0) & p(X2,X1)) & ~p(X2,intersection(X0,X1)))[ennf transformation 44]
44. ~! [X0,X1,X2] : ((po(X0,X1) & p(X2,X0) & p(X2,X1)) => p(X2,intersection(X0,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : ((po(X0,X1) & p(X2,X0) & p(X2,X1)) => p(X2,intersection(X0,X1)))[negated conjecture]
300. ~cont(sK2,sK0) {1} (1:3) [backtracking splitting component]
298. ~eqdim(sK2,intersection(sK0,sK1)) | ~cont(sK2,sK1) | ~cont(sK2,sK0) (1:11) [resolution 217,290]
290. p(X2,intersection(X0,X1)) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:16) [cnf transformation 210]
210. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 209]
209. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[nnf transformation 116]
116. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[flattening 115]
115. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[rectify 23]
23. ! [X49] : ! [X50] : ! [X51] : ((cont(X51,X49) & cont(X51,X50) & eqdim(X51,intersection(X49,X50))) <=> p(X51,intersection(X49,X50)))[input]
217. ~p(sK2,intersection(sK0,sK1)) (0:5) [cnf transformation 169]
2287. ~cont(sK2,sK0) {107} (3:3) [backtracking splitting component]
2280. ~cont(sK2,sK1) | ~cont(sK2,sK0) {93} (3:6) [resolution 2102,286]
286. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 159]
159. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 158]
158. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 108]
108. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[flattening 107]
107. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[rectify 22]
22. ! [X46] : ! [X47] : ! [X48] : ((cont(X48,X46) & cont(X48,X47)) => leq(X48,intersection(X46,X47)))[input]
326. cont(sK2,sK1) (1:3) [backtracking split refutation 304,306,324]
324. $false {3} (0:0) [subsumption resolution 216,314]
314. ~p(sK2,sK1) {3} (2:3) [resolution 306,272]
306. ~cont(sK2,sK1) {3} (1:3) [backtracking splitting component]
304. ~eqdim(sK2,intersection(sK0,sK1)) | ~cont(sK2,sK1) (1:8) [backtracking split refutation 298,300,303]
2102. ~leq(sK2,intersection(sK0,sK1)) {93} (2:5) [backtracking splitting component]
583. ~leq(intersection(sK0,sK1),sK2) | ~leq(sK2,intersection(sK0,sK1)) (2:10) [resolution 325,277]
277. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 205]
205. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 204]
204. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[nnf transformation 96]
96. ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[flattening 95]
95. ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[rectify 37]
37. ! [X80] : ! [X81] : (eqdim(X80,X81) <=> (leq(X80,X81) & leq(X81,X80)))[input]
325. ~eqdim(sK2,intersection(sK0,sK1)) (1:5) [backtracking split refutation 304,306,324]
2294. ~leq(sK2,intersection(sK0,sK1)) {108} (2:5) [backtracking splitting component]
584. lt(sK2,intersection(sK0,sK1)) | ~leq(sK2,intersection(sK0,sK1)) (2:10) [resolution 325,280]
280. eqdim(X0,X1) | lt(X0,X1) | ~leq(X0,X1) (0:9) [cnf transformation 207]
207. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | lt(X0,X1)))[flattening 206]
206. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | lt(X0,X1)))[nnf transformation 150]
150. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[ennf transformation 98]
98. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[flattening 97]
97. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[rectify 6]
6. ! [X13] : ! [X14] : (lt(X13,X14) <=> (leq(X13,X14) & ~eqdim(X13,X14)))[input]
4419. lt(intersection(sK0,sK1),sK1) (6:5) [backtracking split refutation 3885,3888,4417]
4417. $false {156} (0:0) [subsumption resolution 215,4413]
4413. ~p(sK2,sK0) {156} (9:3) [resolution 4400,273]
4400. ~eqdim(sK2,sK0) {156} (8:3) [resolution 4386,236]
4386. lt(sK2,sK0) {156} (7:3) [resolution 4206,2314]
4206. lt(intersection(sK0,sK1),sK0) {156} (6:5) [backtracking split refutation 3919,4133,4205]
4205. $false {156, 177} (6:0) [subsumption resolution 4133,4201]
4201. cont(intersection(sK0,sK1),sK0) {156} (8:5) [resolution 3981,243]
243. zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0) (0:9) [cnf transformation 145]
145. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 76]
76. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 75]
75. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 21]
21. ! [X44] : ! [X45] : (~zex(intersection(X44,X45)) => cont(intersection(X44,X45),X44))[input]
3981. ~zex(intersection(sK0,sK1)) {156} (7:4) [resolution 3888,230]
230. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 126]
126. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 52]
52. ! [X0,X1] : (zex(X0) => leq(X0,X1))[flattening 51]
51. ! [X0] : ! [X1] : (zex(X0) => leq(X0,X1))[rectify 41]
41. ! [X88] : ! [X89] : (zex(X88) => leq(X88,X89))[input]
4133. ~cont(intersection(sK0,sK1),sK0) {177} (6:5) [backtracking splitting component]
3919. lt(intersection(sK0,sK1),sK0) | ~cont(intersection(sK0,sK1),sK0) (6:10) [resolution 3872,233]
233. eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 129]
129. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 128]
128. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 56]
56. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 55]
55. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 24]
24. ! [X52] : ! [X53] : (cont(X52,X53) => (lt(X52,X53) | eqdim(X52,X53)))[input]
3872. ~eqdim(intersection(sK0,sK1),sK0) (5:5) [resolution 3856,235]
235. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 132]
132. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 60]
60. ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))[flattening 59]
59. ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))[rectify 29]
29. ! [X61] : ! [X62] : (eqdim(X61,X62) => eqdim(X62,X61))[input]
3856. ~eqdim(sK0,intersection(sK0,sK1)) (4:5) [resolution 2321,215]
2321. ~p(sK2,X3) | ~eqdim(X3,intersection(sK0,sK1)) (3:8) [resolution 585,273]
585. ~eqdim(sK2,X0) | ~eqdim(X0,intersection(sK0,sK1)) (2:8) [resolution 325,288]
288. eqdim(X0,X2) | ~eqdim(X1,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 163]
163. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~eqdim(X1,X2) | eqdim(X0,X2))[flattening 162]
162. ! [X0,X1,X2] : ((~eqdim(X0,X1) | ~eqdim(X1,X2)) | eqdim(X0,X2))[ennf transformation 112]
112. ! [X0,X1,X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))[flattening 111]
111. ! [X0] : ! [X1] : ! [X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))[rectify 30]
30. ! [X63] : ! [X64] : ! [X65] : ((eqdim(X63,X64) & eqdim(X64,X65)) => eqdim(X63,X65))[input]
3888. ~leq(intersection(sK0,sK1),sK1) {156} (6:5) [backtracking splitting component]
3885. lt(intersection(sK0,sK1),sK1) | ~leq(intersection(sK0,sK1),sK1) (6:10) [resolution 3862,280]
3862. ~eqdim(intersection(sK0,sK1),sK1) (5:5) [resolution 3855,235]
3855. ~eqdim(sK1,intersection(sK0,sK1)) (4:5) [resolution 2321,216]
216. p(sK2,sK1) (0:3) [cnf transformation 169]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 881
Passive clauses: 1886
Generated clauses: 4536
Final active clauses: 333
Final passive clauses: 344
Input formulas: 43
Initial clauses: 84

Unused predicate definitions: 2

Duplicate literals: 9
Bw subsumption resolutions: 142
Fw demodulations: 79

Simple tautologies: 20
Equational tautologies: 40
Forward subsumptions: 2144

Binary resolution: 3534
Forward superposition: 5
Backward superposition: 73

Backtracking splits: 215
Backtracking splits refuted: 106
Backtracking splits refuted at zero level: 19

Memory used: 767KB
Time elapsed: 0.662 s
------------------------------
% Success in time 3.082 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 23:34:04 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal12.all.tptp"
Input read from 
============================ end of footer ===========================
