Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal03.all
Refutation found. Thanks to Tanya!
5230. $false (4:0) [subsumption resolution 5229,215]
215. zex(intersection(sK1,sK0)) (0:4) [cnf transformation 168]
168. zex(intersection(sK1,sK0)) & ~zex(intersection(sK0,sK1))[skolemisation 120]
120. ? [X0,X1] : (zex(intersection(X1,X0)) & ~zex(intersection(X0,X1)))[ennf transformation 44]
44. ~! [X0,X1] : (zex(intersection(X1,X0)) => zex(intersection(X0,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : (zex(intersection(X1,X0)) => zex(intersection(X0,X1)))[negated conjecture]
5229. ~zex(intersection(sK1,sK0)) (4:4) [resolution 5210,285]
285. ~c(X0,X1) | ~zex(intersection(X0,X1)) (0:7) [cnf transformation 209]
209. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[nnf transformation 150]
150. ! [X0,X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[ennf transformation 100]
100. ! [X0,X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[flattening 99]
99. ! [X0] : ! [X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[rectify 29]
29. ! [X57] : ! [X58] : (~c(X57,X58) <=> zex(intersection(X57,X58)))[input]
5210. c(sK1,sK0) (3:3) [resolution 5207,501]
501. ~cont(intersection(sK0,sK1),X1) | c(X1,sK0) (2:8) [resolution 274,476]
476. cont(intersection(sK0,sK1),sK0) (1:5) [resolution 242,216]
216. ~zex(intersection(sK0,sK1)) (0:4) [cnf transformation 168]
242. zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0) (0:9) [cnf transformation 144]
144. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 76]
76. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 75]
75. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 30]
30. ! [X59] : ! [X60] : (~zex(intersection(X59,X60)) => cont(intersection(X59,X60),X59))[input]
274. ~cont(X3,X1) | ~cont(X3,X0) | c(X0,X1) (0:9) [cnf transformation 202]
202. ! [X0,X1] : ((~c(X0,X1) | (cont(sK11(X1,X0),X0) & cont(sK11(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 201]
201. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[rectify 200]
200. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | c(X0,X1)))[nnf transformation 92]
92. ! [X0,X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))[flattening 91]
91. ! [X0] : ! [X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))[rectify 15]
15. ! [X27] : ! [X28] : (c(X27,X28) <=> ? [X29] : (cont(X29,X27) & cont(X29,X28)))[input]
5207. cont(intersection(sK0,sK1),sK1) (2:5) [subsumption resolution 5191,218]
218. eqdim(X0,X0) (0:3) [cnf transformation 46]
46. ! [X0] : eqdim(X0,X0)[rectify 5]
5. ! [X6] : eqdim(X6,X6)[input]
5191. ~eqdim(intersection(sK0,sK1),intersection(sK0,sK1)) | cont(intersection(sK0,sK1),sK1) (2:12) [resolution 552,311]
311. cont(intersection(sK0,sK1),intersection(sK0,sK1)) (1:7) [resolution 226,216]
226. zex(X0) | cont(X0,X0) (0:5) [cnf transformation 177]
177. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[nnf transformation 123]
123. ! [X0] : (~zex(X0) <=> cont(X0,X0))[ennf transformation 49]
49. ! [X0] : (~zex(X0) <=> cont(X0,X0))[rectify 23]
23. ! [X45] : (~zex(X45) <=> cont(X45,X45))[input]
552. ~cont(X4,intersection(X5,X6)) | ~eqdim(X4,intersection(X5,X6)) | cont(X4,X6) (1:13) [resolution 277,295]
295. ~p(X2,intersection(X0,X1)) | cont(X2,X1) (0:8) [cnf transformation 211]
211. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 210]
210. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[nnf transformation 116]
116. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[flattening 115]
115. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[rectify 32]
32. ! [X64] : ! [X65] : ! [X66] : ((cont(X66,X64) & cont(X66,X65) & eqdim(X66,intersection(X64,X65))) <=> p(X66,intersection(X64,X65)))[input]
277. p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 204]
204. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 203]
203. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[nnf transformation 94]
94. ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[flattening 93]
93. ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[rectify 19]
19. ! [X36] : ! [X37] : (p(X36,X37) <=> (cont(X36,X37) & eqdim(X36,X37)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 569
Passive clauses: 2266
Generated clauses: 5013
Final active clauses: 565
Final passive clauses: 1568
Input formulas: 43
Initial clauses: 89

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 70
Fw subsumption resolutions: 112

Simple tautologies: 57
Forward subsumptions: 2552

Binary resolution: 4621
Factoring: 28

Backtracking splits: 68
Backtracking splits refuted: 12
Backtracking splits refuted at zero level: 2

Memory used: 1023KB
Time elapsed: 0.340 s
------------------------------
% Success in time 0.374 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 00:57:08 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal03.all.tptp"
Input read from 
============================ end of footer ===========================
