Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_down_theoremsT6-T10_goal09.all
Refutation found. Thanks to Tanya!
13111. $false (9:0) [subsumption resolution 13096,9389]
9389. cont(sK0,sK1) (11:3) [backtracking split refutation 9357,9368,9387]
9387. $false {79} (12:0) [subsumption resolution 9382,240]
240. zex(difference(sK0,sK1)) (0:4) [cnf transformation 222]
222. zex(difference(sK0,sK1)) & zex(difference(sK1,sK0)) & ~zex(sK0) & ~zex(sK1) & ~p(sK1,sK0)[skolemisation 138]
138. ? [X0,X1] : (zex(difference(X0,X1)) & zex(difference(X1,X0)) & ~zex(X0) & ~zex(X1) & ~p(X1,X0))[flattening 137]
137. ? [X0,X1] : ((zex(difference(X0,X1)) & zex(difference(X1,X0)) & ~zex(X0) & ~zex(X1)) & ~p(X1,X0))[ennf transformation 68]
68. ~! [X0,X1] : (~(zex(difference(X0,X1)) & zex(difference(X1,X0)) & ~zex(X0) & ~zex(X1)) | p(X1,X0))[flattening 1]
1. ~! [X0] : ! [X1] : (~(zex(difference(X0,X1)) & zex(difference(X1,X0)) & ~zex(X0) & ~zex(X1)) | p(X1,X0))[negated conjecture]
9382. ~zex(difference(sK0,sK1)) {79} (12:4) [resolution 9368,961]
961. cont(sK0,X11) | ~zex(difference(sK0,X11)) (1:7) [resolution 299,242]
242. ~zex(sK0) (0:2) [cnf transformation 222]
299. zex(X0) | cont(X0,X1) | ~zex(difference(X0,X1)) (0:9) [cnf transformation 169]
169. ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))[ennf transformation 96]
96. ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & (~(zex(X0) | cont(X0,X1)) | zex(difference(X0,X1))))[flattening 12]
12. ! [X0] : ! [X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & (~(zex(X0) | cont(X0,X1)) | zex(difference(X0,X1))))[input]
9368. ~cont(sK0,sK1) {79} (11:3) [backtracking splitting component]
9357. ~zex(difference(sK1,difference(sK2(sK0),sK2(sK0)))) | ~cont(sK0,sK1) {3} (11:11) [resolution 4029,1139]
1139. ~cont(X8,sK2(sK0)) | ~cont(sK0,X8) {3} (10:7) [resolution 344,1095]
1095. ~cont(sK0,sK2(sK0)) {3} (9:4) [subsumption resolution 1083,551]
551. ~eqdim(sK0,sK2(sK0)) {3} (8:4) [resolution 536,290]
290. leq(X0,X1) | ~eqdim(X0,X1) (0:6) [cnf transformation 164]
164. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 163]
163. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[ennf transformation 93]
93. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[flattening 34]
34. ! [X0] : ! [X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[input]
536. ~leq(sK0,sK2(sK0)) {3} (7:4) [subsumption resolution 534,242]
534. ~leq(sK0,sK2(sK0)) | zex(sK0) {3} (7:6) [resolution 531,251]
251. mindim(X0) | ~leq(X0,sK2(X0)) | zex(X0) (0:8) [cnf transformation 223]
223. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & (zex(X0) | (~zex(sK2(X0)) & ~leq(X0,sK2(X0))) | mindim(X0)))[skolemisation 141]
141. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & (zex(X0) | ? [X2] : (~zex(X2) & ~leq(X0,X2)) | mindim(X0)))[flattening 140]
140. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & ((zex(X0) | ? [X2] : (~zex(X2) & ~leq(X0,X2))) | mindim(X0)))[ennf transformation 69]
69. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & (~(~zex(X0) & ! [X2] : (zex(X2) | leq(X0,X2))) | mindim(X0)))[rectify 31]
31. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & (~(~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1))) | mindim(X0)))[input]
531. ~mindim(sK0) {3} (6:2) [backtracking splitting component]
1083. eqdim(sK0,sK2(sK0)) | ~cont(sK0,sK2(sK0)) {3} (9:8) [resolution 317,549]
549. ~lt(sK0,sK2(sK0)) {3} (8:4) [resolution 536,293]
293. leq(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 166]
166. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 165]
165. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[ennf transformation 94]
94. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[flattening 32]
32. ! [X0] : ! [X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[input]
317. lt(X0,X1) | eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 177]
177. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[ennf transformation 111]
111. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 17]
17. ! [X0] : ! [X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[input]
344. cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1) (0:9) [cnf transformation 217]
217. ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))[flattening 216]
216. ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))[ennf transformation 135]
135. ! [X0,X1,X2] : (~(cont(X0,X1) & cont(X1,X2)) | cont(X0,X2))[flattening 21]
21. ! [X0] : ! [X1] : ! [X2] : (~(cont(X0,X1) & cont(X1,X2)) | cont(X0,X2))[input]
4029. cont(sK1,X14) | ~zex(difference(sK1,difference(X14,X14))) (2:9) [resolution 2017,962]
962. cont(sK1,X12) | ~zex(difference(sK1,X12)) (1:7) [resolution 299,243]
243. ~zex(sK1) (0:2) [cnf transformation 222]
2017. ~cont(X0,difference(X1,X1)) | cont(X0,X1) (1:8) [resolution 341,245]
245. leq(X0,X0) (0:3) [cnf transformation 26]
26. ! [X0] : leq(X0,X0)[input]
341. ~leq(X0,X1) | ~cont(X2,difference(X0,X1)) | cont(X2,X0) (0:11) [cnf transformation 211]
211. ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,difference(X0,X1)) | cont(X2,X0))[flattening 210]
210. ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,difference(X0,X1))) | cont(X2,X0))[ennf transformation 132]
132. ! [X0,X1,X2] : (~(leq(X0,X1) & cont(X2,difference(X0,X1))) | cont(X2,X0))[flattening 10]
10. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & cont(X2,difference(X0,X1))) | cont(X2,X0))[input]
13096. ~cont(sK0,sK1) (9:3) [resolution 13077,1091]
1091. ~lt(X6,X5) | ~cont(X5,X6) (2:6) [subsumption resolution 1070,376]
376. ~lt(X5,X6) | ~eqdim(X6,X5) (1:6) [resolution 294,290]
294. ~leq(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 166]
1070. eqdim(X5,X6) | ~cont(X5,X6) | ~lt(X6,X5) (2:9) [resolution 317,374]
374. ~lt(X1,X2) | ~lt(X2,X1) (1:6) [resolution 294,293]
13077. lt(sK1,sK0) (8:3) [backtracking split refutation 13025,13033,13075]
13075. $false {149} (10:0) [subsumption resolution 13074,9273]
9273. cont(sK1,sK0) (11:3) [backtracking split refutation 9247,9258,9271]
9271. $false {71} (12:0) [subsumption resolution 9266,241]
241. zex(difference(sK1,sK0)) (0:4) [cnf transformation 222]
9266. ~zex(difference(sK1,sK0)) {71} (12:4) [resolution 9258,962]
9258. ~cont(sK1,sK0) {71} (11:3) [backtracking splitting component]
9247. ~zex(difference(sK0,difference(sK2(sK1),sK2(sK1)))) | ~cont(sK1,sK0) {4} (11:11) [resolution 4028,1140]
1140. ~cont(X9,sK2(sK1)) | ~cont(sK1,X9) {4} (10:7) [resolution 344,1096]
1096. ~cont(sK1,sK2(sK1)) {4} (9:4) [subsumption resolution 1085,556]
556. ~eqdim(sK1,sK2(sK1)) {4} (8:4) [resolution 540,290]
540. ~leq(sK1,sK2(sK1)) {4} (7:4) [subsumption resolution 538,243]
538. ~leq(sK1,sK2(sK1)) | zex(sK1) {4} (7:6) [resolution 532,251]
532. ~mindim(sK1) {4} (6:2) [backtracking splitting component]
1085. eqdim(sK1,sK2(sK1)) | ~cont(sK1,sK2(sK1)) {4} (9:8) [resolution 317,554]
554. ~lt(sK1,sK2(sK1)) {4} (8:4) [resolution 540,293]
4028. cont(sK0,X13) | ~zex(difference(sK0,difference(X13,X13))) (2:9) [resolution 2017,961]
13074. ~cont(sK1,sK0) {149} (10:3) [subsumption resolution 13073,244]
244. ~p(sK1,sK0) (0:3) [cnf transformation 222]
13073. p(sK1,sK0) | ~cont(sK1,sK0) {149} (10:6) [resolution 13057,298]
298. ~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 168]
168. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 167]
167. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 95]
95. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 3]
3. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
13057. eqdim(sK1,sK0) {149} (9:3) [subsumption resolution 13054,9273]
13054. eqdim(sK1,sK0) | ~cont(sK1,sK0) {149} (9:6) [resolution 13033,317]
13033. ~lt(sK1,sK0) {149} (8:3) [backtracking splitting component]
13025. ~mindim(sK4(sK0,sK0)) | ~lt(sK1,sK0) {126} (8:7) [resolution 12993,6918]
6918. lt(X15,X14) | ~mindim(X15) | ~lt(sK1,X14) (3:8) [resolution 4589,786]
786. leq(X1,X0) | lt(X0,X1) (0:6) [subsumption resolution 295,302]
302. leq(X1,X0) | leq(X0,X1) (0:6) [cnf transformation 97]
97. ! [X0,X1] : (leq(X0,X1) | leq(X1,X0))[flattening 27]
27. ! [X0] : ! [X1] : (leq(X0,X1) | leq(X1,X0))[input]
295. lt(X0,X1) | leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 166]
4589. ~leq(X66,X67) | ~lt(sK1,X66) | ~mindim(X67) (2:8) [resolution 1123,388]
388. leq(X3,sK1) | ~mindim(X3) (1:5) [resolution 249,243]
249. zex(X1) | leq(X0,X1) | ~mindim(X0) (0:7) [cnf transformation 223]
1123. ~leq(X8,X6) | ~leq(X6,X7) | ~lt(X7,X8) (1:9) [resolution 338,294]
338. leq(X0,X2) | ~leq(X1,X2) | ~leq(X0,X1) (0:9) [cnf transformation 205]
205. ! [X0,X1,X2] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2))[flattening 204]
204. ! [X0,X1,X2] : ((~leq(X0,X1) | ~leq(X1,X2)) | leq(X0,X2))[ennf transformation 129]
129. ! [X0,X1,X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[flattening 28]
28. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[input]
12993. ~lt(sK4(sK0,sK0),sK0) {126} (7:5) [backtracking split refutation 12930,12933,12979]
12979. $false {146} (12:0) [resolution 12933,9273]
12933. ~cont(X1,sK0) {146} (7:3) [backtracking splitting component]
12930. ~lt(sK4(sK0,sK0),sK0) | ~cont(X1,sK0) {126} (7:8) [duplicate literal removal 12929]
12929. ~lt(sK4(sK0,sK0),sK0) | ~lt(sK4(sK0,sK0),sK0) | ~cont(X1,sK0) | ~cont(X1,sK0) {126} (7:16) [resolution 12917,272]
272. sc(X0,X1) | ~lt(sK4(X1,X0),X1) | ~lt(sK4(X1,X0),X0) | ~cont(X4,X1) | ~cont(X4,X0) (0:19) [cnf transformation 230]
230. ! [X0,X1] : ((~sc(X0,X1) | ((cont(sK3(X1,X0),X0) & cont(sK3(X1,X0),X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (lt(X3,X0) & lt(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | (cont(sK4(X1,X0),X0) & cont(sK4(X1,X0),X1) & (~lt(sK4(X1,X0),X0) | ~lt(sK4(X1,X0),X1))) | sc(X0,X1)))[skolemisation 157]
157. ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (lt(X3,X0) & lt(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | ? [X5] : (cont(X5,X0) & cont(X5,X1) & (~lt(X5,X0) | ~lt(X5,X1))) | sc(X0,X1)))[flattening 156]
156. ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((~cont(X3,X0) | ~cont(X3,X1)) | (lt(X3,X0) & lt(X3,X1))))) & ((! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | ? [X5] : ((cont(X5,X0) & cont(X5,X1)) & (~lt(X5,X0) | ~lt(X5,X1)))) | sc(X0,X1)))[ennf transformation 85]
85. ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~(cont(X3,X0) & cont(X3,X1)) | (lt(X3,X0) & lt(X3,X1))))) & (~(? [X4] : (cont(X4,X0) & cont(X4,X1)) & ! [X5] : (~(cont(X5,X0) & cont(X5,X1)) | (lt(X5,X0) & lt(X5,X1)))) | sc(X0,X1)))[flattening 84]
84. ! [X0] : ! [X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~(cont(X3,X0) & cont(X3,X1)) | (lt(X3,X0) & lt(X3,X1))))) & (~(? [X4] : (cont(X4,X0) & cont(X4,X1)) & ! [X5] : (~(cont(X5,X0) & cont(X5,X1)) | (lt(X5,X0) & lt(X5,X1)))) | sc(X0,X1)))[rectify 6]
6. ! [X0] : ! [X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | (lt(X2,X0) & lt(X2,X1))))) & (~(? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | (lt(X2,X0) & lt(X2,X1)))) | sc(X0,X1)))[input]
12917. ~sc(sK0,sK0) {126} (6:3) [resolution 12905,313]
313. lt(intersection(X0,X1),X0) | ~sc(X0,X1) (0:8) [cnf transformation 173]
173. ! [X0,X1] : (~sc(X0,X1) | lt(intersection(X0,X1),X0))[ennf transformation 107]
107. ! [X0,X1] : (~sc(X0,X1) | lt(intersection(X0,X1),X0))[flattening 45]
45. ! [X0] : ! [X1] : (~sc(X0,X1) | lt(intersection(X0,X1),X0))[input]
12905. ~lt(intersection(sK0,sK0),sK0) {126} (5:5) [subsumption resolution 12904,245]
12904. ~lt(intersection(sK0,sK0),sK0) | ~leq(sK0,sK0) {126} (5:8) [subsumption resolution 12899,358]
358. cont(sK0,sK0) (1:3) [resolution 246,242]
246. zex(X0) | cont(X0,X0) (0:5) [cnf transformation 139]
139. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[ennf transformation 19]
19. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[input]
12899. ~lt(intersection(sK0,sK0),sK0) | ~cont(sK0,sK0) | ~leq(sK0,sK0) {126} (5:11) [resolution 12886,339]
339. cont(X2,difference(X0,X1)) | ~lt(intersection(X2,X1),X2) | ~cont(X2,X0) | ~leq(X0,X1) (0:16) [cnf transformation 207]
207. ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X0) | ~lt(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)))[flattening 206]
206. ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X0) | ~lt(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))[ennf transformation 130]
130. ! [X0,X1,X2] : (~(leq(X0,X1) & cont(X2,X0) & lt(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))[flattening 9]
9. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & cont(X2,X0) & lt(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))[input]
12886. ~cont(sK0,difference(sK0,sK0)) {126} (4:5) [resolution 11684,301]
301. zex(difference(X0,X1)) | ~cont(X0,X1) (0:7) [cnf transformation 169]
11684. ~zex(difference(sK0,difference(sK0,sK0))) {126} (3:6) [backtracking splitting component]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1078
Passive clauses: 5993
Generated clauses: 12887
Final active clauses: 1038
Final passive clauses: 4475
Input formulas: 67
Initial clauses: 111

Selected by SInE selection: 67
SInE iterations: 1

Duplicate literals: 125
Fw subsumption resolutions: 522

Simple tautologies: 49
Forward subsumptions: 6154

Binary resolution: 11880
Factoring: 21

Backtracking splits: 150
Backtracking splits refuted: 37
Backtracking splits refuted at zero level: 6

Memory used: 2302KB
Time elapsed: 0.284 s
------------------------------
% Success in time 0.308 s
