Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f933,plain,(
  $false),
  inference(subsumption_resolution,[],[f931,f652])).
fof(f652,plain,(
  p(sK1,sK1)),
  inference(subsumption_resolution,[],[f643,f537])).
fof(f537,plain,(
  ( ! [X0] : (eqdim(X0,X0)) )),
  inference(subsumption_resolution,[],[f525,f263])).
fof(f263,plain,(
  ( ! [X0] : (~less(X0,X0)) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X53] : ~less(X53,X53)),
  file(unknown,unknown)).
fof(f525,plain,(
  ( ! [X0] : (eqdim(X0,X0) | less(X0,X0)) )),
  inference(resolution,[],[f307,f263])).
fof(f307,plain,(
  ( ! [X0,X1] : (less(X1,X0) | eqdim(X0,X1) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X64] : ! [X65] : (eqdim(X64,X65) <=> (~less(X64,X65) & ~less(X65,X64)))),
  file(unknown,unknown)).
fof(f643,plain,(
  ~eqdim(sK1,sK1) | p(sK1,sK1)),
  inference(resolution,[],[f322,f354])).
fof(f354,plain,(
  cont(sK1,sK1)),
  inference(resolution,[],[f261,f244])).
fof(f244,plain,(
  ~zex(sK1)),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ~zex(sK0) & ~zex(sK1) & ~p(sK1,sK0) & eqdim(sK0,sK1) & ~po(sK0,sK1) & ! [X2] : (~p(X2,sK1) | po(X2,sK0))),
  inference(skolemisation,[status(esa)],[f134])).
fof(f134,plain,(
  ? [X0,X1] : (~zex(X0) & ~zex(X1) & ~p(X1,X0) & eqdim(X0,X1) & ~po(X0,X1) & ! [X2] : (~p(X2,X1) | po(X2,X0)))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ? [X0,X1] : ((~zex(X0) & ~zex(X1) & ~p(X1,X0) & eqdim(X0,X1) & ~po(X0,X1)) & ! [X2] : (~p(X2,X1) | po(X2,X0)))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ~! [X0,X1] : ((~zex(X0) & ~zex(X1) & ~p(X1,X0) & eqdim(X0,X1) & ~po(X0,X1)) => ? [X2] : (p(X2,X1) & ~po(X2,X0)))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : ((~zex(X0) & ~zex(X1) & ~p(X1,X0) & eqdim(X0,X1) & ~po(X0,X1)) => ? [X2] : (p(X2,X1) & ~po(X2,X0)))),
  file(unknown,unknown)).
fof(f261,plain,(
  ( ! [X0] : (zex(X0) | cont(X0,X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X104] : (~zex(X104) <=> cont(X104,X104))),
  file(unknown,unknown)).
fof(f322,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X75] : ! [X76] : (p(X75,X76) <=> (cont(X75,X76) & eqdim(X75,X76)))),
  file(unknown,unknown)).
fof(f931,plain,(
  ~p(sK1,sK1)),
  inference(resolution,[],[f925,f248])).
fof(f248,plain,(
  ( ! [X2] : (po(X2,sK0) | ~p(X2,sK1)) )),
  inference(cnf_transformation,[],[f187])).
fof(f925,plain,(
  ~po(sK1,sK0)),
  inference(resolution,[],[f924,f315])).
fof(f315,plain,(
  ( ! [X0,X1] : (p(sK14(X1,X0),X1) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1] : ((~po(X0,X1) | (p(sK14(X1,X0),X0) & p(sK14(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(skolemisation,[status(esa)],[f230])).
fof(f230,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))),
  inference(rectify,[],[f229])).
fof(f229,plain,(
  ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X2] : (~p(X2,X0) | ~p(X2,X1)) | po(X0,X1)))),
  inference(nnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1] : (po(X0,X1) <=> ? [X2] : (p(X2,X0) & p(X2,X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X79] : ! [X80] : (po(X79,X80) <=> ? [X81] : (p(X81,X79) & p(X81,X80)))),
  file(unknown,unknown)).
fof(f924,plain,(
  ~p(sK14(sK0,sK1),sK0)),
  inference(subsumption_resolution,[],[f922,f652])).
fof(f922,plain,(
  ~p(sK14(sK0,sK1),sK0) | ~p(sK1,sK1)),
  inference(resolution,[],[f608,f248])).
fof(f608,plain,(
  ( ! [X3] : (~po(sK1,X3) | ~p(sK14(X3,sK1),sK0)) )),
  inference(resolution,[],[f602,f314])).
fof(f314,plain,(
  ( ! [X0,X1] : (p(sK14(X1,X0),X0) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f602,plain,(
  ( ! [X30] : (~p(X30,sK1) | ~p(X30,sK0)) )),
  inference(resolution,[],[f316,f247])).
fof(f247,plain,(
  ~po(sK0,sK1)),
  inference(cnf_transformation,[],[f187])).
fof(f316,plain,(
  ( ! [X0,X3,X1] : (po(X0,X1) | ~p(X3,X1) | ~p(X3,X0)) )),
  inference(cnf_transformation,[],[f231])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 236
Passive clauses: 390
Generated clauses: 685
Final active clauses: 235
Final passive clauses: 135
Input formulas: 47
Initial clauses: 102

Selected by SInE selection: 47
SInE iterations: 1

Duplicate literals: 8
Fw subsumption resolutions: 26

Simple tautologies: 19
Forward subsumptions: 252

Binary resolution: 539
Factoring: 1

Backtracking splits: 6
Backtracking splits refuted: 1
Backtracking splits refuted at zero level: 1

Memory used: 383KB
Time elapsed: 0.011 s
------------------------------
% Success in time 0.039 s
