% ott+11_2:1_add=large:afp=40000:afq=2.0:amm=sco:anc=none:br=off:cond=on:irw=on:nwc=1:sd=2:ss=axioms:st=2.0:sos=all:urr=on:updr=off_3 on codi_down_sum_theorems_goal06.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.4.0 (commit f9be976 on 2020-01-20 19:56:10 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7803
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% dis+1010_10_afr=on:afp=100000:afq=2.0:amm=sco:anc=none:ep=RS:fde=unused:gs=on:nwc=1:sos=on:sac=on:sp=occurrence_2 on codi_down_sum_theorems_goal06.all
% Refutation not found, incomplete strategy% ------------------------------
% Version: Vampire 4.4.0 (commit f9be976 on 2020-01-20 19:56:10 +0000)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 9978
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
% ott+1010_5:4_av=off:bd=off:fde=none:irw=on:lma=on:nm=32:nwc=2.5:sd=2:ss=axioms:st=3.0:urr=on_5 on codi_down_sum_theorems_goal06.all
% Refutation found. Thanks to Tanya!
% SZS status Theorem for codi_down_sum_theorems_goal06.all
% SZS output start Proof for codi_down_sum_theorems_goal06.all
fof(f5766,plain,(
  $false),
  inference(subsumption_resolution,[],[f5724,f5666])).
fof(f5666,plain,(
  sK0 != sum(sK0,sK0)),
  inference(superposition,[],[f325,f5460])).
fof(f5460,plain,(
  sK0 = sK1),
  inference(unit_resulting_resolution,[],[f324,f5398,f335])).
fof(f335,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : (X0 = X1 | (~zex(X1) | ~zex(X0)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : (X0 = X1 | ~(zex(X1) & zex(X0)))),
  inference(flattening,[],[f83])).
fof(f83,axiom,(
  ! [X0] : ! [X1] : (X0 = X1 | ~(zex(X1) & zex(X0)))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal06.all.tptp',axiom89490)).
fof(f5398,plain,(
  zex(sK0)),
  inference(unit_resulting_resolution,[],[f324,f5381,f356])).
fof(f356,plain,(
  ( ! [X0,X1] : (lt(X0,X1) | zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1] : (lt(X0,X1) | zex(X1) | ~zex(X0))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0,X1] : (lt(X0,X1) | (zex(X1) | ~zex(X0)))),
  inference(ennf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : (lt(X0,X1) | ~(~zex(X1) & zex(X0)))),
  inference(flattening,[],[f84])).
fof(f84,axiom,(
  ! [X0] : ! [X1] : (lt(X0,X1) | ~(~zex(X1) & zex(X0)))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal06.all.tptp',axiom89500)).
fof(f5381,plain,(
  ~lt(sK1,sK0)),
  inference(unit_resulting_resolution,[],[f325,f621])).
fof(f621,plain,(
  ( ! [X2,X3] : (sum(X3,X2) = X3 | ~lt(X2,X3)) )),
  inference(superposition,[],[f340,f345])).
fof(f345,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = sum(X1,X0)) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : sum(X0,X1) = sum(X1,X0)),
  inference(flattening,[],[f3])).
fof(f3,axiom,(
  ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal06.all.tptp',axiom10)).
fof(f340,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = X1 | ~lt(X0,X1)) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : (sum(X0,X1) = X1 | ~lt(X0,X1))),
  inference(flattening,[],[f4])).
fof(f4,axiom,(
  ! [X0] : ! [X1] : (sum(X0,X1) = X1 | ~lt(X0,X1))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal06.all.tptp',axiom20)).
fof(f324,plain,(
  zex(sK1)),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  sK0 != sum(sK0,sK1) & zex(sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f206,f298])).
fof(f298,plain,(
  ? [X0,X1] : (sum(X0,X1) != X0 & zex(X1)) => (sK0 != sum(sK0,sK1) & zex(sK1))),
  introduced(choice_axiom,[])).
fof(f206,plain,(
  ? [X0,X1] : (sum(X0,X1) != X0 & zex(X1))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ~! [X0,X1] : (sum(X0,X1) = X0 | ~zex(X1))),
  inference(flattening,[],[f2])).
fof(f2,negated_conjecture,(
  ~! [X0] : ! [X1] : (sum(X0,X1) = X0 | ~zex(X1))),
  inference(negated_conjecture,[],[f1])).
fof(f1,conjecture,(
  ! [X0] : ! [X1] : (sum(X0,X1) = X0 | ~zex(X1))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal06.all.tptp',conjecture60)).
fof(f325,plain,(
  sK0 != sum(sK0,sK1)),
  inference(cnf_transformation,[],[f299])).
fof(f5724,plain,(
  sK0 = sum(sK0,sK0)),
  inference(superposition,[],[f1968,f5460])).
fof(f1968,plain,(
  sK1 = sum(sK1,sK1)),
  inference(unit_resulting_resolution,[],[f464,f1897,f326])).
fof(f326,plain,(
  ( ! [X0,X1] : (cont(sK2(X0,X1),X1) | X0 = X1 | cont(sK2(X0,X1),X0)) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1] : (X0 = X1 | ((~cont(sK2(X0,X1),X0) & cont(sK2(X0,X1),X1)) | (~cont(sK2(X0,X1),X1) & cont(sK2(X0,X1),X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f207,f300])).
fof(f300,plain,(
  ! [X1,X0] : (? [X2] : ((~cont(X2,X0) & cont(X2,X1)) | (~cont(X2,X1) & cont(X2,X0))) => ((~cont(sK2(X0,X1),X0) & cont(sK2(X0,X1),X1)) | (~cont(sK2(X0,X1),X1) & cont(sK2(X0,X1),X0))))),
  introduced(choice_axiom,[])).
fof(f207,plain,(
  ! [X0,X1] : (X0 = X1 | ? [X2] : ((~cont(X2,X0) & cont(X2,X1)) | (~cont(X2,X1) & cont(X2,X0))))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : (X0 = X1 | ~! [X2] : ((cont(X2,X0) | ~cont(X2,X1)) & (cont(X2,X1) | ~cont(X2,X0))))),
  inference(flattening,[],[f90])).
fof(f90,axiom,(
  ! [X0] : ! [X1] : (X0 = X1 | ~! [X2] : ((cont(X2,X0) | ~cont(X2,X1)) & (cont(X2,X1) | ~cont(X2,X0))))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal06.all.tptp',axiom89560)).
fof(f1897,plain,(
  ( ! [X0] : (~cont(X0,sum(sK1,sK1))) )),
  inference(forward_demodulation,[],[f1896,f555])).
fof(f555,plain,(
  ( ! [X0] : (sK1 = difference(sK1,X0)) )),
  inference(forward_demodulation,[],[f550,f552])).
fof(f552,plain,(
  sK1 = sK5),
  inference(unit_resulting_resolution,[],[f324,f350,f335])).
fof(f350,plain,(
  zex(sK5)),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  zex(sK5)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f107,f306])).
fof(f306,plain,(
  ? [X0] : zex(X0) => zex(sK5)),
  introduced(choice_axiom,[])).
fof(f107,axiom,(
  ? [X0] : zex(X0)),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal06.all.tptp',axiom87750)).
fof(f550,plain,(
  ( ! [X0] : (sK5 = difference(sK1,X0)) )),
  inference(unit_resulting_resolution,[],[f489,f350,f335])).
fof(f489,plain,(
  ( ! [X0] : (zex(difference(sK1,X0))) )),
  inference(unit_resulting_resolution,[],[f324,f379])).
fof(f379,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1] : ((zex(difference(X0,X1)) | (~cont(X0,X1) & ~zex(X0))) & (cont(X0,X1) | zex(X0) | ~zex(difference(X0,X1))))),
  inference(ennf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : ((zex(difference(X0,X1)) | ~(cont(X0,X1) | zex(X0))) & (cont(X0,X1) | zex(X0) | ~zex(difference(X0,X1))))),
  inference(flattening,[],[f12])).
fof(f12,axiom,(
  ! [X0] : ! [X1] : ((zex(difference(X0,X1)) | ~(cont(X0,X1) | zex(X0))) & (cont(X0,X1) | zex(X0) | ~zex(difference(X0,X1))))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal06.all.tptp',axiom15200)).
fof(f1896,plain,(
  ( ! [X0,X1] : (~cont(X0,sum(sK1,difference(sK1,X1)))) )),
  inference(forward_demodulation,[],[f1895,f345])).
fof(f1895,plain,(
  ( ! [X0,X1] : (~cont(X0,sum(difference(sK1,X1),sK1))) )),
  inference(forward_demodulation,[],[f1882,f555])).
fof(f1882,plain,(
  ( ! [X2,X0,X1] : (~cont(X0,sum(difference(sK1,X1),difference(sK1,X2)))) )),
  inference(unit_resulting_resolution,[],[f495,f495,f346])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (~cont(X2,sum(X0,X1)) | cont(X2,X0) | cont(difference(X2,X0),X1)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | (cont(X2,X0) | ~cont(X2,sum(X0,X1))))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  inference(flattening,[],[f6])).
fof(f6,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal06.all.tptp',axiom40)).
fof(f495,plain,(
  ( ! [X0,X1] : (~cont(X0,difference(sK1,X1))) )),
  inference(unit_resulting_resolution,[],[f489,f374])).
fof(f374,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  inference(flattening,[],[f88])).
fof(f88,axiom,(
  ! [X0] : ! [X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal06.all.tptp',axiom89540)).
fof(f464,plain,(
  ( ! [X0] : (~cont(X0,sK1)) )),
  inference(unit_resulting_resolution,[],[f324,f374])).
% SZS output end Proof for codi_down_sum_theorems_goal06.all
% ------------------------------
% Version: Vampire 4.4.0 (commit f9be976 on 2020-01-20 19:56:10 +0000)
% Termination reason: Refutation

% Memory used [KB]: 2046
% Time elapsed: 0.327 s
% ------------------------------
% ------------------------------
% Success in time 0.978 s
