Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal11.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 915
Passive clauses: 5510
Generated clauses: 11514
Final active clauses: 818
Final passive clauses: 3993
Input formulas: 44
Initial clauses: 92

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 168
Fw subsumption resolutions: 419

Simple tautologies: 105
Forward subsumptions: 5011

Binary resolution: 10291
Factoring: 56

Backtracking splits: 270
Backtracking splits refuted: 59
Backtracking splits refuted at zero level: 29

Memory used: 1918KB
Time elapsed: 0.404 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal11.all
Refutation found. Thanks to Tanya!
22897. $false (5:0) [subsumption resolution 22870,443]
443. ~zex(sK0) (2:2) [resolution 438,162]
162. sc(sK0,sK1) (0:3) [cnf transformation 146]
146. sc(sK0,sK1) & ~lt(intersection(sK0,sK1),sK0)[skolemisation 87]
87. ? [X0,X1] : (sc(X0,X1) & ~lt(intersection(X0,X1),X0))[ennf transformation 45]
45. ~! [X0,X1] : (~sc(X0,X1) | lt(intersection(X0,X1),X0))[flattening 1]
1. ~! [X0] : ! [X1] : (~sc(X0,X1) | lt(intersection(X0,X1),X0))[negated conjecture]
438. ~sc(X2,X3) | ~zex(X2) (1:5) [resolution 198,220]
220. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 112]
112. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 67]
67. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[flattening 25]
25. ! [X0] : ! [X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[input]
198. cont(sK8(X1,X0),X0) | ~sc(X0,X1) (0:8) [cnf transformation 156]
156. ! [X0,X1] : ((~sc(X0,X1) | ((cont(sK8(X1,X0),X0) & cont(sK8(X1,X0),X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (lt(X3,X0) & lt(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | (cont(sK9(X1,X0),X0) & cont(sK9(X1,X0),X1) & (~lt(sK9(X1,X0),X0) | ~lt(sK9(X1,X0),X1))) | sc(X0,X1)))[skolemisation 103]
103. ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~cont(X3,X0) | ~cont(X3,X1) | (lt(X3,X0) & lt(X3,X1))))) & (! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | ? [X5] : (cont(X5,X0) & cont(X5,X1) & (~lt(X5,X0) | ~lt(X5,X1))) | sc(X0,X1)))[flattening 102]
102. ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : ((~cont(X3,X0) | ~cont(X3,X1)) | (lt(X3,X0) & lt(X3,X1))))) & ((! [X4] : (~cont(X4,X0) | ~cont(X4,X1)) | ? [X5] : ((cont(X5,X0) & cont(X5,X1)) & (~lt(X5,X0) | ~lt(X5,X1)))) | sc(X0,X1)))[ennf transformation 58]
58. ! [X0,X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~(cont(X3,X0) & cont(X3,X1)) | (lt(X3,X0) & lt(X3,X1))))) & (~(? [X4] : (cont(X4,X0) & cont(X4,X1)) & ! [X5] : (~(cont(X5,X0) & cont(X5,X1)) | (lt(X5,X0) & lt(X5,X1)))) | sc(X0,X1)))[flattening 57]
57. ! [X0] : ! [X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X3] : (~(cont(X3,X0) & cont(X3,X1)) | (lt(X3,X0) & lt(X3,X1))))) & (~(? [X4] : (cont(X4,X0) & cont(X4,X1)) & ! [X5] : (~(cont(X5,X0) & cont(X5,X1)) | (lt(X5,X0) & lt(X5,X1)))) | sc(X0,X1)))[rectify 8]
8. ! [X0] : ! [X1] : ((~sc(X0,X1) | (? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | (lt(X2,X0) & lt(X2,X1))))) & (~(? [X2] : (cont(X2,X0) & cont(X2,X1)) & ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | (lt(X2,X0) & lt(X2,X1)))) | sc(X0,X1)))[input]
22870. zex(sK0) (5:2) [resolution 22736,299]
299. ~zex(intersection(sK0,sK1)) | zex(sK0) (1:6) [resolution 230,163]
163. ~lt(intersection(sK0,sK1),sK0) (0:5) [cnf transformation 146]
230. lt(X0,X1) | zex(X1) | ~zex(X0) (0:7) [cnf transformation 125]
125. ! [X0,X1] : (~zex(X0) | zex(X1) | lt(X0,X1))[flattening 124]
124. ! [X0,X1] : ((~zex(X0) | zex(X1)) | lt(X0,X1))[ennf transformation 76]
76. ! [X0,X1] : (~(zex(X0) & ~zex(X1)) | lt(X0,X1))[flattening 21]
21. ! [X0] : ! [X1] : (~(zex(X0) & ~zex(X1)) | lt(X0,X1))[input]
22736. zex(intersection(sK0,sK1)) (4:4) [subsumption resolution 22668,162]
22668. zex(intersection(sK0,sK1)) | ~sc(sK0,sK1) (4:7) [resolution 7807,1433]
1433. ~cont(intersection(sK0,sK1),X0) | ~sc(sK0,X0) (2:8) [subsumption resolution 1423,221]
221. ~cont(X0,X1) | ~zex(X0) (0:5) [cnf transformation 112]
1423. ~cont(intersection(sK0,sK1),X0) | ~sc(sK0,X0) | zex(intersection(sK0,sK1)) (2:12) [resolution 889,218]
218. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 65]
65. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 3]
3. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
889. ~cont(intersection(sK0,sK1),sK0) | ~cont(intersection(sK0,sK1),X19) | ~sc(sK0,X19) (1:13) [resolution 200,163]
200. lt(X3,X0) | ~cont(X3,X1) | ~cont(X3,X0) | ~sc(X0,X1) (0:12) [cnf transformation 156]
7807. cont(intersection(X0,X1),X1) | zex(intersection(X0,X1)) (3:9) [resolution 2034,166]
166. cont(X0,X0) | zex(X0) (0:5) [cnf transformation 88]
88. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[ennf transformation 22]
22. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[input]
2034. ~cont(intersection(X0,X1),intersection(X0,X1)) | cont(intersection(X0,X1),X1) (2:12) [resolution 541,164]
164. eqdim(X0,X0) (0:3) [cnf transformation 13]
13. ! [X0] : eqdim(X0,X0)[input]
541. ~eqdim(X9,intersection(X10,X11)) | ~cont(X9,intersection(X10,X11)) | cont(X9,X11) (1:13) [resolution 210,238]
238. ~p(X2,intersection(X0,X1)) | cont(X2,X1) (0:8) [cnf transformation 130]
130. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 129]
129. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[ennf transformation 79]
79. ! [X0,X1,X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 5]
5. ! [X0] : ! [X1] : ! [X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[input]
210. p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 106]
106. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 105]
105. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 61]
61. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 29]
29. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1502
Passive clauses: 7712
Generated clauses: 22728
Final active clauses: 1502
Final passive clauses: 5822
Input formulas: 44
Initial clauses: 91

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 823
Fw subsumption resolutions: 1567

Simple tautologies: 262
Equational tautologies: 5
Forward subsumptions: 12761

Binary resolution: 15912
Factoring: 172
Forward superposition: 4146
Backward superposition: 56

Unique components: 2

Memory used: 2686KB
Time elapsed: 1.140 s
------------------------------
% Success in time 1.632 s
