Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal18.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 562
Passive clauses: 1800
Generated clauses: 4890
Final active clauses: 560
Final passive clauses: 1128
Input formulas: 43
Initial clauses: 89

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 40
Fw subsumption resolutions: 34

Simple tautologies: 29
Forward subsumptions: 2925

Binary resolution: 4608
Factoring: 35

Backtracking splits: 71
Backtracking splits refuted: 6

Memory used: 895KB
Time elapsed: 0.405 s
------------------------------
% remaining time: 5992 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal18.all
Refutation found. Thanks to Tanya!
3152. $false (7:0) [resolution 3151,326]
326. c(intersection(sK0,sK0),y) (1:5) [resolution 277,213]
213. ~zex(intersection(intersection(sK0,sK0),y)) (0:6) [cnf transformation 167]
167. zex(intersection(sK0,y)) & ~zex(intersection(intersection(sK0,sK0),y))[skolemisation 119]
119. ? [X0] : (zex(intersection(X0,y)) & ~zex(intersection(intersection(X0,X0),y)))[ennf transformation 1]
1. ~! [X0] : (zex(intersection(X0,y)) => zex(intersection(intersection(X0,X0),y)))[negated conjecture]
277. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 206]
206. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[nnf transformation 149]
149. ! [X0,X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[ennf transformation 99]
99. ! [X0,X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[flattening 98]
98. ! [X0] : ! [X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[rectify 24]
24. ! [X48] : ! [X49] : (~c(X48,X49) <=> zex(intersection(X48,X49)))[input]
3151. ~c(intersection(sK0,X3),y) (6:5) [subsumption resolution 3130,495]
495. ~c(X2,X3) | ~zex(X2) (1:5) [resolution 265,227]
227. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 125]
125. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 53]
53. ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))[flattening 52]
52. ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))[rectify 20]
20. ! [X39] : ! [X40] : (zex(X39) => (~cont(X40,X39) & ~cont(X39,X40)))[input]
265. cont(sK10(X1,X0),X0) | ~c(X0,X1) (0:8) [cnf transformation 199]
199. ! [X0,X1] : ((~c(X0,X1) | (cont(sK10(X1,X0),X0) & cont(sK10(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 198]
198. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[rectify 197]
197. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | c(X0,X1)))[nnf transformation 91]
91. ! [X0,X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))[flattening 90]
90. ! [X0] : ! [X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))[rectify 2]
2. ! [X1] : ! [X2] : (c(X1,X2) <=> ? [X3] : (cont(X3,X1) & cont(X3,X2)))[input]
3130. ~c(intersection(sK0,X3),y) | zex(intersection(sK0,X3)) (6:9) [resolution 3125,239]
239. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 143]
143. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 75]
75. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 74]
74. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 25]
25. ! [X50] : ! [X51] : (~zex(intersection(X50,X51)) => cont(intersection(X50,X51),X50))[input]
3125. ~cont(X0,sK0) | ~c(X0,y) (5:6) [duplicate literal removal 3115]
3115. ~cont(X0,sK0) | ~c(X0,y) | ~c(X0,y) (5:9) [resolution 976,265]
976. ~cont(sK10(y,X18),X19) | ~cont(X19,sK0) | ~c(X18,y) (4:11) [resolution 687,266]
266. cont(sK10(X1,X0),X1) | ~c(X0,X1) (0:8) [cnf transformation 199]
687. ~cont(X9,y) | ~cont(X9,X8) | ~cont(X8,sK0) (3:9) [resolution 281,569]
569. ~cont(X0,sK0) | ~cont(X0,y) (2:6) [resolution 267,327]
327. ~c(sK0,y) (1:3) [resolution 278,212]
212. zex(intersection(sK0,y)) (0:4) [cnf transformation 167]
278. ~zex(intersection(X0,X1)) | ~c(X0,X1) (0:7) [cnf transformation 206]
267. c(X0,X1) | ~cont(X3,X1) | ~cont(X3,X0) (0:9) [cnf transformation 199]
281. cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1) (0:9) [cnf transformation 155]
155. ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))[flattening 154]
154. ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))[ennf transformation 105]
105. ! [X0,X1,X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))[flattening 104]
104. ! [X0] : ! [X1] : ! [X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))[rectify 19]
19. ! [X37] : ! [X38] : ! [X0] : ((cont(X37,X38) & cont(X38,X0)) => cont(X37,X0))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 407
Passive clauses: 1333
Generated clauses: 2942
Final active clauses: 407
Final passive clauses: 870
Input formulas: 43
Initial clauses: 84

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 120
Fw subsumption resolutions: 124

Simple tautologies: 44
Equational tautologies: 3
Forward subsumptions: 1375

Binary resolution: 1431
Factoring: 24
Forward superposition: 1124
Backward superposition: 44

Unique components: 2

Memory used: 639KB
Time elapsed: 0.309 s
------------------------------
% Success in time 1.158 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 23:35:34 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal18.all.tptp"
Input read from 
============================ end of footer ===========================
