Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal10.all
Refutation found. Thanks to Tanya!
4178. $false (13:0) [subsumption resolution 4177,164]
164. lt(sK0,sK1) (0:3) [cnf transformation 147]
147. inc(sK0,sK1) & lt(sK0,sK1) & ~lt(intersection(sK0,sK1),sK1)[skolemisation 88]
88. ? [X0,X1] : (inc(X0,X1) & lt(X0,X1) & ~lt(intersection(X0,X1),X1))[flattening 87]
87. ? [X0,X1] : ((inc(X0,X1) & lt(X0,X1)) & ~lt(intersection(X0,X1),X1))[ennf transformation 45]
45. ~! [X0,X1] : (~(inc(X0,X1) & lt(X0,X1)) | lt(intersection(X0,X1),X1))[flattening 1]
1. ~! [X0] : ! [X1] : (~(inc(X0,X1) & lt(X0,X1)) | lt(intersection(X0,X1),X1))[negated conjecture]
4177. ~lt(sK0,sK1) (13:3) [resolution 4150,217]
217. ~leq(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 111]
111. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 110]
110. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[ennf transformation 63]
63. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[flattening 40]
40. ! [X0] : ! [X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[input]
4150. leq(sK1,sK0) (12:3) [backtracking split refutation 1478,1488,4148]
4148. $false {25} (15:0) [subsumption resolution 4146,165]
165. ~lt(intersection(sK0,sK1),sK1) (0:5) [cnf transformation 147]
4146. lt(intersection(sK0,sK1),sK1) {25} (15:5) [resolution 4090,583]
583. leq(X1,X0) | lt(X0,X1) (0:6) [subsumption resolution 218,219]
219. leq(X1,X0) | leq(X0,X1) (0:6) [cnf transformation 64]
64. ! [X0,X1] : (leq(X0,X1) | leq(X1,X0))[flattening 33]
33. ! [X0] : ! [X1] : (leq(X0,X1) | leq(X1,X0))[input]
218. lt(X0,X1) | leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 111]
4090. ~leq(sK1,intersection(sK0,sK1)) {25} (14:5) [resolution 1490,1665]
1665. leq(intersection(sK0,sK1),sK0) (6:5) [backtracking split refutation 1507,1517,1663]
1663. $false {26} (7:0) [subsumption resolution 1644,644]
644. cont(intersection(sK0,sK1),sK0) (5:5) [resolution 220,311]
311. ~zex(intersection(sK0,sK1)) (4:4) [resolution 305,165]
305. lt(X2,sK1) | ~zex(X2) (3:5) [resolution 232,280]
280. ~zex(sK1) (2:2) [resolution 278,164]
278. ~lt(X3,X4) | ~zex(X4) (1:5) [resolution 217,221]
221. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 112]
112. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 66]
66. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[flattening 36]
36. ! [X0] : ! [X1] : (~zex(X0) | leq(X0,X1))[input]
232. zex(X1) | lt(X0,X1) | ~zex(X0) (0:7) [cnf transformation 126]
126. ! [X0,X1] : (~zex(X0) | zex(X1) | lt(X0,X1))[flattening 125]
125. ! [X0,X1] : ((~zex(X0) | zex(X1)) | lt(X0,X1))[ennf transformation 76]
76. ! [X0,X1] : (~(zex(X0) & ~zex(X1)) | lt(X0,X1))[flattening 21]
21. ! [X0] : ! [X1] : (~(zex(X0) & ~zex(X1)) | lt(X0,X1))[input]
220. zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0) (0:9) [cnf transformation 65]
65. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 3]
3. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
1644. ~cont(intersection(sK0,sK1),sK0) {26} (7:5) [resolution 680,1517]
680. leq(X6,X7) | ~cont(X6,X7) (1:6) [subsumption resolution 678,216]
216. leq(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 111]
678. lt(X6,X7) | ~cont(X6,X7) | leq(X6,X7) (1:9) [resolution 227,213]
213. ~eqdim(X0,X1) | leq(X0,X1) (0:6) [cnf transformation 109]
109. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 108]
108. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[ennf transformation 62]
62. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[flattening 43]
43. ! [X0] : ! [X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[input]
227. eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 117]
117. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[ennf transformation 71]
71. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 30]
30. ! [X0] : ! [X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[input]
1517. ~leq(intersection(sK0,sK1),sK0) {26} (6:5) [backtracking splitting component]
1507. ~lt(sK2(sK0),sK1) | ~leq(intersection(sK0,sK1),sK0) (6:9) [resolution 857,817]
817. lt(X2,sK2(sK0)) | ~leq(X2,sK0) (5:7) [resolution 805,583]
805. ~leq(sK2(sK0),X3) | ~leq(X3,sK0) (4:7) [resolution 245,378]
378. ~leq(sK2(sK0),sK0) (3:4) [backtracking split refutation 354,355,377]
377. $false {1} (6:0) [subsumption resolution 376,355]
376. ~zex(sK0) (6:2) [resolution 375,315]
315. c(sK0,sK1) (5:3) [resolution 311,178]
178. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 95]
95. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[ennf transformation 48]
48. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[flattening 2]
2. ! [X0] : ! [X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[input]
375. ~c(X2,X3) | ~zex(X2) (1:5) [resolution 207,222]
222. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 113]
113. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 67]
67. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[flattening 25]
25. ! [X0] : ! [X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[input]
207. cont(sK10(X1,X0),X0) | ~c(X0,X1) (0:8) [cnf transformation 158]
158. ! [X0,X1] : ((~c(X0,X1) | (cont(sK10(X1,X0),X0) & cont(sK10(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 105]
105. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 60]
60. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 59]
59. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 28]
28. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
355. zex(sK0) {1} (3:2) [backtracking splitting component]
354. ~leq(sK2(sK0),sK0) | zex(sK0) (3:6) [resolution 172,285]
285. ~maxdim(sK0) (2:2) [resolution 279,164]
279. ~lt(X5,X6) | ~maxdim(X5) (1:5) [resolution 217,171]
171. leq(X1,X0) | ~maxdim(X0) (0:5) [cnf transformation 148]
148. ! [X0] : ((~maxdim(X0) | (~zex(X0) & ! [X1] : leq(X1,X0))) & (zex(X0) | ~leq(sK2(X0),X0) | maxdim(X0)))[skolemisation 91]
91. ! [X0] : ((~maxdim(X0) | (~zex(X0) & ! [X1] : leq(X1,X0))) & (zex(X0) | ? [X2] : ~leq(X2,X0) | maxdim(X0)))[flattening 90]
90. ! [X0] : ((~maxdim(X0) | (~zex(X0) & ! [X1] : leq(X1,X0))) & ((zex(X0) | ? [X2] : ~leq(X2,X0)) | maxdim(X0)))[ennf transformation 46]
46. ! [X0] : ((~maxdim(X0) | (~zex(X0) & ! [X1] : leq(X1,X0))) & (~(~zex(X0) & ! [X2] : leq(X2,X0)) | maxdim(X0)))[rectify 37]
37. ! [X0] : ((~maxdim(X0) | (~zex(X0) & ! [X1] : leq(X1,X0))) & (~(~zex(X0) & ! [X1] : leq(X1,X0)) | maxdim(X0)))[input]
172. maxdim(X0) | ~leq(sK2(X0),X0) | zex(X0) (0:8) [cnf transformation 148]
245. leq(X0,X2) | ~leq(X1,X2) | ~leq(X0,X1) (0:9) [cnf transformation 137]
137. ! [X0,X1,X2] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2))[flattening 136]
136. ! [X0,X1,X2] : ((~leq(X0,X1) | ~leq(X1,X2)) | leq(X0,X2))[ennf transformation 83]
83. ! [X0,X1,X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[flattening 34]
34. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[input]
857. ~lt(intersection(sK0,sK1),X16) | ~lt(X16,sK1) (1:8) [resolution 246,165]
246. lt(X0,X2) | ~lt(X1,X2) | ~lt(X0,X1) (0:9) [cnf transformation 139]
139. ! [X0,X1,X2] : (~lt(X0,X1) | ~lt(X1,X2) | lt(X0,X2))[flattening 138]
138. ! [X0,X1,X2] : ((~lt(X0,X1) | ~lt(X1,X2)) | lt(X0,X2))[ennf transformation 84]
84. ! [X0,X1,X2] : (~(lt(X0,X1) & lt(X1,X2)) | lt(X0,X2))[flattening 12]
12. ! [X0] : ! [X1] : ! [X2] : (~(lt(X0,X1) & lt(X1,X2)) | lt(X0,X2))[input]
1490. ~leq(X0,sK0) | ~leq(sK1,X0) {25} (13:6) [resolution 1488,245]
1488. ~leq(sK1,sK0) {25} (12:3) [backtracking splitting component]
1478. ~leq(sK2(sK0),sK14) | ~leq(sK1,sK0) (12:7) [resolution 1430,818]
818. leq(X3,sK2(sK0)) | ~leq(X3,sK0) (5:7) [resolution 805,219]
1430. ~leq(X0,sK14) | ~leq(sK1,X0) (11:6) [resolution 1428,245]
1428. ~leq(sK1,sK14) (10:3) [subsumption resolution 1427,434]
434. ~zex(sK3(sK1)) (7:3) [subsumption resolution 432,280]
432. ~zex(sK3(sK1)) | zex(sK1) (7:5) [resolution 427,175]
175. mindim(X0) | ~zex(sK3(X0)) | zex(X0) (0:7) [cnf transformation 149]
149. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & (zex(X0) | (~zex(sK3(X0)) & ~leq(X0,sK3(X0))) | mindim(X0)))[skolemisation 93]
93. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & (zex(X0) | ? [X2] : (~zex(X2) & ~leq(X0,X2)) | mindim(X0)))[flattening 92]
92. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & ((zex(X0) | ? [X2] : (~zex(X2) & ~leq(X0,X2))) | mindim(X0)))[ennf transformation 47]
47. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & (~(~zex(X0) & ! [X2] : (zex(X2) | leq(X0,X2))) | mindim(X0)))[rectify 38]
38. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & (~(~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1))) | mindim(X0)))[input]
427. ~mindim(sK1) (6:2) [resolution 410,164]
410. ~lt(sK0,X0) | ~mindim(X0) (5:5) [resolution 381,217]
381. leq(X1,sK0) | ~mindim(X1) (4:5) [resolution 379,174]
174. zex(X1) | leq(X0,X1) | ~mindim(X0) (0:7) [cnf transformation 149]
379. ~zex(sK0) (3:2) [backtracking split refutation 354,355,377]
1427. ~leq(sK1,sK14) | zex(sK3(sK1)) (10:6) [resolution 1256,252]
252. ~lt(X1,sK14) | zex(X1) (0:5) [cnf transformation 162]
162. ~zex(sK14) & ! [X1] : (~lt(X1,sK14) | zex(X1))[skolemisation 144]
144. ? [X0] : (~zex(X0) & ! [X1] : (~lt(X1,X0) | zex(X1)))[ennf transformation 19]
19. ? [X0] : (~zex(X0) & ! [X1] : (~lt(X1,X0) | zex(X1)))[input]
1256. lt(sK3(sK1),X3) | ~leq(sK1,X3) (9:7) [resolution 809,583]
809. ~leq(X7,sK3(sK1)) | ~leq(sK1,X7) (8:7) [resolution 245,433]
433. ~leq(sK1,sK3(sK1)) (7:4) [subsumption resolution 431,280]
431. ~leq(sK1,sK3(sK1)) | zex(sK1) (7:6) [resolution 427,176]
176. mindim(X0) | ~leq(X0,sK3(X0)) | zex(X0) (0:8) [cnf transformation 149]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 553
Passive clauses: 1764
Generated clauses: 4031
Final active clauses: 542
Final passive clauses: 1096
Input formulas: 44
Initial clauses: 93

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 38
Fw subsumption resolutions: 76

Simple tautologies: 25
Forward subsumptions: 2155

Binary resolution: 3709
Factoring: 10

Backtracking splits: 52
Backtracking splits refuted: 15
Backtracking splits refuted at zero level: 15

Memory used: 639KB
Time elapsed: 0.123 s
------------------------------
% Success in time 0.167 s
