Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_down_theoremsT6-T10_goal10.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1332
Passive clauses: 8651
Generated clauses: 17083
Final active clauses: 1305
Final passive clauses: 6861
Input formulas: 67
Initial clauses: 109

Selected by SInE selection: 67
SInE iterations: 1
Splitted inequalities: 1

Duplicate literals: 127
Fw subsumption resolutions: 668

Simple tautologies: 73
Forward subsumptions: 7676

Binary resolution: 16098
Factoring: 26

Backtracking splits: 139
Backtracking splits refuted: 15
Backtracking splits refuted at zero level: 6

Memory used: 3198KB
Time elapsed: 0.406 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_down_theoremsT6-T10_goal10.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1610
Passive clauses: 15989
Generated clauses: 32434
Final active clauses: 1610
Final passive clauses: 14118
Input formulas: 67
Initial clauses: 108

Selected by SInE selection: 67
SInE iterations: 1
Splitted inequalities: 1

Duplicate literals: 884
Fw subsumption resolutions: 2345

Simple tautologies: 465
Equational tautologies: 12
Forward subsumptions: 13001

Binary resolution: 22929
Factoring: 122
Forward superposition: 6032
Backward superposition: 44

Unique components: 3

Memory used: 5628KB
Time elapsed: 1.402 s
------------------------------
% remaining time: 5981 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on theorems\conversions\codi_down_theoremsT6-T10_goal10.all
Refutation found. Thanks to Tanya!
12533. $false (11:0) [subsumption resolution 12103,12531]
12531. difference(sK0,sK1) = sK0 (13:5) [resolution 12501,312]
312. ~lt(X1,X0) | difference(X0,X1) = X0 (0:8) [cnf transformation 175]
175. ! [X0,X1] : (~lt(X1,X0) | difference(X0,X1) = X0)[ennf transformation 110]
110. ! [X0,X1] : (~lt(X1,X0) | difference(X0,X1) = X0)[flattening 8]
8. ! [X0] : ! [X1] : (~lt(X1,X0) | difference(X0,X1) = X0)[input]
12501. lt(sK1,sK0) (12:3) [backtracking split refutation 12494,12499,12500]
12500. $false {189} (12:0) [subsumption resolution 12499,12495]
12495. leq(sK1,sK0) (12:3) [resolution 12174,298]
298. leq(X1,X0) | leq(X0,X1) (0:6) [cnf transformation 97]
97. ! [X0,X1] : (leq(X0,X1) | leq(X1,X0))[flattening 27]
27. ! [X0] : ! [X1] : (leq(X0,X1) | leq(X1,X0))[input]
12174. ~leq(sK0,sK1) (11:3) [backtracking split refutation 12102,12168,12169]
12169. $false {179} (4:0) [subsumption resolution 366,12168]
366. cont(sK0,sK0) (4:3) [resolution 360,242]
242. zex(X0) | cont(X0,X0) (0:5) [cnf transformation 138]
138. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[ennf transformation 19]
19. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[input]
360. ~zex(sK0) (3:2) [resolution 351,305]
305. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 169]
169. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 104]
104. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[flattening 30]
30. ! [X0] : ! [X1] : (~zex(X0) | leq(X0,X1))[input]
351. ~leq(sK0,intersection(sK0,sK1)) (2:5) [resolution 346,290]
290. ~lt(X0,X1) | ~leq(X1,X0) (0:6) [cnf transformation 165]
165. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 164]
164. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[ennf transformation 94]
94. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[flattening 32]
32. ! [X0] : ! [X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[input]
346. lt(intersection(sK0,sK1),sK0) (1:5) [resolution 239,309]
309. ~sc(X0,X1) | lt(intersection(X0,X1),X0) (0:8) [cnf transformation 172]
172. ! [X0,X1] : (~sc(X0,X1) | lt(intersection(X0,X1),X0))[ennf transformation 107]
107. ! [X0,X1] : (~sc(X0,X1) | lt(intersection(X0,X1),X0))[flattening 45]
45. ! [X0] : ! [X1] : (~sc(X0,X1) | lt(intersection(X0,X1),X0))[input]
239. sc(sK0,sK1) (0:3) [cnf transformation 221]
221. sc(sK0,sK1) & difference(sK0,sK1) != sK0[skolemisation 137]
137. ? [X0,X1] : (sc(X0,X1) & difference(X0,X1) != X0)[ennf transformation 68]
68. ~! [X0,X1] : (~sc(X0,X1) | difference(X0,X1) = X0)[flattening 1]
1. ~! [X0] : ! [X1] : (~sc(X0,X1) | difference(X0,X1) = X0)[negated conjecture]
12168. ~cont(sK0,sK0) {179} (11:3) [backtracking splitting component]
12102. ~leq(sK0,sK1) | ~cont(sK0,sK0) (11:6) [backtracking split refutation 11337,11361,12098]
12098. $false {148} (0:0) [subsumption resolution 345,12001]
12001. ~sP1(sK0) {148} (12:2) [superposition 344,11361]
344. ~sP1(difference(sK0,sK1)) (0:4) [inequality splitting name introduction]
345. sP1(sK0) (0:2) [inequality splitting 240,344]
240. difference(sK0,sK1) != sK0 (0:5) [cnf transformation 221]
11361. difference(sK0,sK1) = sK0 {148} (11:5) [backtracking splitting component]
11337. ~leq(sK0,sK1) | ~cont(sK0,sK0) | difference(sK0,sK1) = sK0 (11:11) [resolution 11326,969]
969. ~cont(difference(X1,sK1),sK0) | ~leq(X1,sK1) | ~cont(sK0,X1) | difference(X1,sK1) = sK0 (3:16) [resolution 349,321]
321. ~cont(X0,X1) | ~cont(X1,X0) | X0 = X1 (0:9) [cnf transformation 190]
190. ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1)[flattening 189]
189. ! [X0,X1] : ((~cont(X0,X1) | ~cont(X1,X0)) | X0 = X1)[ennf transformation 119]
119. ! [X0,X1] : (~(cont(X0,X1) & cont(X1,X0)) | X0 = X1)[flattening 20]
20. ! [X0] : ! [X1] : (~(cont(X0,X1) & cont(X1,X0)) | X0 = X1)[input]
349. cont(sK0,difference(X0,sK1)) | ~cont(sK0,X0) | ~leq(X0,sK1) (2:11) [resolution 346,335]
335. ~lt(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)) | ~cont(X2,X0) | ~leq(X0,X1) (0:16) [cnf transformation 206]
206. ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X0) | ~lt(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)))[flattening 205]
205. ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X0) | ~lt(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))[ennf transformation 130]
130. ! [X0,X1,X2] : (~(leq(X0,X1) & cont(X2,X0) & lt(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))[flattening 9]
9. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & cont(X2,X0) & lt(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))[input]
11326. cont(difference(sK0,sK1),sK0) (10:5) [backtracking split refutation 11201,11202,11318]
11318. $false {144} (6:0) [subsumption resolution 575,11260]
11260. ~p(difference(sK0,sK1),sK0) {144} (11:5) [resolution 292,11202]
292. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 167]
167. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 166]
166. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 95]
95. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 3]
3. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
575. p(difference(sK0,sK1),sK0) (6:5) [resolution 568,300]
300. zex(difference(X0,X1)) | p(difference(X0,X1),X0) (0:9) [cnf transformation 99]
99. ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))[flattening 60]
60. ! [X0] : ! [X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))[input]
568. ~zex(difference(sK0,sK1)) (5:4) [resolution 363,368]
368. ~cont(sK0,sK1) (3:3) [backtracking split refutation 352,362,367]
367. $false {1} (3:0) [subsumption resolution 362,366]
362. ~cont(sK0,sK0) {1} (3:3) [backtracking splitting component]
352. ~cont(sK0,sK1) | ~cont(sK0,sK0) (3:6) [resolution 351,341]
341. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 218]
218. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 217]
217. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 136]
136. ! [X0,X1,X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[flattening 15]
15. ! [X0] : ! [X1] : ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[input]
363. cont(sK0,X0) | ~zex(difference(sK0,X0)) (4:7) [resolution 360,295]
295. zex(X0) | cont(X0,X1) | ~zex(difference(X0,X1)) (0:9) [cnf transformation 168]
168. ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))[ennf transformation 96]
96. ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & (~(zex(X0) | cont(X0,X1)) | zex(difference(X0,X1))))[flattening 12]
12. ! [X0] : ! [X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & (~(zex(X0) | cont(X0,X1)) | zex(difference(X0,X1))))[input]
11202. ~cont(difference(sK0,sK1),sK0) {144} (10:5) [backtracking splitting component]
11201. eqdim(difference(sK0,sK1),sK0) | ~cont(difference(sK0,sK1),sK0) {8} (10:10) [resolution 11192,313]
313. lt(X0,X1) | eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 176]
176. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[ennf transformation 111]
111. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 17]
17. ! [X0] : ! [X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[input]
11192. ~lt(difference(sK0,sK1),sK0) {8} (9:5) [resolution 10041,575]
10041. ~p(X0,sK0) | ~lt(X0,sK0) {8} (8:6) [subsumption resolution 523,292]
523. ~p(X0,sK0) | ~cont(X0,sK0) | ~lt(X0,sK0) {8} (8:9) [resolution 510,268]
268. inc(X0,X1) | ~p(X3,X1) | ~cont(X3,X0) | ~lt(X3,X0) (0:12) [cnf transformation 233]
233. ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | (lt(sK4(X1,X0),X1) & cont(sK4(X1,X0),X1) & p(sK4(X1,X0),X0))) & ((! [X3] : (~lt(X3,X0) | ~cont(X3,X0) | ~p(X3,X1)) & ! [X4] : (~lt(X4,X1) | ~cont(X4,X1) | ~p(X4,X0))) | inc(X0,X1)))[skolemisation 232]
232. ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | ? [X2] : (lt(X2,X1) & cont(X2,X1) & p(X2,X0))) & ((! [X3] : (~lt(X3,X0) | ~cont(X3,X0) | ~p(X3,X1)) & ! [X4] : (~lt(X4,X1) | ~cont(X4,X1) | ~p(X4,X0))) | inc(X0,X1)))[rectify 220]
220. ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | ? [X3] : (lt(X3,X1) & cont(X3,X1) & p(X3,X0))) & ((! [X4] : (~lt(X4,X0) | ~cont(X4,X0) | ~p(X4,X1)) & ! [X5] : (~lt(X5,X1) | ~cont(X5,X1) | ~p(X5,X0))) | inc(X0,X1)))[definition folding 219]
219. ! [X1,X0] : (~sP0(X1,X0) | ? [X2] : (lt(X2,X0) & cont(X2,X0) & p(X2,X1)))[predicate definition introduction]
510. ~inc(sK0,sK0) {8} (7:3) [backtracking splitting component]
12499. ~leq(sK1,sK0) {189} (12:3) [backtracking splitting component]
12494. lt(sK1,sK0) | ~leq(sK1,sK0) (12:6) [resolution 12174,291]
291. leq(X1,X0) | lt(X0,X1) | ~leq(X0,X1) (0:9) [cnf transformation 165]
12103. difference(sK0,sK1) != sK0 (11:5) [backtracking split refutation 11337,11361,12098]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1337
Passive clauses: 4166
Generated clauses: 13855
Final active clauses: 494
Final passive clauses: 929
Input formulas: 67
Initial clauses: 106

Splitted inequalities: 1

Duplicate literals: 127
Trivial inequalities: 19
Bw subsumption resolutions: 914
Fw demodulations: 1965
Condensations: 32

Simple tautologies: 112
Equational tautologies: 69
Forward subsumptions: 6886

Binary resolution: 5330
Forward superposition: 938
Backward superposition: 2516

Backtracking splits: 189
Backtracking splits refuted: 87
Backtracking splits refuted at zero level: 11

Memory used: 2942KB
Time elapsed: 0.421 s
------------------------------
% Success in time 2.292 s
