% Starting sequential portfolio solving with schedule "casc"
% spawned child 13440 with time: 4 (total remaining time 2999)
% WARNING: value z3 for option sas not known
% dis+1010_4_afp=10000:afq=1.2:anc=none:irw=on:lma=on:nm=64:nwc=10:sas=z3:sac=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal18.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6908
% Time elapsed: 0.500 s
% spawned child 13442 with time: 10 (total remaining time 2994)
% ott+11_8_amm=off:anc=none:bsr=on:cond=on:irw=on:nm=2:nwc=1.1:ss=axioms:st=2.0:sac=on_8 on codi_down_sum_theorems_goal18.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8059
% Time elapsed: 1.100 s
% ------------------------------
% ------------------------------
% spawned child 13444 with time: 4 (total remaining time 2983)
% dis+1011_10_aac=none:add=large:afp=10000:afq=1.1:amm=off:anc=none:fsr=off:gs=on:gsem=off:irw=on:nm=6:newcnf=on:nwc=2.5:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal18.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit

% Memory used [KB]: 11897
% Time elapsed: 0.518 s
% ------------------------------
% ------------------------------
% spawned child 13448 with time: 6 (total remaining time 2977)
% dis+11_28_av=off:fsr=off:irw=on:lcm=predicate:nm=2:newcnf=on:nwc=5:sp=occurrence:urr=on:updr=off_4 on codi_down_sum_theorems_goal18.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 3965
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% spawned child 13450 with time: 6 (total remaining time 2970)
% lrs+1011_5_add=large:afp=1000:afq=1.2:anc=none:fsr=off:irw=on:lma=on:nm=64:newcnf=on:nwc=1:stl=30:sd=3:ss=axioms:st=2.0:sos=on:sac=on:sp=reverse_arity:urr=on:updr=off_4 on codi_down_sum_theorems_goal18.all
% Refutation found. Thanks to Tanya!
% SZS status Theorem for codi_down_sum_theorems_goal18.all
% SZS output start Proof for codi_down_sum_theorems_goal18.all
fof(f1,conjecture,(
  ! [X0] : ! [X1] : (! [X2] : (cont(X0,X2) | ~cont(sum(X0,X1),X2)) | ~(~zex(X1) & ~zex(X0) & eqdim(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal18.all.tptp',conjecture180)).
fof(f2,negated_conjecture,(
  ~! [X0] : ! [X1] : (! [X2] : (cont(X0,X2) | ~cont(sum(X0,X1),X2)) | ~(~zex(X1) & ~zex(X0) & eqdim(X0,X1)))),
  inference(negated_conjecture,[],[f1])).
fof(f5,axiom,(
  ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal18.all.tptp',axiom10)).
fof(f7,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (cont(X2,sum(X0,X1)) | ~(cont(X2,X1) & leq(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal18.all.tptp',axiom30)).
fof(f87,axiom,(
  ! [X0] : ((~zex(X0) | ~cont(X0,X0)) & (cont(X0,X0) | zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal18.all.tptp',axiom89510)).
fof(f89,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (cont(X0,X2) | ~(cont(X1,X2) & cont(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal18.all.tptp',axiom89530)).
fof(f108,axiom,(
  ! [X0] : ! [X1] : ((eqdim(X0,X1) | ~(leq(X1,X0) & leq(X0,X1))) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal18.all.tptp',axiom89310)).
fof(f120,plain,(
  ~! [X0,X1] : (! [X2] : (cont(X0,X2) | ~cont(sum(X0,X1),X2)) | ~(~zex(X1) & ~zex(X0) & eqdim(X0,X1)))),
  inference(flattening,[],[f2])).
fof(f122,plain,(
  ! [X0,X1,X2] : (cont(X0,X2) | ~(cont(X1,X2) & cont(X0,X1)))),
  inference(flattening,[],[f89])).
fof(f140,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | ~(cont(X2,X1) & leq(X0,X1)))),
  inference(flattening,[],[f7])).
fof(f144,plain,(
  ! [X0,X1] : sum(X0,X1) = sum(X1,X0)),
  inference(flattening,[],[f5])).
fof(f180,plain,(
  ! [X0,X1] : ((eqdim(X0,X1) | ~(leq(X1,X0) & leq(X0,X1))) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1)))),
  inference(flattening,[],[f108])).
fof(f216,plain,(
  ? [X0,X1] : (? [X2] : (~cont(X0,X2) & cont(sum(X0,X1),X2)) & (~zex(X1) & ~zex(X0) & eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f120])).
fof(f217,plain,(
  ? [X0,X1] : (? [X2] : (~cont(X0,X2) & cont(sum(X0,X1),X2)) & ~zex(X1) & ~zex(X0) & eqdim(X0,X1))),
  inference(flattening,[],[f216])).
fof(f220,plain,(
  ! [X0,X1,X2] : (cont(X0,X2) | (~cont(X1,X2) | ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f221,plain,(
  ! [X0,X1,X2] : (cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1))),
  inference(flattening,[],[f220])).
fof(f245,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | (~cont(X2,X1) | ~leq(X0,X1)))),
  inference(ennf_transformation,[],[f140])).
fof(f246,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | ~cont(X2,X1) | ~leq(X0,X1))),
  inference(flattening,[],[f245])).
fof(f270,plain,(
  ! [X0,X1] : ((eqdim(X0,X1) | (~leq(X1,X0) | ~leq(X0,X1))) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f180])).
fof(f271,plain,(
  ! [X0,X1] : ((eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1)) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1)))),
  inference(flattening,[],[f270])).
fof(f310,plain,(
  cont(sum(sK0,sK1),sK2)),
  inference(cnf_transformation,[],[f217])).
fof(f311,plain,(
  ~cont(sK0,sK2)),
  inference(cnf_transformation,[],[f217])).
fof(f312,plain,(
  eqdim(sK0,sK1)),
  inference(cnf_transformation,[],[f217])).
fof(f313,plain,(
  ~zex(sK0)),
  inference(cnf_transformation,[],[f217])).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f221])).
fof(f346,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f87])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (~leq(X0,X1) | ~cont(X2,X1) | cont(X2,sum(X0,X1))) )),
  inference(cnf_transformation,[],[f246])).
fof(f353,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = sum(X1,X0)) )),
  inference(cnf_transformation,[],[f144])).
fof(f392,plain,(
  ( ! [X0,X1] : (leq(X1,X0) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f271])).
fof(f447,plain,(
  cont(sK0,sK0)),
  inference(unit_resulting_resolution,[],[f313,f346])).
fof(f497,plain,(
  leq(sK1,sK0)),
  inference(unit_resulting_resolution,[],[f312,f392])).
fof(f530,plain,(
  ~cont(sK0,sum(sK0,sK1))),
  inference(unit_resulting_resolution,[],[f311,f310,f316])).
fof(f679,plain,(
  cont(sK0,sum(sK1,sK0))),
  inference(unit_resulting_resolution,[],[f447,f497,f348])).
fof(f691,plain,(
  cont(sK0,sum(sK0,sK1))),
  inference(forward_demodulation,[],[f679,f353])).
fof(f3002,plain,(
  $false),
  inference(resolution,[],[f691,f530])).
% SZS output end Proof for codi_down_sum_theorems_goal18.all
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 6268
% Time elapsed: 0.276 s
% ------------------------------
% ------------------------------
% Success in time 3.305 s
