Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal09.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 494
Passive clauses: 1682
Generated clauses: 3256
Final active clauses: 485
Final passive clauses: 1076
Input formulas: 43
Initial clauses: 90

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 30
Fw subsumption resolutions: 89

Simple tautologies: 40
Forward subsumptions: 1441

Binary resolution: 2986
Factoring: 18

Backtracking splits: 57
Backtracking splits refuted: 13

Memory used: 767KB
Time elapsed: 0.405 s
------------------------------
% remaining time: 5993 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal09.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 916
Passive clauses: 3975
Generated clauses: 9627
Final active clauses: 916
Final passive clauses: 2822
Input formulas: 43
Initial clauses: 85

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 191
Fw subsumption resolutions: 478

Simple tautologies: 137
Equational tautologies: 5
Forward subsumptions: 5091

Binary resolution: 6317
Factoring: 99
Forward superposition: 2414
Backward superposition: 62

Unique components: 3

Memory used: 1535KB
Time elapsed: 1.415 s
------------------------------
% remaining time: 5972 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal09.all
Refutation found. Thanks to Tanya!
12614. $false (5:0) [subsumption resolution 12567,12607]
12607. ~lt(sK0,sK10(sK0,sK1)) (3:5) [resolution 12569,236]
236. ~lt(X0,X1) | ~lt(X1,X0) (0:6) [cnf transformation 134]
134. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[ennf transformation 64]
64. ! [X0,X1] : (lt(X0,X1) => ~lt(X1,X0))[flattening 63]
63. ! [X0] : ! [X1] : (lt(X0,X1) => ~lt(X1,X0))[rectify 3]
3. ! [X3] : ! [X4] : (lt(X3,X4) => ~lt(X4,X3))[input]
12569. lt(sK10(sK0,sK1),sK0) (2:5) [resolution 12553,265]
265. ~sP0(X0,X1) | lt(sK10(X1,X0),X1) (0:8) [cnf transformation 199]
199. ! [X0,X1] : ((~sP0(X0,X1) | (lt(sK10(X1,X0),X1) & cont(sK10(X1,X0),X1) & p(sK10(X1,X0),X0))) & (! [X3] : (~lt(X3,X1) | ~cont(X3,X1) | ~p(X3,X0)) | sP0(X0,X1)))[skolemisation 198]
198. ! [X0,X1] : ((~sP0(X0,X1) | ? [X2] : (lt(X2,X1) & cont(X2,X1) & p(X2,X0))) & (! [X3] : (~lt(X3,X1) | ~cont(X3,X1) | ~p(X3,X0)) | sP0(X0,X1)))[rectify 197]
197. ! [X1,X0] : ((~sP0(X1,X0) | ? [X2] : (lt(X2,X0) & cont(X2,X0) & p(X2,X1))) & (! [X2] : (~lt(X2,X0) | ~cont(X2,X0) | ~p(X2,X1)) | sP0(X1,X0)))[nnf transformation 167]
167. ! [X1,X0] : (sP0(X1,X0) <=> ? [X2] : (lt(X2,X0) & cont(X2,X0) & p(X2,X1)))[predicate definition introduction]
12553. sP0(sK1,sK0) (1:3) [backtracking split refutation 298,301,12550]
12550. $false {2} (5:0) [subsumption resolution 12199,12543]
12543. ~lt(sK0,sK10(sK0,sK1)) {2} (3:5) [resolution 12201,236]
12201. lt(sK10(sK0,sK1),sK0) {2} (2:5) [resolution 12150,265]
12150. sP0(sK1,sK0) {2} (1:3) [backtracking split refutation 299,302,12147]
12147. $false {2, 3} (1:0) [subsumption resolution 302,11650]
11650. ~p(sK11(sK1,sK0),sK0) {2, 3} (4:5) [resolution 11122,263]
263. eqdim(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 196]
196. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 195]
195. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[nnf transformation 92]
92. ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[flattening 91]
91. ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[rectify 19]
19. ! [X36] : ! [X37] : (p(X36,X37) <=> (cont(X36,X37) & eqdim(X36,X37)))[input]
11122. ~eqdim(sK11(sK1,sK0),sK0) {2, 3} (3:5) [backtracking split refutation 453,459,11118]
11118. $false {2, 3, 14} (5:0) [subsumption resolution 11049,11114]
11114. lt(sK11(sK1,sK0),intersection(sK0,sK1)) {2, 3, 14} (5:7) [backtracking split refutation 11030,11034,11113]
11113. $false {2, 3, 302} (1:0) [subsumption resolution 301,11111]
11111. ~cont(sK11(sK1,sK0),sK1) {3, 302} (6:5) [backtracking split refutation 11102,11109,11110]
11110. $false {3, 306} (2:0) [subsumption resolution 399,11109]
399. cont(sK11(sK1,sK0),sK0) {3} (2:5) [resolution 302,262]
262. ~p(X0,X1) | cont(X0,X1) (0:6) [cnf transformation 196]
11109. ~cont(sK11(sK1,sK0),sK0) {306} (6:5) [backtracking splitting component]
11102. ~cont(sK11(sK1,sK0),sK1) | ~cont(sK11(sK1,sK0),sK0) {302} (6:10) [resolution 11034,285]
285. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 159]
159. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 158]
158. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 108]
108. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[flattening 107]
107. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[rectify 31]
31. ! [X61] : ! [X62] : ! [X63] : ((cont(X63,X61) & cont(X63,X62)) => leq(X63,intersection(X61,X62)))[input]
11034. ~leq(sK11(sK1,sK0),intersection(sK0,sK1)) {302} (5:7) [backtracking splitting component]
11030. lt(sK11(sK1,sK0),intersection(sK0,sK1)) | ~leq(sK11(sK1,sK0),intersection(sK0,sK1)) {14} (5:14) [resolution 1942,279]
279. eqdim(X0,X1) | lt(X0,X1) | ~leq(X0,X1) (0:9) [cnf transformation 207]
207. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | lt(X0,X1)))[flattening 206]
206. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | lt(X0,X1)))[nnf transformation 150]
150. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[ennf transformation 98]
98. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[flattening 97]
97. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[rectify 18]
18. ! [X34] : ! [X35] : (lt(X34,X35) <=> (leq(X34,X35) & ~eqdim(X34,X35)))[input]
1942. ~eqdim(sK11(sK1,sK0),intersection(sK0,sK1)) {14} (4:7) [resolution 728,459]
728. ~eqdim(X5,sK0) | ~eqdim(X5,intersection(sK0,sK1)) (2:8) [resolution 360,234]
234. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 132]
132. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 60]
60. ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))[flattening 59]
59. ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))[rectify 6]
6. ! [X9] : ! [X10] : (eqdim(X9,X10) => eqdim(X10,X9))[input]
360. ~eqdim(intersection(sK0,sK1),X0) | ~eqdim(X0,sK0) (1:8) [resolution 216,287]
287. eqdim(X0,X2) | ~eqdim(X1,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 163]
163. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~eqdim(X1,X2) | eqdim(X0,X2))[flattening 162]
162. ! [X0,X1,X2] : ((~eqdim(X0,X1) | ~eqdim(X1,X2)) | eqdim(X0,X2))[ennf transformation 112]
112. ! [X0,X1,X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))[flattening 111]
111. ! [X0] : ! [X1] : ! [X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))[rectify 7]
7. ! [X11] : ! [X12] : ! [X13] : ((eqdim(X11,X12) & eqdim(X12,X13)) => eqdim(X11,X13))[input]
216. ~eqdim(intersection(sK0,sK1),sK0) (0:5) [cnf transformation 169]
169. inc(sK0,sK1) & lt(sK0,sK1) & ~eqdim(intersection(sK0,sK1),sK0)[skolemisation 121]
121. ? [X0,X1] : (inc(X0,X1) & lt(X0,X1) & ~eqdim(intersection(X0,X1),X0))[flattening 120]
120. ? [X0,X1] : ((inc(X0,X1) & lt(X0,X1)) & ~eqdim(intersection(X0,X1),X0))[ennf transformation 44]
44. ~! [X0,X1] : ((inc(X0,X1) & lt(X0,X1)) => eqdim(intersection(X0,X1),X0))[flattening 1]
1. ~! [X0] : ! [X1] : ((inc(X0,X1) & lt(X0,X1)) => eqdim(intersection(X0,X1),X0))[negated conjecture]
11049. ~lt(sK11(sK1,sK0),intersection(sK0,sK1)) {14} (5:7) [resolution 10773,236]
10773. lt(intersection(sK0,sK1),sK11(sK1,sK0)) {14} (4:7) [resolution 459,1823]
1823. ~eqdim(X5,sK0) | lt(intersection(sK0,sK1),X5) (3:8) [resolution 586,234]
586. ~eqdim(sK0,X3) | lt(intersection(sK0,sK1),X3) (2:8) [resolution 569,283]
283. ~lt(X2,X0) | lt(X2,X1) | ~eqdim(X0,X1) (0:9) [cnf transformation 155]
155. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[flattening 154]
154. ! [X0,X1,X2] : (~eqdim(X0,X1) | (~lt(X2,X0) | lt(X2,X1)))[ennf transformation 104]
104. ! [X0,X1,X2] : (eqdim(X0,X1) => (lt(X2,X0) => lt(X2,X1)))[flattening 103]
103. ! [X0] : ! [X1] : ! [X2] : (eqdim(X0,X1) => (lt(X2,X0) => lt(X2,X1)))[rectify 9]
9. ! [X16] : ! [X17] : ! [X18] : (eqdim(X16,X17) => (lt(X18,X16) => lt(X18,X17)))[input]
569. lt(intersection(sK0,sK1),sK0) (1:5) [backtracking split refutation 361,365,566]
566. $false {8} (2:0) [subsumption resolution 434,547]
547. lt(intersection(sK0,sK1),sK0) {8} (4:5) [backtracking split refutation 535,540,541]
541. $false {20} (0:0) [subsumption resolution 216,540]
540. eqdim(intersection(sK0,sK1),sK0) {20} (4:5) [backtracking splitting component]
535. lt(intersection(sK0,sK1),sK0) | eqdim(intersection(sK0,sK1),sK0) {8} (4:10) [resolution 440,232]
232. ~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1) (0:9) [cnf transformation 129]
129. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 128]
128. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 56]
56. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 55]
55. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 34]
34. ! [X70] : ! [X71] : (cont(X70,X71) => (lt(X70,X71) | eqdim(X70,X71)))[input]
440. cont(intersection(sK0,sK1),sK0) {8} (3:5) [resolution 437,242]
242. zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0) (0:9) [cnf transformation 145]
145. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 76]
76. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 75]
75. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 30]
30. ! [X59] : ! [X60] : (~zex(intersection(X59,X60)) => cont(intersection(X59,X60),X59))[input]
437. ~zex(intersection(sK0,sK1)) {8} (2:4) [resolution 365,229]
229. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 126]
126. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 52]
52. ! [X0,X1] : (zex(X0) => leq(X0,X1))[flattening 51]
51. ! [X0] : ! [X1] : (zex(X0) => leq(X0,X1))[rectify 40]
40. ! [X83] : ! [X84] : (zex(X83) => leq(X83,X84))[input]
434. ~lt(intersection(sK0,sK1),sK0) {8} (2:5) [resolution 365,277]
277. leq(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 207]
365. ~leq(intersection(sK0,sK1),sK0) {8} (1:5) [backtracking splitting component]
361. lt(intersection(sK0,sK1),sK0) | ~leq(intersection(sK0,sK1),sK0) (1:10) [resolution 216,279]
459. eqdim(sK11(sK1,sK0),sK0) {14} (3:5) [backtracking splitting component]
453. lt(sK11(sK1,sK0),sK0) | eqdim(sK11(sK1,sK0),sK0) {3} (3:10) [resolution 399,232]
302. p(sK11(sK1,sK0),sK0) {3} (1:5) [backtracking splitting component]
299. sP0(sK1,sK0) | p(sK11(sK1,sK0),sK0) (1:8) [resolution 214,271]
271. ~inc(X0,X1) | sP0(X1,X0) | p(sK11(X1,X0),X0) (0:11) [cnf transformation 203]
203. ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | (lt(sK11(X1,X0),X1) & cont(sK11(X1,X0),X1) & p(sK11(X1,X0),X0))) & ((~sP0(X1,X0) & ! [X3] : (~lt(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))[skolemisation 202]
202. ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | ? [X2] : (lt(X2,X1) & cont(X2,X1) & p(X2,X0))) & ((~sP0(X1,X0) & ! [X3] : (~lt(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))[rectify 201]
201. ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | ? [X3] : (lt(X3,X1) & cont(X3,X1) & p(X3,X0))) & ((~sP0(X1,X0) & ! [X3] : (~lt(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))[flattening 200]
200. ! [X0,X1] : ((~inc(X0,X1) | (sP0(X1,X0) | ? [X3] : (lt(X3,X1) & cont(X3,X1) & p(X3,X0)))) & ((~sP0(X1,X0) & ! [X3] : (~lt(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))[nnf transformation 168]
168. ! [X0,X1] : (inc(X0,X1) <=> (sP0(X1,X0) | ? [X3] : (lt(X3,X1) & cont(X3,X1) & p(X3,X0))))[definition folding 167]
214. inc(sK0,sK1) (0:3) [cnf transformation 169]
12199. lt(sK0,sK10(sK0,sK1)) {2} (5:5) [resolution 12150,784]
784. ~sP0(sK1,X3) | lt(sK0,sK10(X3,sK1)) (4:8) [resolution 756,267]
267. p(sK10(X1,X0),X0) | ~sP0(X0,X1) (0:8) [cnf transformation 199]
756. ~p(X4,sK1) | lt(sK0,X4) (3:6) [resolution 594,263]
594. ~eqdim(X5,sK1) | lt(sK0,X5) (2:6) [resolution 309,234]
309. ~eqdim(sK1,X3) | lt(sK0,X3) (1:6) [resolution 215,283]
215. lt(sK0,sK1) (0:3) [cnf transformation 169]
301. cont(sK11(sK1,sK0),sK1) {2} (1:5) [backtracking splitting component]
298. sP0(sK1,sK0) | cont(sK11(sK1,sK0),sK1) (1:8) [resolution 214,270]
270. ~inc(X0,X1) | sP0(X1,X0) | cont(sK11(X1,X0),X1) (0:11) [cnf transformation 203]
12567. lt(sK0,sK10(sK0,sK1)) (5:5) [resolution 12553,784]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1995
Passive clauses: 4607
Generated clauses: 13679
Final active clauses: 521
Final passive clauses: 540
Input formulas: 43
Initial clauses: 83

Unused predicate definitions: 2

Duplicate literals: 26
Trivial inequalities: 5
Bw subsumption resolutions: 277
Fw demodulations: 1157
Condensations: 2

Simple tautologies: 69
Equational tautologies: 137
Forward subsumptions: 7010

Binary resolution: 9527
Factoring: 2
Forward superposition: 28
Backward superposition: 528

Backtracking splits: 384
Backtracking splits refuted: 191
Backtracking splits refuted at zero level: 5

Memory used: 2430KB
Time elapsed: 2.561 s
------------------------------
% Success in time 6.015 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 00:58:06 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal09.all.tptp"
Input read from 
============================ end of footer ===========================
