% Starting sequential portfolio solving with schedule "casc"
% spawned child 12807 with time: 4 (total remaining time 2999)
% WARNING: value z3 for option sas not known
% dis+1010_4_afp=10000:afq=1.2:anc=none:irw=on:lma=on:nm=64:nwc=10:sas=z3:sac=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal10.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7036
% Time elapsed: 0.500 s
% spawned child 12808 with time: 10 (total remaining time 2994)
% ott+11_8_amm=off:anc=none:bsr=on:cond=on:irw=on:nm=2:nwc=1.1:ss=axioms:st=2.0:sac=on_8 on codi_down_sum_theorems_goal10.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8699
% Time elapsed: 1.100 s
% ------------------------------
% ------------------------------
% spawned child 12809 with time: 4 (total remaining time 2983)
% dis+1011_10_aac=none:add=large:afp=10000:afq=1.1:amm=off:anc=none:fsr=off:gs=on:gsem=off:irw=on:nm=6:newcnf=on:nwc=2.5:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal10.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 12153
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 12810 with time: 6 (total remaining time 2977)
% dis+11_28_av=off:fsr=off:irw=on:lcm=predicate:nm=2:newcnf=on:nwc=5:sp=occurrence:urr=on:updr=off_4 on codi_down_sum_theorems_goal10.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 4221
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% spawned child 12811 with time: 6 (total remaining time 2970)
% lrs+1011_5_add=large:afp=1000:afq=1.2:anc=none:fsr=off:irw=on:lma=on:nm=64:newcnf=on:nwc=1:stl=30:sd=3:ss=axioms:st=2.0:sos=on:sac=on:sp=reverse_arity:urr=on:updr=off_4 on codi_down_sum_theorems_goal10.all
% Refutation found. Thanks to Tanya!
% SZS status Theorem for codi_down_sum_theorems_goal10.all
% SZS output start Proof for codi_down_sum_theorems_goal10.all
fof(f1,conjecture,(
  ! [X0] : ! [X1] : ! [X2] : (cont(intersection(X2,X0),X0) | ~(~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',conjecture100)).
fof(f2,negated_conjecture,(
  ~! [X0] : ! [X1] : ! [X2] : (cont(intersection(X2,X0),X0) | ~(~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)))),
  inference(negated_conjecture,[],[f1])).
fof(f6,axiom,(
  ! [X0] : ! [X1] : ! [X3] : (cont(difference(X3,X0),X1) | ~(~cont(X3,X0) & cont(X3,sum(X0,X1))))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom40)).
fof(f8,axiom,(
  ! [X0] : ! [X1] : (difference(X0,X1) = X0 | ~lt(X1,X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom15160)).
fof(f31,axiom,(
  ! [X0] : ! [X1] : ! [X3] : (p(X3,difference(X0,X1)) | ~p(X3,difference(X0,intersection(X0,X1))))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom74380)).
fof(f36,axiom,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) | X0 != X1)),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom74430)).
fof(f37,axiom,(
  ! [X0] : ! [X1] : (p(X0,X1) | ~(~zex(X1) & ~zex(X0) & zex(difference(X1,X0)) & zex(difference(X0,X1))))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom74440)).
fof(f41,axiom,(
  ! [X0] : ! [X1] : (p(intersection(X1,X0),intersection(X0,X1)) | zex(intersection(X1,X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom86950)).
fof(f68,axiom,(
  ! [X0] : ! [X1] : ! [X3] : (((eqdim(X3,intersection(X0,X1)) & cont(X3,X1) & cont(X3,X0)) | ~p(X3,intersection(X0,X1))) & (p(X3,intersection(X0,X1)) | ~(eqdim(X3,intersection(X0,X1)) & cont(X3,X1) & cont(X3,X0))))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom87220)).
fof(f84,axiom,(
  ! [X0] : ! [X1] : (lt(X0,X1) | ~(~zex(X1) & zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom89500)).
fof(f87,axiom,(
  ! [X0] : ! [X1] : ! [X3] : (cont(X0,X3) | ~(cont(X1,X3) & cont(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom89530)).
fof(f88,axiom,(
  ! [X0] : ! [X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom89540)).
fof(f92,axiom,(
  ! [X0] : ! [X1] : ((p(X0,X1) | ~(eqdim(X0,X1) & cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom88840)).
fof(f98,axiom,(
  ! [X0] : ! [X1] : (X0 = X1 | ~(zex(X1) & zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom89130)).
fof(f107,axiom,(
  ? [X0] : zex(X0)),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal10.all.tptp',axiom87750)).
fof(f118,plain,(
  ~! [X0,X1,X2] : (cont(intersection(X2,X0),X0) | ~(~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)))),
  inference(flattening,[],[f2])).
fof(f123,plain,(
  ! [X0] : ! [X1] : ! [X2] : (p(X2,difference(X0,X1)) | ~p(X2,difference(X0,intersection(X0,X1))))),
  inference(rectify,[],[f31])).
fof(f124,plain,(
  ! [X0,X1,X2] : (p(X2,difference(X0,X1)) | ~p(X2,difference(X0,intersection(X0,X1))))),
  inference(flattening,[],[f123])).
fof(f133,plain,(
  ! [X0,X1] : (p(intersection(X1,X0),intersection(X0,X1)) | zex(intersection(X1,X0)))),
  inference(flattening,[],[f41])).
fof(f152,plain,(
  ! [X0,X1] : (p(X0,X1) | ~(~zex(X1) & ~zex(X0) & zex(difference(X1,X0)) & zex(difference(X0,X1))))),
  inference(flattening,[],[f37])).
fof(f156,plain,(
  ! [X0] : ! [X1] : ! [X2] : (cont(X0,X2) | ~(cont(X1,X2) & cont(X0,X1)))),
  inference(rectify,[],[f87])).
fof(f157,plain,(
  ! [X0,X1,X2] : (cont(X0,X2) | ~(cont(X1,X2) & cont(X0,X1)))),
  inference(flattening,[],[f156])).
fof(f168,plain,(
  ! [X0] : ! [X1] : ! [X2] : (((eqdim(X2,intersection(X0,X1)) & cont(X2,X1) & cont(X2,X0)) | ~p(X2,intersection(X0,X1))) & (p(X2,intersection(X0,X1)) | ~(eqdim(X2,intersection(X0,X1)) & cont(X2,X1) & cont(X2,X0))))),
  inference(rectify,[],[f68])).
fof(f169,plain,(
  ! [X0,X1,X2] : (((eqdim(X2,intersection(X0,X1)) & cont(X2,X1) & cont(X2,X0)) | ~p(X2,intersection(X0,X1))) & (p(X2,intersection(X0,X1)) | ~(eqdim(X2,intersection(X0,X1)) & cont(X2,X1) & cont(X2,X0))))),
  inference(flattening,[],[f168])).
fof(f175,plain,(
  ! [X0,X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  inference(flattening,[],[f88])).
fof(f178,plain,(
  ! [X0,X1] : ((p(X0,X1) | ~(eqdim(X0,X1) & cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  inference(flattening,[],[f92])).
fof(f179,plain,(
  ! [X0] : ! [X1] : ! [X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  inference(rectify,[],[f6])).
fof(f180,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  inference(flattening,[],[f179])).
fof(f215,plain,(
  ! [X0,X1] : (lt(X0,X1) | ~(~zex(X1) & zex(X0)))),
  inference(flattening,[],[f84])).
fof(f217,plain,(
  ! [X0,X1] : (difference(X0,X1) = X0 | ~lt(X1,X0))),
  inference(flattening,[],[f8])).
fof(f228,plain,(
  ! [X0,X1] : (X0 = X1 | ~(zex(X1) & zex(X0)))),
  inference(flattening,[],[f98])).
fof(f229,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) | X0 != X1)),
  inference(flattening,[],[f36])).
fof(f249,plain,(
  ? [X0,X1,X2] : (~cont(intersection(X2,X0),X0) & (~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f118])).
fof(f250,plain,(
  ? [X0,X1,X2] : (~cont(intersection(X2,X0),X0) & ~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1))),
  inference(flattening,[],[f249])).
fof(f265,plain,(
  ! [X0,X1] : (p(X0,X1) | (zex(X1) | zex(X0) | ~zex(difference(X1,X0)) | ~zex(difference(X0,X1))))),
  inference(ennf_transformation,[],[f152])).
fof(f266,plain,(
  ! [X0,X1] : (p(X0,X1) | zex(X1) | zex(X0) | ~zex(difference(X1,X0)) | ~zex(difference(X0,X1)))),
  inference(flattening,[],[f265])).
fof(f269,plain,(
  ! [X0,X1,X2] : (cont(X0,X2) | (~cont(X1,X2) | ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f157])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1))),
  inference(flattening,[],[f269])).
fof(f281,plain,(
  ! [X0,X1,X2] : (((eqdim(X2,intersection(X0,X1)) & cont(X2,X1) & cont(X2,X0)) | ~p(X2,intersection(X0,X1))) & (p(X2,intersection(X0,X1)) | (~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0))))),
  inference(ennf_transformation,[],[f169])).
fof(f282,plain,(
  ! [X0,X1,X2] : (((eqdim(X2,intersection(X0,X1)) & cont(X2,X1) & cont(X2,X0)) | ~p(X2,intersection(X0,X1))) & (p(X2,intersection(X0,X1)) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0)))),
  inference(flattening,[],[f281])).
fof(f288,plain,(
  ! [X0,X1] : ((p(X0,X1) | (~eqdim(X0,X1) | ~cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  inference(ennf_transformation,[],[f178])).
fof(f289,plain,(
  ! [X0,X1] : ((p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1)) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  inference(flattening,[],[f288])).
fof(f290,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | (cont(X2,X0) | ~cont(X2,sum(X0,X1))))),
  inference(ennf_transformation,[],[f180])).
fof(f291,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f290])).
fof(f314,plain,(
  ! [X0,X1] : (lt(X0,X1) | (zex(X1) | ~zex(X0)))),
  inference(ennf_transformation,[],[f215])).
fof(f315,plain,(
  ! [X0,X1] : (lt(X0,X1) | zex(X1) | ~zex(X0))),
  inference(flattening,[],[f314])).
fof(f325,plain,(
  ! [X0,X1] : (X0 = X1 | (~zex(X1) | ~zex(X0)))),
  inference(ennf_transformation,[],[f228])).
fof(f326,plain,(
  ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0))),
  inference(flattening,[],[f325])).
fof(f343,plain,(
  cont(sK2,sum(sK0,sK1))),
  inference(cnf_transformation,[],[f250])).
fof(f344,plain,(
  ~cont(sK2,sK0)),
  inference(cnf_transformation,[],[f250])).
fof(f345,plain,(
  ~cont(sK2,sK1)),
  inference(cnf_transformation,[],[f250])).
fof(f348,plain,(
  ~cont(intersection(sK2,sK0),sK0)),
  inference(cnf_transformation,[],[f250])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (p(X2,difference(X0,X1)) | ~p(X2,difference(X0,intersection(X0,X1)))) )),
  inference(cnf_transformation,[],[f124])).
fof(f358,plain,(
  ( ! [X0,X1] : (p(intersection(X1,X0),intersection(X0,X1)) | zex(intersection(X1,X0))) )),
  inference(cnf_transformation,[],[f133])).
fof(f374,plain,(
  zex(sK4)),
  inference(cnf_transformation,[],[f107])).
fof(f376,plain,(
  ( ! [X0,X1] : (~zex(difference(X0,X1)) | ~zex(difference(X1,X0)) | zex(X0) | zex(X1) | p(X0,X1)) )),
  inference(cnf_transformation,[],[f266])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f270])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (cont(X2,X0) | ~p(X2,intersection(X0,X1))) )),
  inference(cnf_transformation,[],[f282])).
fof(f395,plain,(
  ( ! [X0,X1] : (~zex(X0) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f175])).
fof(f402,plain,(
  ( ! [X0,X1] : (cont(X0,X1) | ~p(X0,X1)) )),
  inference(cnf_transformation,[],[f289])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1))) )),
  inference(cnf_transformation,[],[f291])).
fof(f437,plain,(
  ( ! [X0,X1] : (lt(X0,X1) | zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f315])).
fof(f439,plain,(
  ( ! [X0,X1] : (~lt(X1,X0) | difference(X0,X1) = X0) )),
  inference(cnf_transformation,[],[f217])).
fof(f453,plain,(
  ( ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1) )),
  inference(cnf_transformation,[],[f326])).
fof(f454,plain,(
  ( ! [X0,X1] : (X0 != X1 | zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f229])).
fof(f474,plain,(
  ( ! [X1] : (zex(difference(X1,X1))) )),
  inference(equality_resolution,[],[f454])).
fof(f556,plain,(
  cont(difference(sK2,sK0),sK1)),
  inference(unit_resulting_resolution,[],[f344,f343,f407])).
fof(f562,plain,(
  ~zex(sK2)),
  inference(unit_resulting_resolution,[],[f343,f395])).
fof(f576,plain,(
  p(sK2,sK2)),
  inference(unit_resulting_resolution,[],[f562,f474,f474,f562,f376])).
fof(f585,plain,(
  lt(sK4,sK2)),
  inference(unit_resulting_resolution,[],[f374,f562,f437])).
fof(f645,plain,(
  ( ! [X0] : (~p(intersection(sK2,sK0),intersection(sK0,X0))) )),
  inference(unit_resulting_resolution,[],[f348,f385])).
fof(f840,plain,(
  difference(sK2,sK4) = sK2),
  inference(unit_resulting_resolution,[],[f585,f439])).
fof(f1385,plain,(
  ~cont(sK2,difference(sK2,sK0))),
  inference(unit_resulting_resolution,[],[f345,f556,f379])).
fof(f2175,plain,(
  zex(intersection(sK2,sK0))),
  inference(unit_resulting_resolution,[],[f645,f358])).
fof(f2677,plain,(
  intersection(sK2,sK0) = sK4),
  inference(unit_resulting_resolution,[],[f374,f2175,f453])).
fof(f3497,plain,(
  ~p(sK2,difference(sK2,sK0))),
  inference(unit_resulting_resolution,[],[f1385,f402])).
fof(f4836,plain,(
  ~p(sK2,difference(sK2,intersection(sK2,sK0)))),
  inference(unit_resulting_resolution,[],[f3497,f351])).
fof(f4851,plain,(
  ~p(sK2,difference(sK2,sK4))),
  inference(forward_demodulation,[],[f4836,f2677])).
fof(f4852,plain,(
  ~p(sK2,sK2)),
  inference(forward_demodulation,[],[f4851,f840])).
fof(f4867,plain,(
  $false),
  inference(resolution,[],[f4852,f576])).
% SZS output end Proof for codi_down_sum_theorems_goal10.all
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 6780
% Time elapsed: 0.357 s
% ------------------------------
% ------------------------------
% Success in time 3.338 s
