% Starting sequential portfolio solving with schedule "casc"
% spawned child 10016 with time: 4 (total remaining time 2999)
% WARNING: value z3 for option sas not known
% dis+1010_4_afp=10000:afq=1.2:anc=none:irw=on:lma=on:nm=64:nwc=10:sas=z3:sac=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal04.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6652
% Time elapsed: 0.500 s
% spawned child 10017 with time: 10 (total remaining time 2994)
% ott+11_8_amm=off:anc=none:bsr=on:cond=on:irw=on:nm=2:nwc=1.1:ss=axioms:st=2.0:sac=on_8 on codi_down_sum_theorems_goal04.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7547
% Time elapsed: 1.100 s
% ------------------------------
% ------------------------------
% spawned child 10018 with time: 4 (total remaining time 2984)
% dis+1011_10_aac=none:add=large:afp=10000:afq=1.1:amm=off:anc=none:fsr=off:gs=on:gsem=off:irw=on:nm=6:newcnf=on:nwc=2.5:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal04.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 11129
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 10019 with time: 6 (total remaining time 2979)
% dis+11_28_av=off:fsr=off:irw=on:lcm=predicate:nm=2:newcnf=on:nwc=5:sp=occurrence:urr=on:updr=off_4 on codi_down_sum_theorems_goal04.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 4221
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% spawned child 10020 with time: 6 (total remaining time 2972)
% lrs+1011_5_add=large:afp=1000:afq=1.2:anc=none:fsr=off:irw=on:lma=on:nm=64:newcnf=on:nwc=1:stl=30:sd=3:ss=axioms:st=2.0:sos=on:sac=on:sp=reverse_arity:urr=on:updr=off_4 on codi_down_sum_theorems_goal04.all
% Refutation found. Thanks to Tanya!
% SZS status Theorem for codi_down_sum_theorems_goal04.all
% SZS output start Proof for codi_down_sum_theorems_goal04.all
fof(f1,conjecture,(
  ! [X0] : ! [X1] : (zex(sum(X0,X1)) | ~(zex(X1) & zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal04.all.tptp',conjecture40)).
fof(f2,negated_conjecture,(
  ~! [X0] : ! [X1] : (zex(sum(X0,X1)) | ~(zex(X1) & zex(X0)))),
  inference(negated_conjecture,[],[f1])).
fof(f6,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal04.all.tptp',axiom40)).
fof(f85,axiom,(
  ! [X0] : ((~zex(X0) | ~cont(X0,X0)) & (cont(X0,X0) | zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal04.all.tptp',axiom89510)).
fof(f88,axiom,(
  ! [X0] : ! [X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal04.all.tptp',axiom89540)).
fof(f98,axiom,(
  ! [X0] : ! [X1] : (X0 = X1 | ~(zex(X1) & zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal04.all.tptp',axiom89130)).
fof(f118,plain,(
  ~! [X0,X1] : (zex(sum(X0,X1)) | ~(zex(X1) & zex(X0)))),
  inference(flattening,[],[f2])).
fof(f119,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  inference(flattening,[],[f6])).
fof(f173,plain,(
  ! [X0,X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  inference(flattening,[],[f88])).
fof(f183,plain,(
  ! [X0,X1] : (X0 = X1 | ~(zex(X1) & zex(X0)))),
  inference(flattening,[],[f98])).
fof(f212,plain,(
  ? [X0,X1] : (~zex(sum(X0,X1)) & (zex(X1) & zex(X0)))),
  inference(ennf_transformation,[],[f118])).
fof(f213,plain,(
  ? [X0,X1] : (~zex(sum(X0,X1)) & zex(X1) & zex(X0))),
  inference(flattening,[],[f212])).
fof(f214,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | (cont(X2,X0) | ~cont(X2,sum(X0,X1))))),
  inference(ennf_transformation,[],[f119])).
fof(f215,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f214])).
fof(f276,plain,(
  ! [X0,X1] : (X0 = X1 | (~zex(X1) | ~zex(X0)))),
  inference(ennf_transformation,[],[f183])).
fof(f277,plain,(
  ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0))),
  inference(flattening,[],[f276])).
fof(f305,plain,(
  zex(sK0)),
  inference(cnf_transformation,[],[f213])).
fof(f306,plain,(
  zex(sK1)),
  inference(cnf_transformation,[],[f213])).
fof(f307,plain,(
  ~zex(sum(sK0,sK1))),
  inference(cnf_transformation,[],[f213])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1))) )),
  inference(cnf_transformation,[],[f215])).
fof(f370,plain,(
  ( ! [X0,X1] : (~zex(X0) | ~cont(X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f376,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f85])).
fof(f391,plain,(
  ( ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1) )),
  inference(cnf_transformation,[],[f277])).
fof(f439,plain,(
  ( ! [X0] : (~cont(X0,sK0)) )),
  inference(unit_resulting_resolution,[],[f305,f370])).
fof(f495,plain,(
  sK0 = sK1),
  inference(unit_resulting_resolution,[],[f305,f306,f391])).
fof(f513,plain,(
  ~zex(sum(sK0,sK0))),
  inference(backward_demodulation,[],[f495,f307])).
fof(f543,plain,(
  ( ! [X0] : (~cont(X0,sum(sK0,sK0))) )),
  inference(unit_resulting_resolution,[],[f439,f439,f308])).
fof(f1125,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f513,f543,f376])).
% SZS output end Proof for codi_down_sum_theorems_goal04.all
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.209 s
% ------------------------------
% ------------------------------
% Success in time 2.983 s
