Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal07.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 418
Passive clauses: 1512
Generated clauses: 4147
Final active clauses: 416
Final passive clauses: 1012
Input formulas: 43
Initial clauses: 92

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 50
Fw subsumption resolutions: 100

Simple tautologies: 48
Forward subsumptions: 2477

Binary resolution: 3874
Factoring: 29

Backtracking splits: 35
Backtracking splits refuted: 8

Memory used: 767KB
Time elapsed: 0.419 s
------------------------------
% remaining time: 5993 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal07.all
Refutation found. Thanks to Tanya!
2210. $false (9:0) [subsumption resolution 2205,215]
215. eqdim(sK2,sK0) (0:3) [cnf transformation 168]
168. cont(sK2,sK0) & cont(sK2,sK1) & eqdim(sK2,sK0) & eqdim(sK2,sK1) & ~cont(sK2,intersection(sK0,sK1))[skolemisation 120]
120. ? [X0,X1,X2] : (cont(X2,X0) & cont(X2,X1) & eqdim(X2,X0) & eqdim(X2,X1) & ~cont(X2,intersection(X0,X1)))[flattening 119]
119. ? [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,X0) & eqdim(X2,X1)) & ~cont(X2,intersection(X0,X1)))[ennf transformation 44]
44. ~! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,X0) & eqdim(X2,X1)) => cont(X2,intersection(X0,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,X0) & eqdim(X2,X1)) => cont(X2,intersection(X0,X1)))[negated conjecture]
2205. ~eqdim(sK2,sK0) (9:3) [resolution 2197,237]
237. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 133]
133. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 63]
63. ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))[flattening 62]
62. ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))[rectify 6]
6. ! [X10] : ! [X11] : (eqdim(X10,X11) => eqdim(X11,X10))[input]
2197. ~eqdim(sK0,sK2) (8:3) [subsumption resolution 2187,1577]
1577. ~zex(intersection(sK0,sK1)) (5:4) [resolution 1572,230]
230. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 125]
125. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 52]
52. ! [X0,X1] : (zex(X0) => leq(X0,X1))[flattening 51]
51. ! [X0] : ! [X1] : (zex(X0) => leq(X0,X1))[rectify 40]
40. ! [X84] : ! [X85] : (zex(X84) => leq(X84,X85))[input]
1572. ~leq(intersection(sK0,sK1),sK2) (4:5) [subsumption resolution 1571,213]
213. cont(sK2,sK0) (0:3) [cnf transformation 168]
1571. ~leq(intersection(sK0,sK1),sK2) | ~cont(sK2,sK0) (4:8) [subsumption resolution 1555,214]
214. cont(sK2,sK1) (0:3) [cnf transformation 168]
1555. ~leq(intersection(sK0,sK1),sK2) | ~cont(sK2,sK1) | ~cont(sK2,sK0) (4:11) [resolution 1180,287]
287. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 160]
160. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 159]
159. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 109]
109. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[flattening 108]
108. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[rectify 31]
31. ! [X62] : ! [X63] : ! [X64] : ((cont(X64,X62) & cont(X64,X63)) => leq(X64,intersection(X62,X63)))[input]
1180. ~leq(sK2,intersection(sK0,sK1)) | ~leq(intersection(sK0,sK1),sK2) (3:10) [resolution 1176,280]
280. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 206]
206. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 205]
205. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[nnf transformation 97]
97. ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[flattening 96]
96. ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[rectify 28]
28. ! [X56] : ! [X57] : (eqdim(X56,X57) <=> (leq(X56,X57) & leq(X57,X56)))[input]
1176. ~eqdim(sK2,intersection(sK0,sK1)) (2:5) [subsumption resolution 1175,213]
1175. ~eqdim(sK2,intersection(sK0,sK1)) | ~cont(sK2,sK0) (2:8) [subsumption resolution 1167,214]
1167. ~eqdim(sK2,intersection(sK0,sK1)) | ~cont(sK2,sK1) | ~cont(sK2,sK0) (2:11) [resolution 290,319]
319. ~p(sK2,intersection(sK0,sK1)) (1:5) [resolution 272,217]
217. ~cont(sK2,intersection(sK0,sK1)) (0:5) [cnf transformation 168]
272. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 202]
202. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 201]
201. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[nnf transformation 93]
93. ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[flattening 92]
92. ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[rectify 19]
19. ! [X37] : ! [X38] : (p(X37,X38) <=> (cont(X37,X38) & eqdim(X37,X38)))[input]
290. p(X2,intersection(X0,X1)) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:16) [cnf transformation 209]
209. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 208]
208. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[nnf transformation 115]
115. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[flattening 114]
114. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[rectify 32]
32. ! [X65] : ! [X66] : ! [X67] : ((cont(X67,X65) & cont(X67,X66) & eqdim(X67,intersection(X65,X66))) <=> p(X67,intersection(X65,X66)))[input]
2187. ~eqdim(sK0,sK2) | zex(intersection(sK0,sK1)) (8:7) [resolution 1873,243]
243. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 144]
144. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 75]
75. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 74]
74. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 30]
30. ! [X60] : ! [X61] : (~zex(intersection(X60,X61)) => cont(intersection(X60,X61),X60))[input]
1873. ~cont(intersection(sK0,sK1),X0) | ~eqdim(X0,sK2) (7:8) [subsumption resolution 1857,1191]
1191. ~eqdim(intersection(sK0,sK1),X0) | ~eqdim(X0,sK2) (4:8) [resolution 1182,289]
289. eqdim(X0,X2) | ~eqdim(X1,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 164]
164. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~eqdim(X1,X2) | eqdim(X0,X2))[flattening 163]
163. ! [X0,X1,X2] : ((~eqdim(X0,X1) | ~eqdim(X1,X2)) | eqdim(X0,X2))[ennf transformation 113]
113. ! [X0,X1,X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))[flattening 112]
112. ! [X0] : ! [X1] : ! [X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))[rectify 7]
7. ! [X12] : ! [X13] : ! [X14] : ((eqdim(X12,X13) & eqdim(X13,X14)) => eqdim(X12,X14))[input]
1182. ~eqdim(intersection(sK0,sK1),sK2) (3:5) [resolution 1176,237]
1857. ~eqdim(X0,sK2) | eqdim(intersection(sK0,sK1),X0) | ~cont(intersection(sK0,sK1),X0) (7:13) [resolution 1638,233]
233. lt(X0,X1) | eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 128]
128. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 127]
127. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 56]
56. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 55]
55. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 36]
36. ! [X76] : ! [X77] : (cont(X76,X77) => (lt(X76,X77) | eqdim(X76,X77)))[input]
1638. ~lt(intersection(sK0,sK1),X3) | ~eqdim(X3,sK2) (6:8) [resolution 1578,284]
284. lt(X2,X1) | ~lt(X2,X0) | ~eqdim(X0,X1) (0:9) [cnf transformation 154]
154. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[flattening 153]
153. ! [X0,X1,X2] : (~eqdim(X0,X1) | (~lt(X2,X0) | lt(X2,X1)))[ennf transformation 103]
103. ! [X0,X1,X2] : (eqdim(X0,X1) => (lt(X2,X0) => lt(X2,X1)))[flattening 102]
102. ! [X0] : ! [X1] : ! [X2] : (eqdim(X0,X1) => (lt(X2,X0) => lt(X2,X1)))[rectify 9]
9. ! [X15] : ! [X16] : ! [X17] : (eqdim(X15,X16) => (lt(X17,X15) => lt(X17,X16)))[input]
1578. ~lt(intersection(sK0,sK1),sK2) (5:5) [resolution 1572,275]
275. leq(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 204]
204. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | lt(X0,X1)))[flattening 203]
203. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | lt(X0,X1)))[nnf transformation 149]
149. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[ennf transformation 95]
95. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[flattening 94]
94. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[rectify 18]
18. ! [X35] : ! [X36] : (lt(X35,X36) <=> (leq(X35,X36) & ~eqdim(X35,X36)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 287
Passive clauses: 863
Generated clauses: 1992
Final active clauses: 287
Final passive clauses: 548
Input formulas: 43
Initial clauses: 87

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 53
Fw subsumption resolutions: 120

Simple tautologies: 42
Equational tautologies: 3
Forward subsumptions: 937

Binary resolution: 995
Factoring: 14
Forward superposition: 668
Backward superposition: 58

Unique components: 5

Memory used: 511KB
Time elapsed: 0.148 s
------------------------------
% Success in time 0.892 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 00:57:39 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal07.all.tptp"
Input read from 
============================ end of footer ===========================
