% Starting sequential portfolio solving with schedule "casc"
% spawned child 9953 with time: 4 (total remaining time 2999)
% WARNING: value z3 for option sas not known
% dis+1010_4_afp=10000:afq=1.2:anc=none:irw=on:lma=on:nm=64:nwc=10:sas=z3:sac=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal01.all
% Refutation found. Thanks to Tanya!
% SZS status Theorem for codi_down_sum_theorems_goal01.all
% SZS output start Proof for codi_down_sum_theorems_goal01.all
fof(f1,conjecture,(
  ! [X0] : ! [X1] : ! [X2] : (? [X3] : ((cont(X3,X1) & p(X3,X2)) | (cont(X3,X0) & p(X3,X2))) | ~cont(X2,sum(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal01.all.tptp',conjecture10)).
fof(f2,negated_conjecture,(
  ~! [X0] : ! [X1] : ! [X2] : (? [X3] : ((cont(X3,X1) & p(X3,X2)) | (cont(X3,X0) & p(X3,X2))) | ~cont(X2,sum(X0,X1)))),
  inference(negated_conjecture,[],[f1])).
fof(f6,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal01.all.tptp',axiom40)).
fof(f21,axiom,(
  ! [X0] : ! [X1] : (p(difference(X0,X1),X0) | zex(difference(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal01.all.tptp',axiom64980)).
fof(f76,axiom,(
  ! [X0] : eqdim(X0,X0)),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal01.all.tptp',axiom89420)).
fof(f85,axiom,(
  ! [X0] : ((~zex(X0) | ~cont(X0,X0)) & (cont(X0,X0) | zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal01.all.tptp',axiom89510)).
fof(f88,axiom,(
  ! [X0] : ! [X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal01.all.tptp',axiom89540)).
fof(f92,axiom,(
  ! [X0] : ! [X1] : ((p(X0,X1) | ~(eqdim(X0,X1) & cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal01.all.tptp',axiom88840)).
fof(f118,plain,(
  ~! [X0,X1,X2] : (? [X3] : ((cont(X3,X1) & p(X3,X2)) | (cont(X3,X0) & p(X3,X2))) | ~cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f2])).
fof(f152,plain,(
  ! [X0,X1] : ((p(X0,X1) | ~(eqdim(X0,X1) & cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  inference(flattening,[],[f92])).
fof(f158,plain,(
  ! [X0,X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  inference(flattening,[],[f88])).
fof(f179,plain,(
  ! [X0,X1] : (p(difference(X0,X1),X0) | zex(difference(X0,X1)))),
  inference(flattening,[],[f21])).
fof(f233,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  inference(flattening,[],[f6])).
fof(f236,plain,(
  ? [X0,X1,X2] : (! [X3] : ((~cont(X3,X1) | ~p(X3,X2)) & (~cont(X3,X0) | ~p(X3,X2))) & cont(X2,sum(X0,X1)))),
  inference(ennf_transformation,[],[f118])).
fof(f264,plain,(
  ! [X0,X1] : ((p(X0,X1) | (~eqdim(X0,X1) | ~cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  inference(ennf_transformation,[],[f152])).
fof(f265,plain,(
  ! [X0,X1] : ((p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1)) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  inference(flattening,[],[f264])).
fof(f344,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | (cont(X2,X0) | ~cont(X2,sum(X0,X1))))),
  inference(ennf_transformation,[],[f233])).
fof(f345,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f344])).
fof(f350,plain,(
  ? [X0,X1,X2] : (! [X3] : ((~cont(X3,X1) | ~p(X3,X2)) & (~cont(X3,X0) | ~p(X3,X2))) & cont(X2,sum(X0,X1))) => (! [X3] : ((~cont(X3,sK1) | ~p(X3,sK2)) & (~cont(X3,sK0) | ~p(X3,sK2))) & cont(sK2,sum(sK0,sK1)))),
  introduced(choice_axiom,[])).
fof(f351,plain,(
  ! [X3] : ((~cont(X3,sK1) | ~p(X3,sK2)) & (~cont(X3,sK0) | ~p(X3,sK2))) & cont(sK2,sum(sK0,sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f236,f350])).
fof(f402,plain,(
  cont(sK2,sum(sK0,sK1))),
  inference(cnf_transformation,[],[f351])).
fof(f403,plain,(
  ( ! [X3] : (~p(X3,sK2) | ~cont(X3,sK0)) )),
  inference(cnf_transformation,[],[f351])).
fof(f404,plain,(
  ( ! [X3] : (~p(X3,sK2) | ~cont(X3,sK1)) )),
  inference(cnf_transformation,[],[f351])).
fof(f406,plain,(
  ( ! [X0] : (eqdim(X0,X0)) )),
  inference(cnf_transformation,[],[f76])).
fof(f421,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f85])).
fof(f479,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f491,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f158])).
fof(f513,plain,(
  ( ! [X0,X1] : (p(difference(X0,X1),X0) | zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f179])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (~cont(X2,sum(X0,X1)) | cont(X2,X0) | cont(difference(X2,X0),X1)) )),
  inference(cnf_transformation,[],[f345])).
fof(f593,plain,(
  ~zex(sK2)),
  inference(resolution,[],[f491,f402])).
fof(f905,plain,(
  ( ! [X0] : (~cont(X0,X0) | p(X0,X0)) )),
  inference(resolution,[],[f479,f406])).
fof(f909,plain,(
  ( ! [X0] : (p(X0,X0) | zex(X0)) )),
  inference(resolution,[],[f905,f421])).
fof(f918,plain,(
  zex(sK2) | ~cont(sK2,sK0)),
  inference(resolution,[],[f909,f403])).
fof(f920,plain,(
  ~cont(sK2,sK0)),
  inference(subsumption_resolution,[],[f918,f593])).
fof(f1076,plain,(
  ( ! [X18] : (zex(difference(sK2,X18)) | ~cont(difference(sK2,X18),sK1)) )),
  inference(resolution,[],[f513,f404])).
fof(f1080,plain,(
  ( ! [X18] : (~cont(difference(sK2,X18),sK1)) )),
  inference(subsumption_resolution,[],[f1076,f491])).
fof(f2881,plain,(
  cont(sK2,sK0) | cont(difference(sK2,sK0),sK1)),
  inference(resolution,[],[f580,f402])).
fof(f2894,plain,(
  cont(difference(sK2,sK0),sK1)),
  inference(subsumption_resolution,[],[f2881,f920])).
fof(f2895,plain,(
  $false),
  inference(subsumption_resolution,[],[f2894,f1080])).
% SZS output end Proof for codi_down_sum_theorems_goal01.all
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 6012
% Time elapsed: 0.280 s
% ------------------------------
% ------------------------------
% Success in time 0.384 s
