Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal22.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 663
Passive clauses: 2775
Generated clauses: 5676
Final active clauses: 661
Final passive clauses: 2002
Input formulas: 43
Initial clauses: 89

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 43
Fw subsumption resolutions: 66

Simple tautologies: 37
Forward subsumptions: 2780

Binary resolution: 5350
Factoring: 39

Backtracking splits: 76
Backtracking splits refuted: 6

Memory used: 1151KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal22.all
Refutation found. Thanks to Tanya!
707. $false (4:0) [subsumption resolution 706,495]
495. c(y,sK0) (3:3) [resolution 491,277]
277. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 206]
206. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[nnf transformation 149]
149. ! [X0,X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[ennf transformation 99]
99. ! [X0,X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[flattening 98]
98. ! [X0] : ! [X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[rectify 20]
20. ! [X41] : ! [X42] : (~c(X41,X42) <=> zex(intersection(X41,X42)))[input]
491. ~zex(intersection(y,sK0)) (2:4) [resolution 486,326]
326. c(intersection(y,sK0),sK0) (1:5) [resolution 277,213]
213. ~zex(intersection(intersection(y,sK0),sK0)) (0:6) [cnf transformation 167]
167. zex(intersection(sK0,y)) & ~zex(intersection(intersection(y,sK0),sK0))[skolemisation 119]
119. ? [X0] : (zex(intersection(X0,y)) & ~zex(intersection(intersection(y,X0),X0)))[ennf transformation 1]
1. ~! [X0] : (zex(intersection(X0,y)) => zex(intersection(intersection(y,X0),X0)))[negated conjecture]
486. ~c(X2,X3) | ~zex(X2) (1:5) [resolution 265,227]
227. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 125]
125. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 53]
53. ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))[flattening 52]
52. ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))[rectify 16]
16. ! [X32] : ! [X33] : (zex(X32) => (~cont(X33,X32) & ~cont(X32,X33)))[input]
265. cont(sK10(X1,X0),X0) | ~c(X0,X1) (0:8) [cnf transformation 199]
199. ! [X0,X1] : ((~c(X0,X1) | (cont(sK10(X1,X0),X0) & cont(sK10(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 198]
198. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[rectify 197]
197. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | c(X0,X1)))[nnf transformation 91]
91. ! [X0,X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))[flattening 90]
90. ! [X0] : ! [X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))[rectify 2]
2. ! [X1] : ! [X2] : (c(X1,X2) <=> ? [X3] : (cont(X3,X1) & cont(X3,X2)))[input]
706. ~c(y,sK0) (4:3) [duplicate literal removal 697]
697. ~c(y,sK0) | ~c(y,sK0) (4:6) [resolution 574,265]
574. ~cont(sK10(sK0,X6),y) | ~c(X6,sK0) (3:8) [resolution 562,266]
266. cont(sK10(X1,X0),X1) | ~c(X0,X1) (0:8) [cnf transformation 199]
562. ~cont(X0,sK0) | ~cont(X0,y) (2:6) [resolution 267,327]
327. ~c(sK0,y) (1:3) [resolution 278,212]
212. zex(intersection(sK0,y)) (0:4) [cnf transformation 167]
278. ~zex(intersection(X0,X1)) | ~c(X0,X1) (0:7) [cnf transformation 206]
267. c(X0,X1) | ~cont(X3,X1) | ~cont(X3,X0) (0:9) [cnf transformation 199]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 110
Passive clauses: 271
Generated clauses: 490
Final active clauses: 110
Final passive clauses: 141
Input formulas: 43
Initial clauses: 84

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 26
Fw subsumption resolutions: 33

Simple tautologies: 25
Equational tautologies: 1
Forward subsumptions: 146

Binary resolution: 166
Factoring: 4
Forward superposition: 134
Backward superposition: 44

Unique components: 2

Memory used: 255KB
Time elapsed: 0.022 s
------------------------------
% Success in time 0.487 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 23:39:29 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal22.all.tptp"
Input read from 
============================ end of footer ===========================
