Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_down_sum_theorems_goal15.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1102
Passive clauses: 10543
Generated clauses: 14814
Final active clauses: 1092
Final passive clauses: 9238
Input formulas: 78
Initial clauses: 123

Selected by SInE selection: 78
SInE iterations: 1

Duplicate literals: 53
Fw subsumption resolutions: 273

Simple tautologies: 47
Forward subsumptions: 3939

Binary resolution: 14233
Factoring: 29

Backtracking splits: 119
Backtracking splits refuted: 7
Backtracking splits refuted at zero level: 7

Memory used: 3070KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_down_sum_theorems_goal15.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1431
Passive clauses: 12127
Generated clauses: 34168
Final active clauses: 1431
Final passive clauses: 10234
Input formulas: 78
Initial clauses: 122

Selected by SInE selection: 78
SInE iterations: 1

Duplicate literals: 674
Fw subsumption resolutions: 1805
Fw demodulations: 12

Simple tautologies: 1160
Equational tautologies: 11
Forward subsumptions: 18869

Binary resolution: 23073
Factoring: 139
Forward superposition: 7703
Backward superposition: 761
Self superposition: 2

Unique components: 6

Memory used: 4861KB
Time elapsed: 1.403 s
------------------------------
% remaining time: 5981 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on theorems\conversions\codi_down_sum_theorems_goal15.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 868
Passive clauses: 14082
Generated clauses: 133694
Final active clauses: 262
Final passive clauses: 650
Input formulas: 78
Initial clauses: 120

Duplicate literals: 131
Trivial inequalities: 486
Bw subsumption resolutions: 6361
Fw demodulations: 74834
Condensations: 131

Simple tautologies: 264
Equational tautologies: 53
Forward subsumptions: 38532

Binary resolution: 7138
Forward superposition: 731
Backward superposition: 1882

Backtracking splits: 1227
Backtracking splits refuted: 659
Backtracking splits refuted at zero level: 37

Memory used: 12665KB
Time elapsed: 3.301 s
------------------------------
% remaining time: 5948 next slice time: 95
dis+2_4_bs=off:ep=on:n=on:nicw=on:nwc=1.5:ptb=off:sac=on:sio=off:spl=backtracking:ssec=off_95 on theorems\conversions\codi_down_sum_theorems_goal15.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 3559
Passive clauses: 87301
Generated clauses: 212678
Final active clauses: 3391
Final passive clauses: 81414
Input formulas: 78
Initial clauses: 122

Duplicate literals: 2697
Fw subsumption resolutions: 12524
Fw demodulations: 6445
Bw demodulations: 165

Simple tautologies: 1620
Equational tautologies: 41
Forward subsumptions: 104020

Binary resolution: 181717
Factoring: 47
Forward superposition: 7292
Backward superposition: 1871

Backtracking splits: 14
Backtracking splits refuted: 3
Backtracking splits refuted at zero level: 2

Memory used: 38378KB
Time elapsed: 9.625 s
------------------------------
% remaining time: 5851 next slice time: 4
lrs-1003_12_bms=on:bs=unit_only:bsr=on:cond=fast:drc=off:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=10:sfv=off:sgo=on:sio=off:sp=occurrence:spo=on:ssec=off:stl=81:updr=off_4 on theorems\conversions\codi_down_sum_theorems_goal15.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 738
Passive clauses: 2445
Generated clauses: 6889
Final active clauses: 736
Final passive clauses: 1697
Input formulas: 78
Initial clauses: 122

Duplicate literals: 27
Fw subsumption resolutions: 364
Bw subsumption resolutions: 2
Fw demodulations: 375
Condensations: 23

Simple tautologies: 114
Equational tautologies: 1
Forward subsumptions: 3535
Backward subsumptions: 10

Binary resolution: 5173
Factoring: 78
Forward superposition: 379
Backward superposition: 346

Unique components: 2043

Memory used: 1918KB
Time elapsed: 0.507 s
------------------------------
% remaining time: 5846 next slice time: 22
lrs+1010_50_bs=off:drc=off:flr=on:lcm=predicate:n=on:nwc=1.2:sfv=off:sgo=on:stl=81_22 on theorems\conversions\codi_down_sum_theorems_goal15.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1594
Passive clauses: 3655
Generated clauses: 55553
Final active clauses: 1594
Final passive clauses: 2061
Input formulas: 78
Initial clauses: 120

Duplicate literals: 1562
Fw subsumption resolutions: 5085
Fw literal rewrites: 9520

Simple tautologies: 1262
Equational tautologies: 4
Forward subsumptions: 34393

Binary resolution: 35990
Factoring: 16
Forward superposition: 3183
Backward superposition: 98
Equality resolution: 1

Unique components: 6

Memory used: 3965KB
Time elapsed: 2.304 s
------------------------------
% remaining time: 5822 next slice time: 41
dis-1010_5_bd=off:bs=off:cond=fast:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=1.3:ptb=off:sac=on:sgo=on:sio=off:sos=on:spl=backtracking:ssec=off:updr=off_41 on theorems\conversions\codi_down_sum_theorems_goal15.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 3128
Passive clauses: 15412
Generated clauses: 68695
Final active clauses: 3079
Final passive clauses: 11362
Input formulas: 78
Initial clauses: 122

Duplicate literals: 770
Fw subsumption resolutions: 5138
Fw demodulations: 544
Condensations: 578

Simple tautologies: 927
Forward subsumptions: 45946

Binary resolution: 56339
Factoring: 402
Forward superposition: 4889

Backtracking splits: 29
Backtracking splits refuted: 15
Backtracking splits refuted at zero level: 3

Memory used: 13432KB
Time elapsed: 4.203 s
------------------------------
% remaining time: 5780 next slice time: 237
lrs+1_2:1_bs=off:ep=on:gsp=input_only:lcm=predicate:n=on:nicw=on:nwc=1.1:ptb=off:sac=on:sgo=on:sio=off:sos=on:sp=occurrence:spl=backtracking:spo=on:ssec=off:stl=81:updr=off_237 on theorems\conversions\codi_down_sum_theorems_goal15.all
Refutation found. Thanks to Tanya!
9376. $false (6:0) [subsumption resolution 9369,266]
266. ~po(sK2,sum(sK0,sK1)) (0:5) [cnf transformation 246]
246. eqdim(sK0,sK1) & ~zex(sK0) & ~zex(sK1) & ~zex(sK2) & po(sK2,sK0) & ~po(sK2,sum(sK0,sK1))[skolemisation 147]
147. ? [X0,X1,X2] : (eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & ~zex(X2) & po(X2,X0) & ~po(X2,sum(X0,X1)))[flattening 146]
146. ? [X0,X1,X2] : ((eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & ~zex(X2) & po(X2,X0)) & ~po(X2,sum(X0,X1)))[ennf transformation 79]
79. ~! [X0,X1,X2] : (~(eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & ~zex(X2) & po(X2,X0)) | po(X2,sum(X0,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : (~(eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & ~zex(X2) & po(X2,X0)) | po(X2,sum(X0,X1)))[negated conjecture]
9369. po(sK2,sum(sK0,sK1)) (6:5) [resolution 8323,314]
314. ~po(X0,X1) | po(X1,X0) (0:6) [cnf transformation 174]
174. ! [X0,X1] : (~po(X0,X1) | po(X1,X0))[ennf transformation 100]
100. ! [X0,X1] : (~po(X0,X1) | po(X1,X0))[flattening 3]
3. ! [X0] : ! [X1] : (~po(X0,X1) | po(X1,X0))[input]
8323. po(sum(sK0,sK1),sK2) (5:5) [resolution 4132,486]
486. ~p(sK4(sK0,sK2),X1) | po(X1,sK2) (2:8) [resolution 412,290]
290. ~p(X3,X1) | po(X0,X1) | ~p(X3,X0) (0:9) [cnf transformation 248]
248. ! [X0,X1] : ((~po(X0,X1) | (p(sK4(X1,X0),X0) & p(sK4(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[skolemisation 162]
162. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[ennf transformation 86]
86. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[flattening 85]
85. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[rectify 47]
47. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X2] : (p(X2,X0) & p(X2,X1)) | po(X0,X1)))[input]
412. p(sK4(sK0,sK2),sK2) (1:5) [resolution 265,288]
288. ~po(X0,X1) | p(sK4(X1,X0),X0) (0:8) [cnf transformation 248]
265. po(sK2,sK0) (0:3) [cnf transformation 246]
4132. p(sK4(sK0,sK2),sum(sK0,sK1)) (4:7) [subsumption resolution 4120,875]
875. eqdim(sK4(sK0,sK2),sum(sK0,sK1)) (3:7) [resolution 464,490]
490. eqdim(sK4(sK0,sK2),sK0) (2:5) [resolution 413,298]
298. ~p(X0,X1) | eqdim(X0,X1) (0:6) [cnf transformation 168]
168. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 167]
167. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 89]
89. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 68]
68. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
413. p(sK4(sK0,sK2),sK0) (1:5) [resolution 265,289]
289. ~po(X0,X1) | p(sK4(X1,X0),X1) (0:8) [cnf transformation 248]
464. ~eqdim(X0,sK0) | eqdim(X0,sum(sK0,sK1)) (2:8) [resolution 407,373]
373. ~eqdim(X1,X2) | eqdim(X0,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 236]
236. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~eqdim(X1,X2) | eqdim(X0,X2))[flattening 235]
235. ! [X0,X1,X2] : ((~eqdim(X0,X1) | ~eqdim(X1,X2)) | eqdim(X0,X2))[ennf transformation 141]
141. ! [X0,X1,X2] : (~(eqdim(X0,X1) & eqdim(X1,X2)) | eqdim(X0,X2))[flattening 54]
54. ! [X0] : ! [X1] : ! [X2] : (~(eqdim(X0,X1) & eqdim(X1,X2)) | eqdim(X0,X2))[input]
407. eqdim(sK0,sum(sK0,sK1)) (1:5) [subsumption resolution 406,262]
262. ~zex(sK0) (0:2) [cnf transformation 246]
406. zex(sK0) | eqdim(sK0,sum(sK0,sK1)) (1:7) [subsumption resolution 400,263]
263. ~zex(sK1) (0:2) [cnf transformation 246]
400. zex(sK1) | zex(sK0) | eqdim(sK0,sum(sK0,sK1)) (1:9) [resolution 261,334]
334. ~eqdim(X0,X1) | zex(X1) | zex(X0) | eqdim(X0,sum(X0,X1)) (0:12) [cnf transformation 200]
200. ! [X0,X1] : (~eqdim(X0,X1) | zex(X0) | zex(X1) | eqdim(X0,sum(X0,X1)))[flattening 199]
199. ! [X0,X1] : ((~eqdim(X0,X1) | zex(X0) | zex(X1)) | eqdim(X0,sum(X0,X1)))[ennf transformation 119]
119. ! [X0,X1] : (~(eqdim(X0,X1) & ~zex(X0) & ~zex(X1)) | eqdim(X0,sum(X0,X1)))[flattening 11]
11. ! [X0] : ! [X1] : (~(eqdim(X0,X1) & ~zex(X0) & ~zex(X1)) | eqdim(X0,sum(X0,X1)))[input]
261. eqdim(sK0,sK1) (0:3) [cnf transformation 246]
4120. ~eqdim(sK4(sK0,sK2),sum(sK0,sK1)) | p(sK4(sK0,sK2),sum(sK0,sK1)) (4:14) [resolution 835,299]
299. ~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1) (0:9) [cnf transformation 168]
835. cont(sK4(sK0,sK2),sum(sK0,sK1)) (3:7) [resolution 444,489]
489. cont(sK4(sK0,sK2),sK0) (2:5) [resolution 413,297]
297. ~p(X0,X1) | cont(X0,X1) (0:6) [cnf transformation 168]
444. ~cont(X2,sK0) | cont(X2,sum(sK0,sK1)) (2:8) [forward demodulation 441,279]
279. sum(X0,X1) = sum(X1,X0) (0:7) [cnf transformation 80]
80. ! [X0,X1] : sum(X0,X1) = sum(X1,X0)[flattening 12]
12. ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)[input]
441. ~cont(X2,sK0) | cont(X2,sum(sK1,sK0)) (2:8) [resolution 403,370]
370. ~leq(X0,X1) | ~cont(X2,X1) | cont(X2,sum(X0,X1)) (0:11) [cnf transformation 230]
230. ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X1) | cont(X2,sum(X0,X1)))[flattening 229]
229. ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X1)) | cont(X2,sum(X0,X1)))[ennf transformation 138]
138. ! [X0,X1,X2] : (~(leq(X0,X1) & cont(X2,X1)) | cont(X2,sum(X0,X1)))[flattening 14]
14. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & cont(X2,X1)) | cont(X2,sum(X0,X1)))[input]
403. leq(sK1,sK0) (1:3) [resolution 261,295]
295. ~eqdim(X0,X1) | leq(X1,X0) (0:6) [cnf transformation 166]
166. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 165]
165. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[ennf transformation 88]
88. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[flattening 76]
76. ! [X0] : ! [X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1002
Passive clauses: 4730
Generated clauses: 10122
Final active clauses: 654
Final passive clauses: 1901
Input formulas: 78
Initial clauses: 122

Duplicate literals: 16
Fw subsumption resolutions: 1076
Fw demodulations: 649
Bw demodulations: 1081

Simple tautologies: 71
Forward subsumptions: 3541

Binary resolution: 5974
Forward superposition: 42

Backtracking splits: 66
Backtracking splits refuted: 38
Backtracking splits refuted at zero level: 16

Memory used: 2046KB
Time elapsed: 0.350 s
------------------------------
% Success in time 22.348 s
