Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 1928
Passive clauses: 20634
Generated clauses: 35165
Input formulas: 58
Initial clauses: 116

Selected by SInE selection: 58
SInE iterations: 1
Splitted inequalities: 1

Duplicate literals: 155
Fw subsumption resolutions: 1366

Simple tautologies: 81
Forward subsumptions: 12328

Binary resolution: 32221
Factoring: 68

Backtracking splits: 698
Backtracking splits refuted: 227
Backtracking splits refuted at zero level: 44

Memory used: 6780KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f7032,plain,(
  $false),
  inference(resolution,[],[f7003,f411])).
fof(f411,plain,(
  zex(sK23)),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  zex(sK23)),
  inference(skolemisation,[status(esa)],[f165])).
fof(f165,plain,(
  ? [X0] : zex(X0)),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ? [X51] : zex(X51)),
  file(unknown,unknown)).
fof(f7003,plain,(
  ( ! [X6] : (~zex(X6)) )),
  inference(subsumption_resolution,[],[f6986,f842])).
fof(f842,plain,(
  ( ! [X1] : (~zex(sum(X1,sK0)) | ~zex(X1)) )),
  inference(superposition,[],[f827,f317])).
fof(f317,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = sum(X1,X0)) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : sum(X0,X1) = sum(X1,X0)),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)),
  inference(rectify,[],[f2])).
fof(f2,axiom,(
  ! [X2] : ! [X3] : sum(X2,X3) = sum(X3,X2)),
  file(unknown,unknown)).
fof(f827,plain,(
  ( ! [X0] : (~zex(sum(sK0,X0)) | ~zex(X0)) )),
  inference(subsumption_resolution,[],[f540,f659])).
fof(f659,plain,(
  ( ! [X0] : (zex(sK0) | ~zex(X0)) )),
  inference(duplicate_literal_removal,[],[f654])).
fof(f654,plain,(
  ( ! [X0] : (~zex(X0) | zex(sK0) | ~zex(X0)) )),
  inference(resolution,[],[f633,f328])).
fof(f328,plain,(
  ( ! [X0,X1] : (less(X0,X1) | zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : (~zex(X0) | zex(X1) | less(X0,X1))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : ((~zex(X0) | zex(X1)) | less(X0,X1))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & ~zex(X1)) => less(X0,X1))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X73] : ! [X74] : ((zex(X73) & ~zex(X74)) => less(X73,X74))),
  file(unknown,unknown)).
fof(f633,plain,(
  ( ! [X6] : (~less(X6,sK0) | ~zex(X6)) )),
  inference(subsumption_resolution,[],[f628,f413])).
fof(f413,plain,(
  sP1(sK0)),
  inference(inequality_splitting,[],[f301,f412])).
fof(f412,plain,(
  ~sP1(sum(sK0,sK1))),
  introduced(inequality_splitting_name_introduction,[])).
fof(f301,plain,(
  sum(sK0,sK1) != sK0),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  zex(sK1) & sum(sK0,sK1) != sK0),
  inference(skolemisation,[status(esa)],[f166])).
fof(f166,plain,(
  ? [X0,X1] : (zex(X1) & sum(X0,X1) != X0)),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ~! [X0,X1] : (zex(X1) => sum(X0,X1) = X0)),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : (zex(X1) => sum(X0,X1) = X0)),
  file(unknown,unknown)).
fof(f628,plain,(
  ( ! [X6] : (~sP1(sK0) | ~zex(X6) | ~less(X6,sK0)) )),
  inference(superposition,[],[f545,f324])).
fof(f324,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = X1 | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : (~less(X0,X1) | sum(X0,X1) = X1)),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : (less(X0,X1) => sum(X0,X1) = X1)),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => sum(X0,X1) = X1)),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X4] : ! [X5] : (less(X4,X5) => sum(X4,X5) = X5)),
  file(unknown,unknown)).
fof(f545,plain,(
  ( ! [X0] : (~sP1(sum(X0,sK0)) | ~zex(X0)) )),
  inference(superposition,[],[f526,f317])).
fof(f526,plain,(
  ( ! [X0] : (~sP1(sum(sK0,X0)) | ~zex(X0)) )),
  inference(subsumption_resolution,[],[f482,f300])).
fof(f300,plain,(
  zex(sK1)),
  inference(cnf_transformation,[],[f237])).
fof(f482,plain,(
  ( ! [X0] : (~sP1(sum(sK0,X0)) | ~zex(X0) | ~zex(sK1)) )),
  inference(superposition,[],[f412,f456])).
fof(f456,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0)) )),
  inference(resolution,[],[f414,f423])).
fof(f423,plain,(
  ( ! [X0,X1] : (~sQ1_eqProxy(X0,X1) | X0 = X1) )),
  introduced(equality_proxy_axiom,[])).
fof(f414,plain,(
  ( ! [X0,X1] : (sQ1_eqProxy(X0,X1) | ~zex(X1) | ~zex(X0)) )),
  inference(equality_proxy_replacement,[],[f327])).
fof(f327,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => X0 = X1)),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X71] : ! [X72] : ((zex(X71) & zex(X72)) => X71 = X72)),
  file(unknown,unknown)).
fof(f540,plain,(
  ( ! [X0] : (~zex(X0) | ~zex(sum(sK0,X0)) | ~zex(sK0)) )),
  inference(resolution,[],[f526,f481])).
fof(f481,plain,(
  ( ! [X0] : (sP1(X0) | ~zex(X0) | ~zex(sK0)) )),
  inference(superposition,[],[f413,f456])).
fof(f6986,plain,(
  ( ! [X6] : (zex(sum(X6,sK0)) | ~zex(X6)) )),
  inference(resolution,[],[f6980,f328])).
fof(f6980,plain,(
  ( ! [X37] : (~less(X37,sum(X37,sK0))) )),
  inference(subsumption_resolution,[],[f6968,f4283])).
fof(f4283,plain,(
  ( ! [X2,X3] : (less(sK1,X3) | ~less(X2,X3)) )),
  inference(subsumption_resolution,[],[f4276,f3459])).
fof(f3459,plain,(
  ( ! [X16] : (~less(X16,sK1)) )),
  inference(resolution,[],[f3173,f632])).
fof(f632,plain,(
  ~less(sK0,sK1)),
  inference(subsumption_resolution,[],[f631,f300])).
fof(f631,plain,(
  ~zex(sK1) | ~less(sK0,sK1)),
  inference(subsumption_resolution,[],[f623,f454])).
fof(f454,plain,(
  ( ! [X4,X3] : (~less(X3,X4) | ~zex(X4) | zex(X3)) )),
  inference(resolution,[],[f328,f325])).
fof(f325,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : (~less(X0,X1) | ~less(X1,X0))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1] : (less(X0,X1) => ~less(X1,X0))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X66] : ! [X67] : (less(X66,X67) => ~less(X67,X66))),
  file(unknown,unknown)).
fof(f623,plain,(
  ~zex(sK1) | ~zex(sK0) | ~less(sK0,sK1)),
  inference(superposition,[],[f534,f324])).
fof(f534,plain,(
  ~zex(sum(sK0,sK1)) | ~zex(sK0)),
  inference(resolution,[],[f481,f412])).
fof(f3173,plain,(
  ( ! [X2,X3] : (less(sK0,X3) | ~less(X2,X3)) )),
  inference(subsumption_resolution,[],[f3166,f2001])).
fof(f2001,plain,(
  ( ! [X6] : (~less(X6,sK0)) )),
  inference(subsumption_resolution,[],[f2000,f300])).
fof(f2000,plain,(
  ( ! [X6] : (~less(X6,sK0) | ~zex(sK1)) )),
  inference(subsumption_resolution,[],[f1995,f633])).
fof(f1995,plain,(
  ( ! [X6] : (~less(X6,sK0) | zex(X6) | ~zex(sK1)) )),
  inference(resolution,[],[f1968,f328])).
fof(f1968,plain,(
  ( ! [X1] : (~less(sK1,X1) | ~less(X1,sK0)) )),
  inference(resolution,[],[f1954,f376])).
fof(f376,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X78] : ! [X79] : (leq(X78,X79) <=> (less(X78,X79) | eqdim(X78,X79)))),
  file(unknown,unknown)).
fof(f1954,plain,(
  ( ! [X2] : (~leq(X2,sK0) | ~less(sK1,X2)) )),
  inference(resolution,[],[f1949,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (less(X0,X2) | ~leq(X1,X2) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : (~less(X0,X1) | ~leq(X1,X2) | less(X0,X2))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((~less(X0,X1) | ~leq(X1,X2)) | less(X0,X2))),
  inference(ennf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((less(X0,X1) & leq(X1,X2)) => less(X0,X2))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X68] : ! [X69] : ! [X70] : ((less(X68,X69) & leq(X69,X70)) => less(X68,X70))),
  file(unknown,unknown)).
fof(f1949,plain,(
  ~less(sK1,sK0)),
  inference(subsumption_resolution,[],[f1927,f413])).
fof(f1927,plain,(
  ~sP1(sK0) | ~less(sK1,sK0)),
  inference(superposition,[],[f412,f620])).
fof(f620,plain,(
  ( ! [X2,X3] : (sum(X3,X2) = X3 | ~less(X2,X3)) )),
  inference(superposition,[],[f324,f317])).
fof(f3166,plain,(
  ( ! [X2,X3] : (~less(X2,X3) | less(sK0,X3) | less(X3,sK0)) )),
  inference(resolution,[],[f2197,f374])).
fof(f374,plain,(
  ( ! [X0,X1] : (eqdim(X0,X1) | less(X1,X0) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X76] : ! [X77] : (eqdim(X76,X77) <=> (~less(X76,X77) & ~less(X77,X76)))),
  file(unknown,unknown)).
fof(f2197,plain,(
  ( ! [X0,X1] : (~eqdim(X1,sK0) | ~less(X0,X1)) )),
  inference(resolution,[],[f2013,f377])).
fof(f377,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f283])).
fof(f2013,plain,(
  ( ! [X4,X5] : (~leq(X4,sK0) | ~less(X5,X4)) )),
  inference(resolution,[],[f2001,f400])).
fof(f4276,plain,(
  ( ! [X2,X3] : (~less(X2,X3) | less(sK1,X3) | less(X3,sK1)) )),
  inference(resolution,[],[f4075,f374])).
fof(f4075,plain,(
  ( ! [X0,X1] : (~eqdim(X1,sK1) | ~less(X0,X1)) )),
  inference(resolution,[],[f3519,f377])).
fof(f3519,plain,(
  ( ! [X4,X5] : (~leq(X4,sK1) | ~less(X5,X4)) )),
  inference(resolution,[],[f3459,f400])).
fof(f6968,plain,(
  ( ! [X37] : (~less(X37,sum(X37,sK0)) | ~less(sK1,sum(X37,sK0))) )),
  inference(resolution,[],[f4631,f1969])).
fof(f1969,plain,(
  ( ! [X2] : (~cont(X2,sK0) | ~less(sK1,X2)) )),
  inference(resolution,[],[f1954,f326])).
fof(f326,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : (~cont(X0,X1) | leq(X0,X1))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1] : (cont(X0,X1) => leq(X0,X1))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X49] : ! [X50] : (cont(X49,X50) => leq(X49,X50))),
  file(unknown,unknown)).
fof(f4631,plain,(
  ( ! [X6,X7] : (cont(sum(X6,X7),X7) | ~less(X6,sum(X6,X7))) )),
  inference(subsumption_resolution,[],[f4594,f675])).
fof(f675,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~zex(X1)) )),
  inference(subsumption_resolution,[],[f674,f454])).
fof(f674,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~zex(X1) | ~zex(X0)) )),
  inference(subsumption_resolution,[],[f672,f659])).
fof(f672,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~zex(X1) | ~zex(X0) | ~zex(sK0)) )),
  inference(superposition,[],[f640,f456])).
fof(f640,plain,(
  ( ! [X0] : (~less(sK0,X0) | ~zex(X0)) )),
  inference(subsumption_resolution,[],[f636,f300])).
fof(f636,plain,(
  ( ! [X0] : (~less(sK0,X0) | ~zex(X0) | ~zex(sK1)) )),
  inference(superposition,[],[f632,f456])).
fof(f4594,plain,(
  ( ! [X6,X7] : (cont(sum(X6,X7),X7) | ~less(X6,sum(X6,X7)) | zex(sum(X6,X7))) )),
  inference(resolution,[],[f4540,f314])).
fof(f314,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X134] : (~zex(X134) <=> cont(X134,X134))),
  file(unknown,unknown)).
fof(f4540,plain,(
  ( ! [X2,X0,X1] : (~cont(X0,sum(X1,X2)) | cont(X0,X2) | ~less(X1,X0)) )),
  inference(subsumption_resolution,[],[f1313,f617])).
fof(f617,plain,(
  ( ! [X4,X5] : (~cont(X4,X5) | ~less(X5,X4)) )),
  inference(subsumption_resolution,[],[f615,f319])).
fof(f319,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X140] : ! [X141] : (zex(X140) => (~cont(X141,X140) & ~cont(X140,X141)))),
  file(unknown,unknown)).
fof(f615,plain,(
  ( ! [X4,X5] : (zex(X4) | ~cont(X4,X5) | ~less(X5,X4)) )),
  inference(superposition,[],[f392,f323])).
fof(f323,plain,(
  ( ! [X0,X1] : (difference(X0,X1) = X0 | ~less(X1,X0)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : (~less(X1,X0) | difference(X0,X1) = X0)),
  inference(ennf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X14] : ! [X15] : (less(X15,X14) => difference(X14,X15) = X14)),
  file(unknown,unknown)).
fof(f392,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X25] : ! [X26] : (zex(difference(X25,X26)) <=> (zex(X25) | cont(X25,X26)))),
  file(unknown,unknown)).
fof(f1313,plain,(
  ( ! [X2,X0,X1] : (cont(X0,X2) | cont(X0,X1) | ~cont(X0,sum(X1,X2)) | ~less(X1,X0)) )),
  inference(superposition,[],[f405,f323])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : (~cont(X2,sum(X0,X1)) | cont(X2,X0) | cont(difference(X2,X0),X1))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((~cont(X2,sum(X0,X1)) | cont(X2,X0)) | cont(difference(X2,X0),X1))),
  inference(ennf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X2,sum(X0,X1)) & ~cont(X2,X0)) => cont(difference(X2,X0),X1))),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X9] : ! [X10] : ! [X11] : ((cont(X11,sum(X9,X10)) & ~cont(X11,X9)) => cont(difference(X11,X9),X10))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 515
Passive clauses: 2960
Generated clauses: 6699
Final active clauses: 515
Final passive clauses: 2230
Input formulas: 58
Initial clauses: 115

Selected by SInE selection: 58
SInE iterations: 1
Splitted inequalities: 1

Duplicate literals: 101
Fw subsumption resolutions: 858
Fw demodulations: 1

Simple tautologies: 173
Equational tautologies: 9
Forward subsumptions: 2817

Binary resolution: 2863
Factoring: 18
Forward superposition: 2182
Backward superposition: 592

Unique components: 3

Memory used: 1279KB
Time elapsed: 0.095 s
------------------------------
% Success in time 0.533 s
