Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal23.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 807
Passive clauses: 4286
Generated clauses: 10612
Final active clauses: 791
Final passive clauses: 3189
Input formulas: 50
Initial clauses: 86

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 80
Fw subsumption resolutions: 481

Simple tautologies: 58
Forward subsumptions: 5819

Binary resolution: 9847
Factoring: 28

Backtracking splits: 119
Backtracking splits refuted: 12

Memory used: 1918KB
Time elapsed: 0.403 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal23.all
Refutation found. Thanks to Tanya!
1723. $false (11:0) [subsumption resolution 1715,453]
453. ~zex(intersection(sK0,sK0)) (2:4) [resolution 447,304]
304. c(sK1,intersection(sK0,sK0)) (1:5) [resolution 187,173]
173. ~zex(intersection(sK1,intersection(sK0,sK0))) (0:6) [cnf transformation 162]
162. zex(intersection(sK0,intersection(sK0,sK1))) & ~zex(intersection(sK1,intersection(sK0,sK0)))[skolemisation 95]
95. ? [X0,X1] : (zex(intersection(X0,intersection(X0,X1))) & ~zex(intersection(X1,intersection(X0,X0))))[ennf transformation 51]
51. ~! [X0,X1] : (~zex(intersection(X0,intersection(X0,X1))) | zex(intersection(X1,intersection(X0,X0))))[flattening 1]
1. ~! [X0] : ! [X1] : (~zex(intersection(X0,intersection(X0,X1))) | zex(intersection(X1,intersection(X0,X0))))[negated conjecture]
187. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 102]
102. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[ennf transformation 54]
54. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[flattening 11]
11. ! [X0] : ! [X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[input]
447. ~c(X2,X3) | ~zex(X3) (1:5) [resolution 199,213]
213. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 116]
116. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 67]
67. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[flattening 31]
31. ! [X0] : ! [X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[input]
199. cont(sK5(X1,X0),X1) | ~c(X0,X1) (0:8) [cnf transformation 166]
166. ! [X0,X1] : ((~c(X0,X1) | (cont(sK5(X1,X0),X0) & cont(sK5(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 108]
108. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 60]
60. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 59]
59. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 34]
34. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
1715. zex(intersection(sK0,sK0)) (11:4) [resolution 1693,211]
211. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 65]
65. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 12]
12. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
1693. ~cont(intersection(sK0,sK0),sK0) (10:5) [resolution 1474,173]
1474. zex(intersection(sK1,X0)) | ~cont(X0,sK0) (9:7) [resolution 475,785]
785. ~cont(intersection(sK1,X21),X20) | ~cont(X20,sK0) (8:8) [resolution 247,575]
575. ~cont(intersection(sK1,X5),sK0) (7:5) [subsumption resolution 571,214]
214. ~cont(X0,X1) | ~zex(X0) (0:5) [cnf transformation 116]
571. zex(intersection(sK1,X5)) | ~cont(intersection(sK1,X5),sK0) (7:9) [resolution 211,528]
528. ~cont(X1,sK1) | ~cont(X1,sK0) (6:6) [resolution 518,200]
200. c(X0,X1) | ~cont(X3,X1) | ~cont(X3,X0) (0:9) [cnf transformation 166]
518. ~c(sK0,sK1) (5:3) [resolution 509,188]
188. ~zex(intersection(X0,X1)) | ~c(X0,X1) (0:7) [cnf transformation 102]
509. zex(intersection(sK0,sK1)) (4:4) [resolution 495,185]
185. p(X0,X0) | zex(X0) (0:5) [cnf transformation 2]
2. ! [X0] : (zex(X0) | p(X0,X0))[input]
495. ~p(X0,intersection(sK0,sK1)) (3:5) [subsumption resolution 485,235]
235. ~p(X2,intersection(X0,X1)) | cont(X2,X0) (0:8) [cnf transformation 137]
137. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 136]
136. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[ennf transformation 81]
81. ! [X0,X1,X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 14]
14. ! [X0] : ! [X1] : ! [X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[input]
485. ~cont(X0,sK0) | ~p(X0,intersection(sK0,sK1)) (3:8) [resolution 481,201]
201. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 110]
110. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 109]
109. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 61]
61. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 35]
35. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
481. ~cont(X0,intersection(sK0,sK1)) | ~cont(X0,sK0) (2:8) [resolution 200,305]
305. ~c(sK0,intersection(sK0,sK1)) (1:5) [resolution 188,172]
172. zex(intersection(sK0,intersection(sK0,sK1))) (0:6) [cnf transformation 162]
247. cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1) (0:9) [cnf transformation 154]
154. ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))[flattening 153]
153. ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))[ennf transformation 91]
91. ! [X0,X1,X2] : (~(cont(X0,X1) & cont(X1,X2)) | cont(X0,X2))[flattening 30]
30. ! [X0] : ! [X1] : ! [X2] : (~(cont(X0,X1) & cont(X1,X2)) | cont(X0,X2))[input]
475. cont(intersection(X0,X1),X1) | zex(intersection(X0,X1)) (1:9) [resolution 236,185]
236. ~p(X2,intersection(X0,X1)) | cont(X2,X1) (0:8) [cnf transformation 137]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 241
Passive clauses: 673
Generated clauses: 1539
Final active clauses: 241
Final passive clauses: 391
Input formulas: 50
Initial clauses: 85

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 34
Fw subsumption resolutions: 104

Simple tautologies: 55
Equational tautologies: 1
Forward subsumptions: 705

Binary resolution: 740
Factoring: 20
Forward superposition: 504
Backward superposition: 52

Unique components: 2

Memory used: 383KB
Time elapsed: 0.079 s
------------------------------
% Success in time 0.562 s
