Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal18.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 889
Passive clauses: 4285
Generated clauses: 13093
Final active clauses: 877
Final passive clauses: 3189
Input formulas: 44
Initial clauses: 92

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 74
Fw subsumption resolutions: 168

Simple tautologies: 41
Forward subsumptions: 8554

Binary resolution: 12492
Factoring: 55

Backtracking splits: 154
Backtracking splits refuted: 20

Memory used: 1791KB
Time elapsed: 0.404 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal18.all
Refutation found. Thanks to Tanya!
1809. $false (7:0) [resolution 1806,275]
275. c(intersection(sK0,sK0),sK1) (1:5) [resolution 176,163]
163. ~zex(intersection(intersection(sK0,sK0),sK1)) (0:6) [cnf transformation 146]
146. zex(intersection(sK0,sK1)) & ~zex(intersection(intersection(sK0,sK0),sK1))[skolemisation 87]
87. ? [X0,X1] : (zex(intersection(X0,X1)) & ~zex(intersection(intersection(X0,X0),X1)))[ennf transformation 45]
45. ~! [X0,X1] : (~zex(intersection(X0,X1)) | zex(intersection(intersection(X0,X0),X1)))[flattening 1]
1. ~! [X0] : ! [X1] : (~zex(intersection(X0,X1)) | zex(intersection(intersection(X0,X0),X1)))[negated conjecture]
176. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 94]
94. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[ennf transformation 48]
48. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[flattening 2]
2. ! [X0] : ! [X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[input]
1806. ~c(intersection(sK0,X3),sK1) (6:5) [subsumption resolution 1794,454]
454. ~c(X2,X3) | ~zex(X2) (1:5) [resolution 205,220]
220. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 112]
112. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 67]
67. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[flattening 25]
25. ! [X0] : ! [X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[input]
205. cont(sK10(X1,X0),X0) | ~c(X0,X1) (0:8) [cnf transformation 157]
157. ! [X0,X1] : ((~c(X0,X1) | (cont(sK10(X1,X0),X0) & cont(sK10(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 104]
104. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 60]
60. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 59]
59. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 28]
28. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
1794. ~c(intersection(sK0,X3),sK1) | zex(intersection(sK0,X3)) (6:9) [resolution 1790,218]
218. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 65]
65. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 3]
3. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
1790. ~cont(X0,sK0) | ~c(X0,sK1) (5:6) [duplicate literal removal 1780]
1780. ~cont(X0,sK0) | ~c(X0,sK1) | ~c(X0,sK1) (5:9) [resolution 749,205]
749. ~cont(sK10(sK1,X28),X27) | ~cont(X27,sK0) | ~c(X28,sK1) (4:11) [resolution 245,521]
521. ~cont(sK10(sK1,X5),sK0) | ~c(X5,sK1) (3:8) [resolution 510,206]
206. cont(sK10(X1,X0),X1) | ~c(X0,X1) (0:8) [cnf transformation 157]
510. ~cont(X0,sK1) | ~cont(X0,sK0) (2:6) [resolution 207,276]
276. ~c(sK0,sK1) (1:3) [resolution 177,162]
162. zex(intersection(sK0,sK1)) (0:4) [cnf transformation 146]
177. ~zex(intersection(X0,X1)) | ~c(X0,X1) (0:7) [cnf transformation 94]
207. c(X0,X1) | ~cont(X3,X1) | ~cont(X3,X0) (0:9) [cnf transformation 157]
245. cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1) (0:9) [cnf transformation 140]
140. ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))[flattening 139]
139. ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))[ennf transformation 85]
85. ! [X0,X1,X2] : (~(cont(X0,X1) & cont(X1,X2)) | cont(X0,X2))[flattening 24]
24. ! [X0] : ! [X1] : ! [X2] : (~(cont(X0,X1) & cont(X1,X2)) | cont(X0,X2))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 270
Passive clauses: 785
Generated clauses: 1647
Final active clauses: 270
Final passive clauses: 470
Input formulas: 44
Initial clauses: 91

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 70
Fw subsumption resolutions: 95

Simple tautologies: 41
Equational tautologies: 3
Forward subsumptions: 692

Binary resolution: 783
Factoring: 16
Forward superposition: 546
Backward superposition: 48

Unique components: 2

Memory used: 383KB
Time elapsed: 0.075 s
------------------------------
% Success in time 0.559 s
