Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal14.all
Refutation found. Thanks to Tanya!
9297. $false (6:0) [subsumption resolution 9296,265]
265. cont(sK0,sK0) (1:3) [resolution 167,162]
162. ~zex(sK0) (0:2) [cnf transformation 146]
146. ~zex(sK0) & ~zex(intersection(sK0,sK0)) & ~p(sK0,intersection(sK0,sK0))[skolemisation 87]
87. ? [X0] : (~zex(X0) & ~zex(intersection(X0,X0)) & ~p(X0,intersection(X0,X0)))[flattening 86]
86. ? [X0] : ((~zex(X0) & ~zex(intersection(X0,X0))) & ~p(X0,intersection(X0,X0)))[ennf transformation 1]
1. ~! [X0] : (~(~zex(X0) & ~zex(intersection(X0,X0))) | p(X0,intersection(X0,X0)))[negated conjecture]
167. zex(X0) | cont(X0,X0) (0:5) [cnf transformation 88]
88. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[ennf transformation 22]
22. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[input]
9296. ~cont(sK0,sK0) (6:3) [duplicate literal removal 9283]
9283. ~cont(sK0,sK0) | ~cont(sK0,sK0) (6:6) [resolution 9278,247]
247. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 142]
142. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 141]
141. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 85]
85. ! [X0,X1,X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[flattening 4]
4. ! [X0] : ! [X1] : ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[input]
9278. ~leq(sK0,intersection(sK0,sK0)) (5:5) [subsumption resolution 9273,596]
596. cont(intersection(sK0,sK0),sK0) (1:5) [resolution 219,163]
163. ~zex(intersection(sK0,sK0)) (0:4) [cnf transformation 146]
219. zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0) (0:9) [cnf transformation 64]
64. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 3]
3. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
9273. ~cont(intersection(sK0,sK0),sK0) | ~leq(sK0,intersection(sK0,sK0)) (5:10) [resolution 9150,623]
623. eqdim(X8,X9) | ~cont(X8,X9) | ~leq(X9,X8) (1:9) [resolution 226,216]
216. ~lt(X0,X1) | ~leq(X1,X0) (0:6) [cnf transformation 110]
110. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 109]
109. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[ennf transformation 62]
62. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[flattening 40]
40. ! [X0] : ! [X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[input]
226. lt(X0,X1) | eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 116]
116. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[ennf transformation 70]
70. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 30]
30. ! [X0] : ! [X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[input]
9150. ~eqdim(intersection(sK0,sK0),sK0) (4:5) [resolution 9141,223]
223. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 113]
113. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 67]
67. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[flattening 14]
14. ! [X0] : ! [X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[input]
9141. ~eqdim(sK0,intersection(sK0,sK0)) (3:5) [resolution 8953,165]
165. eqdim(X0,X0) (0:3) [cnf transformation 13]
13. ! [X0] : eqdim(X0,X0)[input]
8953. ~eqdim(X77,intersection(sK0,sK0)) | ~eqdim(sK0,X77) (2:8) [subsumption resolution 8952,265]
8952. ~eqdim(sK0,X77) | ~eqdim(X77,intersection(sK0,sK0)) | ~cont(sK0,sK0) (2:11) [duplicate literal removal 8924]
8924. ~cont(sK0,sK0) | ~cont(sK0,sK0) | ~eqdim(X77,intersection(sK0,sK0)) | ~eqdim(sK0,X77) (2:14) [resolution 1519,164]
164. ~p(sK0,intersection(sK0,sK0)) (0:5) [cnf transformation 146]
1519. p(X2,intersection(X3,X4)) | ~cont(X2,X4) | ~cont(X2,X3) | ~eqdim(X5,intersection(X3,X4)) | ~eqdim(X2,X5) (1:19) [resolution 237,243]
243. eqdim(X0,X2) | ~eqdim(X1,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 134]
134. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~eqdim(X1,X2) | eqdim(X0,X2))[flattening 133]
133. ! [X0,X1,X2] : ((~eqdim(X0,X1) | ~eqdim(X1,X2)) | eqdim(X0,X2))[ennf transformation 81]
81. ! [X0,X1,X2] : (~(eqdim(X0,X1) & eqdim(X1,X2)) | eqdim(X0,X2))[flattening 15]
15. ! [X0] : ! [X1] : ! [X2] : (~(eqdim(X0,X1) & eqdim(X1,X2)) | eqdim(X0,X2))[input]
237. ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:16) [cnf transformation 130]
130. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 129]
129. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[ennf transformation 78]
78. ! [X0,X1,X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 5]
5. ! [X0] : ! [X1] : ! [X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 816
Passive clauses: 3707
Generated clauses: 9131
Final active clauses: 813
Final passive clauses: 2732
Input formulas: 44
Initial clauses: 93

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 63
Fw subsumption resolutions: 147

Simple tautologies: 34
Forward subsumptions: 5252

Binary resolution: 8716
Factoring: 37

Backtracking splits: 56
Backtracking splits refuted: 8

Memory used: 1663KB
Time elapsed: 0.309 s
------------------------------
% Success in time 0.35 s
