Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_down_sum_theorems_goal16.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1133
Passive clauses: 10655
Generated clauses: 15690
Final active clauses: 1125
Final passive clauses: 9320
Input formulas: 78
Initial clauses: 123

Selected by SInE selection: 78
SInE iterations: 1

Duplicate literals: 46
Fw subsumption resolutions: 339

Simple tautologies: 49
Forward subsumptions: 4637

Binary resolution: 15038
Factoring: 29

Backtracking splits: 118
Backtracking splits refuted: 9
Backtracking splits refuted at zero level: 7

Memory used: 3198KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_down_sum_theorems_goal16.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1487
Passive clauses: 12530
Generated clauses: 34815
Final active clauses: 1487
Final passive clauses: 10548
Input formulas: 78
Initial clauses: 122

Selected by SInE selection: 78
SInE iterations: 1

Duplicate literals: 573
Fw subsumption resolutions: 1719
Fw demodulations: 11

Simple tautologies: 865
Equational tautologies: 12
Forward subsumptions: 19635

Binary resolution: 25389
Factoring: 148
Forward superposition: 8244
Backward superposition: 523
Self superposition: 2

Split clauses: 1
Split components: 2
Unique components: 7
Introduced splitting names: 1

Memory used: 7036KB
Time elapsed: 1.402 s
------------------------------
% remaining time: 5981 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on theorems\conversions\codi_down_sum_theorems_goal16.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 868
Passive clauses: 14765
Generated clauses: 142321
Final active clauses: 262
Final passive clauses: 795
Input formulas: 78
Initial clauses: 120

Duplicate literals: 131
Trivial inequalities: 490
Bw subsumption resolutions: 6377
Fw demodulations: 80655
Condensations: 131

Simple tautologies: 281
Equational tautologies: 53
Forward subsumptions: 40808

Binary resolution: 7138
Forward superposition: 731
Backward superposition: 1882

Backtracking splits: 1237
Backtracking splits refuted: 667
Backtracking splits refuted at zero level: 35

Memory used: 14200KB
Time elapsed: 3.301 s
------------------------------
% remaining time: 5948 next slice time: 95
dis+2_4_bs=off:ep=on:n=on:nicw=on:nwc=1.5:ptb=off:sac=on:sio=off:spl=backtracking:ssec=off_95 on theorems\conversions\codi_down_sum_theorems_goal16.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 3590
Passive clauses: 92784
Generated clauses: 227691
Final active clauses: 3412
Final passive clauses: 86781
Input formulas: 78
Initial clauses: 122

Duplicate literals: 2826
Fw subsumption resolutions: 15318
Fw demodulations: 6927
Bw demodulations: 171

Simple tautologies: 2095
Equational tautologies: 76
Forward subsumptions: 110192

Binary resolution: 191292
Factoring: 51
Forward superposition: 9592
Backward superposition: 1708

Backtracking splits: 27
Backtracking splits refuted: 7
Backtracking splits refuted at zero level: 4

Memory used: 43751KB
Time elapsed: 9.615 s
------------------------------
% remaining time: 5851 next slice time: 4
lrs-1003_12_bms=on:bs=unit_only:bsr=on:cond=fast:drc=off:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=10:sfv=off:sgo=on:sio=off:sp=occurrence:spo=on:ssec=off:stl=81:updr=off_4 on theorems\conversions\codi_down_sum_theorems_goal16.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 886
Passive clauses: 3656
Generated clauses: 9200
Final active clauses: 884
Final passive clauses: 2758
Input formulas: 78
Initial clauses: 122
Reactivated clauses: 13

Duplicate literals: 62
Fw subsumption resolutions: 321
Bw subsumption resolutions: 2
Fw demodulations: 445
Condensations: 23

Simple tautologies: 136
Equational tautologies: 1
Forward subsumptions: 4520
Backward subsumptions: 18

Binary resolution: 7275
Factoring: 81
Forward superposition: 423
Backward superposition: 451

Split clauses: 1
Split components: 2
Unique components: 3166
Introduced splitting names: 1

Memory used: 2942KB
Time elapsed: 0.502 s
------------------------------
% remaining time: 5846 next slice time: 22
lrs+1010_50_bs=off:drc=off:flr=on:lcm=predicate:n=on:nwc=1.2:sfv=off:sgo=on:stl=81_22 on theorems\conversions\codi_down_sum_theorems_goal16.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 2544
Passive clauses: 5782
Generated clauses: 82179
Final active clauses: 2544
Final passive clauses: 3238
Input formulas: 78
Initial clauses: 120

Duplicate literals: 1587
Fw subsumption resolutions: 6308
Fw literal rewrites: 12089

Simple tautologies: 1758
Equational tautologies: 4
Forward subsumptions: 54713

Binary resolution: 64459
Factoring: 22
Forward superposition: 5073
Backward superposition: 1126
Equality resolution: 1

Split clauses: 1
Split components: 2
Unique components: 7
Introduced splitting names: 1

Memory used: 10874KB
Time elapsed: 2.305 s
------------------------------
% remaining time: 5823 next slice time: 41
dis-1010_5_bd=off:bs=off:cond=fast:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=1.3:ptb=off:sac=on:sgo=on:sio=off:sos=on:spl=backtracking:ssec=off:updr=off_41 on theorems\conversions\codi_down_sum_theorems_goal16.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 3891
Passive clauses: 21550
Generated clauses: 87120
Final active clauses: 3835
Final passive clauses: 15843
Input formulas: 78
Initial clauses: 122

Duplicate literals: 1069
Fw subsumption resolutions: 7072
Fw demodulations: 5187
Condensations: 740

Simple tautologies: 794
Forward subsumptions: 51792

Binary resolution: 65000
Factoring: 411
Forward superposition: 5460
Backward superposition: 2132

Backtracking splits: 38
Backtracking splits refuted: 19
Backtracking splits refuted at zero level: 3

Memory used: 17014KB
Time elapsed: 4.206 s
------------------------------
% remaining time: 5780 next slice time: 237
lrs+1_2:1_bs=off:ep=on:gsp=input_only:lcm=predicate:n=on:nicw=on:nwc=1.1:ptb=off:sac=on:sgo=on:sio=off:sos=on:sp=occurrence:spl=backtracking:spo=on:ssec=off:stl=81:updr=off_237 on theorems\conversions\codi_down_sum_theorems_goal16.all
Refutation found. Thanks to Tanya!
19727. $false (6:0) [subsumption resolution 19723,266]
266. ~po(sK2,sum(sK0,sK1)) (0:5) [cnf transformation 246]
246. eqdim(sK0,sK1) & ~zex(sK0) & ~zex(sK1) & ~zex(sK2) & (po(sK2,sK0) | po(sK2,sK1)) & ~po(sK2,sum(sK0,sK1))[skolemisation 147]
147. ? [X0,X1,X2] : (eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & ~zex(X2) & (po(X2,X0) | po(X2,X1)) & ~po(X2,sum(X0,X1)))[flattening 146]
146. ? [X0,X1,X2] : ((eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & ~zex(X2) & (po(X2,X0) | po(X2,X1))) & ~po(X2,sum(X0,X1)))[ennf transformation 79]
79. ~! [X0,X1,X2] : (~(eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & ~zex(X2) & (po(X2,X0) | po(X2,X1))) | po(X2,sum(X0,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : (~(eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & ~zex(X2) & (po(X2,X0) | po(X2,X1))) | po(X2,sum(X0,X1)))[negated conjecture]
19723. po(sK2,sum(sK0,sK1)) (6:5) [resolution 17469,318]
318. ~po(X0,X1) | po(X1,X0) (0:6) [cnf transformation 178]
178. ! [X0,X1] : (~po(X0,X1) | po(X1,X0))[ennf transformation 104]
104. ! [X0,X1] : (~po(X0,X1) | po(X1,X0))[flattening 3]
3. ! [X0] : ! [X1] : (~po(X0,X1) | po(X1,X0))[input]
17469. po(sum(sK0,sK1),sK2) (5:5) [resolution 15404,14651]
14651. ~p(sK4(sK1,sK2),X1) | po(X1,sK2) (2:8) [resolution 14603,293]
293. ~p(X3,X1) | po(X0,X1) | ~p(X3,X0) (0:9) [cnf transformation 248]
248. ! [X0,X1] : ((~po(X0,X1) | (p(sK4(X1,X0),X0) & p(sK4(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[skolemisation 164]
164. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[ennf transformation 87]
87. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[flattening 86]
86. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[rectify 47]
47. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X2] : (p(X2,X0) & p(X2,X1)) | po(X0,X1)))[input]
14603. p(sK4(sK1,sK2),sK2) (1:5) [resolution 14587,291]
291. ~po(X0,X1) | p(sK4(X1,X0),X0) (0:8) [cnf transformation 248]
14587. po(sK2,sK1) (0:3) [backtracking split refutation 265,419,14586]
14586. $false {1} (6:0) [subsumption resolution 14582,266]
14582. po(sK2,sum(sK0,sK1)) {1} (6:5) [resolution 8220,318]
8220. po(sum(sK0,sK1),sK2) {1} (5:5) [resolution 3204,498]
498. ~p(sK4(sK0,sK2),X1) | po(X1,sK2) {1} (2:8) [resolution 432,293]
432. p(sK4(sK0,sK2),sK2) {1} (1:5) [resolution 419,291]
3204. p(sK4(sK0,sK2),sum(sK0,sK1)) {1} (4:7) [subsumption resolution 3192,859]
859. eqdim(sK4(sK0,sK2),sum(sK0,sK1)) {1} (3:7) [resolution 465,511]
511. eqdim(sK4(sK0,sK2),sK0) {1} (2:5) [resolution 433,298]
298. ~p(X0,X1) | eqdim(X0,X1) (0:6) [cnf transformation 168]
168. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 167]
167. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 89]
89. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 68]
68. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
433. p(sK4(sK0,sK2),sK0) {1} (1:5) [resolution 419,292]
292. ~po(X0,X1) | p(sK4(X1,X0),X1) (0:8) [cnf transformation 248]
465. ~eqdim(X0,sK0) | eqdim(X0,sum(sK0,sK1)) (2:8) [resolution 407,373]
373. ~eqdim(X1,X2) | eqdim(X0,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 236]
236. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~eqdim(X1,X2) | eqdim(X0,X2))[flattening 235]
235. ! [X0,X1,X2] : ((~eqdim(X0,X1) | ~eqdim(X1,X2)) | eqdim(X0,X2))[ennf transformation 141]
141. ! [X0,X1,X2] : (~(eqdim(X0,X1) & eqdim(X1,X2)) | eqdim(X0,X2))[flattening 54]
54. ! [X0] : ! [X1] : ! [X2] : (~(eqdim(X0,X1) & eqdim(X1,X2)) | eqdim(X0,X2))[input]
407. eqdim(sK0,sum(sK0,sK1)) (1:5) [subsumption resolution 406,262]
262. ~zex(sK0) (0:2) [cnf transformation 246]
406. zex(sK0) | eqdim(sK0,sum(sK0,sK1)) (1:7) [subsumption resolution 400,263]
263. ~zex(sK1) (0:2) [cnf transformation 246]
400. zex(sK1) | zex(sK0) | eqdim(sK0,sum(sK0,sK1)) (1:9) [resolution 261,334]
334. ~eqdim(X0,X1) | zex(X1) | zex(X0) | eqdim(X0,sum(X0,X1)) (0:12) [cnf transformation 200]
200. ! [X0,X1] : (~eqdim(X0,X1) | zex(X0) | zex(X1) | eqdim(X0,sum(X0,X1)))[flattening 199]
199. ! [X0,X1] : ((~eqdim(X0,X1) | zex(X0) | zex(X1)) | eqdim(X0,sum(X0,X1)))[ennf transformation 119]
119. ! [X0,X1] : (~(eqdim(X0,X1) & ~zex(X0) & ~zex(X1)) | eqdim(X0,sum(X0,X1)))[flattening 11]
11. ! [X0] : ! [X1] : (~(eqdim(X0,X1) & ~zex(X0) & ~zex(X1)) | eqdim(X0,sum(X0,X1)))[input]
261. eqdim(sK0,sK1) (0:3) [cnf transformation 246]
3192. ~eqdim(sK4(sK0,sK2),sum(sK0,sK1)) | p(sK4(sK0,sK2),sum(sK0,sK1)) {1} (4:14) [resolution 840,299]
299. ~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1) (0:9) [cnf transformation 168]
840. cont(sK4(sK0,sK2),sum(sK0,sK1)) {1} (3:7) [resolution 445,510]
510. cont(sK4(sK0,sK2),sK0) {1} (2:5) [resolution 433,297]
297. ~p(X0,X1) | cont(X0,X1) (0:6) [cnf transformation 168]
445. ~cont(X2,sK0) | cont(X2,sum(sK0,sK1)) (2:8) [forward demodulation 442,279]
279. sum(X0,X1) = sum(X1,X0) (0:7) [cnf transformation 80]
80. ! [X0,X1] : sum(X0,X1) = sum(X1,X0)[flattening 12]
12. ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)[input]
442. ~cont(X2,sK0) | cont(X2,sum(sK1,sK0)) (2:8) [resolution 403,370]
370. ~leq(X0,X1) | ~cont(X2,X1) | cont(X2,sum(X0,X1)) (0:11) [cnf transformation 230]
230. ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X1) | cont(X2,sum(X0,X1)))[flattening 229]
229. ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X1)) | cont(X2,sum(X0,X1)))[ennf transformation 138]
138. ! [X0,X1,X2] : (~(leq(X0,X1) & cont(X2,X1)) | cont(X2,sum(X0,X1)))[flattening 14]
14. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & cont(X2,X1)) | cont(X2,sum(X0,X1)))[input]
403. leq(sK1,sK0) (1:3) [resolution 261,295]
295. ~eqdim(X0,X1) | leq(X1,X0) (0:6) [cnf transformation 166]
166. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 165]
165. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[ennf transformation 88]
88. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[flattening 76]
76. ! [X0] : ! [X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[input]
419. po(sK2,sK0) {1} (0:3) [backtracking splitting component]
265. po(sK2,sK1) | po(sK2,sK0) (0:6) [cnf transformation 246]
15404. p(sK4(sK1,sK2),sum(sK0,sK1)) (4:7) [subsumption resolution 15392,14834]
14834. eqdim(sK4(sK1,sK2),sum(sK0,sK1)) (4:7) [resolution 14655,571]
571. ~eqdim(X0,sK1) | eqdim(X0,sum(sK0,sK1)) (3:8) [resolution 418,373]
418. eqdim(sK1,sum(sK0,sK1)) (2:5) [forward demodulation 417,279]
417. eqdim(sK1,sum(sK1,sK0)) (2:5) [subsumption resolution 416,263]
416. zex(sK1) | eqdim(sK1,sum(sK1,sK0)) (2:7) [subsumption resolution 410,262]
410. zex(sK0) | zex(sK1) | eqdim(sK1,sum(sK1,sK0)) (2:9) [resolution 401,334]
401. eqdim(sK1,sK0) (1:3) [resolution 261,321]
321. ~eqdim(X0,X1) | eqdim(X1,X0) (0:6) [cnf transformation 181]
181. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 107]
107. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[flattening 53]
53. ! [X0] : ! [X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[input]
14655. eqdim(sK4(sK1,sK2),sK1) (2:5) [resolution 14604,298]
14604. p(sK4(sK1,sK2),sK1) (1:5) [resolution 14587,292]
15392. ~eqdim(sK4(sK1,sK2),sum(sK0,sK1)) | p(sK4(sK1,sK2),sum(sK0,sK1)) (4:14) [resolution 14807,299]
14807. cont(sK4(sK1,sK2),sum(sK0,sK1)) (3:7) [resolution 14654,425]
425. ~cont(X2,sK1) | cont(X2,sum(sK0,sK1)) (2:8) [resolution 402,370]
402. leq(sK0,sK1) (1:3) [resolution 261,294]
294. ~eqdim(X0,X1) | leq(X0,X1) (0:6) [cnf transformation 166]
14654. cont(sK4(sK1,sK2),sK1) (2:5) [resolution 14604,297]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1814
Passive clauses: 9619
Generated clauses: 20837
Final active clauses: 919
Final passive clauses: 3296
Input formulas: 78
Initial clauses: 122

Duplicate literals: 34
Fw subsumption resolutions: 2187
Fw demodulations: 1070
Bw demodulations: 1549

Simple tautologies: 153
Forward subsumptions: 7711

Binary resolution: 14121
Forward superposition: 102
Backward superposition: 50

Backtracking splits: 106
Backtracking splits refuted: 56
Backtracking splits refuted at zero level: 16

Memory used: 3709KB
Time elapsed: 0.742 s
------------------------------
% Success in time 22.677 s
