Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal02.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 915
Passive clauses: 4743
Generated clauses: 12919
Final active clauses: 905
Final passive clauses: 3534
Input formulas: 46
Initial clauses: 80

Selected by SInE selection: 46
SInE iterations: 1

Duplicate literals: 85
Fw subsumption resolutions: 235

Simple tautologies: 49
Forward subsumptions: 7891

Binary resolution: 12312
Factoring: 45

Backtracking splits: 127
Backtracking splits refuted: 12
Backtracking splits refuted at zero level: 3

Memory used: 1791KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal02.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1271
Passive clauses: 11164
Generated clauses: 23987
Final active clauses: 1271
Final passive clauses: 9431
Input formulas: 46
Initial clauses: 79

Selected by SInE selection: 46
SInE iterations: 1

Duplicate literals: 223
Fw subsumption resolutions: 1729

Simple tautologies: 458
Equational tautologies: 13
Forward subsumptions: 10839

Binary resolution: 18800
Factoring: 192
Forward superposition: 2910
Backward superposition: 54

Unique components: 5

Memory used: 3070KB
Time elapsed: 1.402 s
------------------------------
% remaining time: 5981 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on theorems\conversions\codi_int_theorems_goal02.all
Refutation found. Thanks to Tanya!
24421. $false (1:0) [subsumption resolution 272,24409]
24409. ~cont(intersection(sK0,sK1),sK0) (6:5) [backtracking split refutation 24396,24406,24407]
24407. $false {283} (2:0) [subsumption resolution 519,24406]
519. cont(intersection(sK0,sK1),sK1) (2:5) [resolution 275,215]
215. ~p(X2,intersection(X0,X1)) | cont(X2,X1) (0:8) [cnf transformation 125]
125. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 124]
124. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[ennf transformation 74]
74. ! [X0,X1,X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 12]
12. ! [X0] : ! [X1] : ! [X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[input]
275. p(intersection(sK0,sK1),intersection(sK0,sK1)) (1:7) [resolution 155,171]
171. zex(X0) | p(X0,X0) (0:5) [cnf transformation 2]
2. ! [X0] : (zex(X0) | p(X0,X0))[input]
155. ~zex(intersection(sK0,sK1)) (0:4) [cnf transformation 146]
146. ~zex(intersection(sK0,sK1)) & ~zex(intersection(sK1,sK0)) & ~zex(sK0) & ~zex(sK1) & ~p(intersection(sK1,sK0),intersection(sK0,sK1))[skolemisation 87]
87. ? [X0,X1] : (~zex(intersection(X0,X1)) & ~zex(intersection(X1,X0)) & ~zex(X0) & ~zex(X1) & ~p(intersection(X1,X0),intersection(X0,X1)))[flattening 86]
86. ? [X0,X1] : ((~zex(intersection(X0,X1)) & ~zex(intersection(X1,X0)) & ~zex(X0) & ~zex(X1)) & ~p(intersection(X1,X0),intersection(X0,X1)))[ennf transformation 47]
47. ~! [X0,X1] : (~(~zex(intersection(X0,X1)) & ~zex(intersection(X1,X0)) & ~zex(X0) & ~zex(X1)) | p(intersection(X1,X0),intersection(X0,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : (~(~zex(intersection(X0,X1)) & ~zex(intersection(X1,X0)) & ~zex(X0) & ~zex(X1)) | p(intersection(X1,X0),intersection(X0,X1)))[negated conjecture]
24406. ~cont(intersection(sK0,sK1),sK1) {283} (6:5) [backtracking splitting component]
24396. ~cont(intersection(sK0,sK1),sK0) | ~cont(intersection(sK0,sK1),sK1) (6:10) [resolution 24274,226]
226. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 142]
142. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 141]
141. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 84]
84. ! [X0,X1,X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[flattening 11]
11. ! [X0] : ! [X1] : ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[input]
24274. ~leq(intersection(sK0,sK1),intersection(sK1,sK0)) (5:7) [backtracking split refutation 23876,23878,24273]
24273. $false {270} (1:0) [subsumption resolution 281,24271]
24271. ~cont(intersection(sK1,sK0),sK1) {270} (6:5) [backtracking split refutation 24238,24267,24268]
24268. $false {280} (2:0) [subsumption resolution 883,24267]
883. cont(intersection(sK1,sK0),sK0) (2:5) [resolution 284,215]
284. p(intersection(sK1,sK0),intersection(sK1,sK0)) (1:7) [resolution 156,171]
156. ~zex(intersection(sK1,sK0)) (0:4) [cnf transformation 146]
24267. ~cont(intersection(sK1,sK0),sK0) {280} (6:5) [backtracking splitting component]
24238. ~cont(intersection(sK1,sK0),sK1) | ~cont(intersection(sK1,sK0),sK0) {270} (6:10) [resolution 23878,226]
281. cont(intersection(sK1,sK0),sK1) (1:5) [resolution 156,190]
190. zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0) (0:9) [cnf transformation 58]
58. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 10]
10. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
23878. ~leq(intersection(sK1,sK0),intersection(sK0,sK1)) {270} (5:7) [backtracking splitting component]
23876. ~leq(intersection(sK0,sK1),intersection(sK1,sK0)) | ~leq(intersection(sK1,sK0),intersection(sK0,sK1)) (5:14) [resolution 23304,188]
188. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 102]
102. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 101]
101. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[ennf transformation 56]
56. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[flattening 45]
45. ! [X0] : ! [X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[input]
23304. ~eqdim(intersection(sK1,sK0),intersection(sK0,sK1)) (4:7) [backtracking split refutation 22764,23299,23300]
23300. $false {251} (0:0) [subsumption resolution 159,23299]
159. ~p(intersection(sK1,sK0),intersection(sK0,sK1)) (0:7) [cnf transformation 146]
23299. p(intersection(sK1,sK0),intersection(sK0,sK1)) {251} (4:7) [backtracking splitting component]
22764. ~eqdim(intersection(sK1,sK0),intersection(sK0,sK1)) | p(intersection(sK1,sK0),intersection(sK0,sK1)) (4:14) [resolution 1200,281]
1200. ~cont(intersection(sK1,sK0),X1) | ~eqdim(intersection(sK1,sK0),intersection(sK0,X1)) | p(intersection(sK1,sK0),intersection(sK0,X1)) (3:19) [resolution 883,213]
213. ~cont(X2,X0) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | p(X2,intersection(X0,X1)) (0:16) [cnf transformation 125]
272. cont(intersection(sK0,sK1),sK0) (1:5) [resolution 155,190]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 2322
Passive clauses: 6052
Generated clauses: 26539
Final active clauses: 1362
Final passive clauses: 1862
Input formulas: 46
Initial clauses: 77

Duplicate literals: 135
Trivial inequalities: 27
Bw subsumption resolutions: 181
Fw demodulations: 3263
Condensations: 20

Simple tautologies: 140
Equational tautologies: 58
Forward subsumptions: 16365

Binary resolution: 18534
Factoring: 7
Forward superposition: 342
Backward superposition: 1169

Backtracking splits: 284
Backtracking splits refuted: 114
Backtracking splits refuted at zero level: 5

Memory used: 4221KB
Time elapsed: 1.142 s
------------------------------
% Success in time 3.065 s
