Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal21.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 946
Passive clauses: 5099
Generated clauses: 13578
Final active clauses: 930
Final passive clauses: 3781
Input formulas: 50
Initial clauses: 86

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 93
Fw subsumption resolutions: 579

Simple tautologies: 67
Forward subsumptions: 7908

Binary resolution: 12623
Factoring: 28

Backtracking splits: 142
Backtracking splits refuted: 12

Memory used: 2174KB
Time elapsed: 0.405 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal21.all
Refutation found. Thanks to Tanya!
1584. $false (7:0) [subsumption resolution 1576,431]
431. ~zex(intersection(sK0,sK0)) (2:4) [resolution 425,304]
304. c(sK1,intersection(sK0,sK0)) (1:5) [resolution 187,173]
173. ~zex(intersection(sK1,intersection(sK0,sK0))) (0:6) [cnf transformation 162]
162. zex(intersection(sK0,sK1)) & ~zex(intersection(sK1,intersection(sK0,sK0)))[skolemisation 95]
95. ? [X0,X1] : (zex(intersection(X0,X1)) & ~zex(intersection(X1,intersection(X0,X0))))[ennf transformation 51]
51. ~! [X0,X1] : (~zex(intersection(X0,X1)) | zex(intersection(X1,intersection(X0,X0))))[flattening 1]
1. ~! [X0] : ! [X1] : (~zex(intersection(X0,X1)) | zex(intersection(X1,intersection(X0,X0))))[negated conjecture]
187. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 102]
102. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[ennf transformation 54]
54. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[flattening 11]
11. ! [X0] : ! [X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[input]
425. ~c(X2,X3) | ~zex(X3) (1:5) [resolution 199,213]
213. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 116]
116. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 67]
67. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[flattening 31]
31. ! [X0] : ! [X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[input]
199. cont(sK5(X1,X0),X1) | ~c(X0,X1) (0:8) [cnf transformation 166]
166. ! [X0,X1] : ((~c(X0,X1) | (cont(sK5(X1,X0),X0) & cont(sK5(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 108]
108. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 60]
60. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 59]
59. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 34]
34. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
1576. zex(intersection(sK0,sK0)) (7:4) [resolution 1562,211]
211. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 65]
65. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 12]
12. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
1562. ~cont(intersection(sK0,sK0),sK0) (6:5) [resolution 1445,173]
1445. zex(intersection(sK1,X0)) | ~cont(X0,sK0) (5:7) [resolution 453,708]
708. ~cont(intersection(sK1,X16),X15) | ~cont(X15,sK0) (4:8) [resolution 247,539]
539. ~cont(intersection(sK1,X5),sK0) (3:5) [subsumption resolution 535,214]
214. ~cont(X0,X1) | ~zex(X0) (0:5) [cnf transformation 116]
535. zex(intersection(sK1,X5)) | ~cont(intersection(sK1,X5),sK0) (3:9) [resolution 211,459]
459. ~cont(X0,sK1) | ~cont(X0,sK0) (2:6) [resolution 200,305]
305. ~c(sK0,sK1) (1:3) [resolution 188,172]
172. zex(intersection(sK0,sK1)) (0:4) [cnf transformation 162]
188. ~zex(intersection(X0,X1)) | ~c(X0,X1) (0:7) [cnf transformation 102]
200. c(X0,X1) | ~cont(X3,X1) | ~cont(X3,X0) (0:9) [cnf transformation 166]
247. cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1) (0:9) [cnf transformation 154]
154. ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))[flattening 153]
153. ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))[ennf transformation 91]
91. ! [X0,X1,X2] : (~(cont(X0,X1) & cont(X1,X2)) | cont(X0,X2))[flattening 30]
30. ! [X0] : ! [X1] : ! [X2] : (~(cont(X0,X1) & cont(X1,X2)) | cont(X0,X2))[input]
453. cont(intersection(X0,X1),X1) | zex(intersection(X0,X1)) (1:9) [resolution 236,185]
185. p(X0,X0) | zex(X0) (0:5) [cnf transformation 2]
2. ! [X0] : (zex(X0) | p(X0,X0))[input]
236. ~p(X2,intersection(X0,X1)) | cont(X2,X1) (0:8) [cnf transformation 137]
137. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 136]
136. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[ennf transformation 81]
81. ! [X0,X1,X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 14]
14. ! [X0] : ! [X1] : ! [X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 229
Passive clauses: 636
Generated clauses: 1400
Final active clauses: 229
Final passive clauses: 366
Input formulas: 50
Initial clauses: 85

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 30
Fw subsumption resolutions: 102

Simple tautologies: 49
Equational tautologies: 1
Forward subsumptions: 615

Binary resolution: 700
Factoring: 19
Forward superposition: 418
Backward superposition: 46

Unique components: 2

Memory used: 383KB
Time elapsed: 0.056 s
------------------------------
% Success in time 0.535 s
