Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal08.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 971
Passive clauses: 4859
Generated clauses: 13474
Final active clauses: 951
Final passive clauses: 3544
Input formulas: 48
Initial clauses: 96

Selected by SInE selection: 48
SInE iterations: 1

Duplicate literals: 158
Fw subsumption resolutions: 353

Simple tautologies: 72
Forward subsumptions: 8226

Binary resolution: 12708
Factoring: 39

Backtracking splits: 89
Backtracking splits refuted: 25
Backtracking splits refuted at zero level: 8

Memory used: 2046KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal08.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1520
Passive clauses: 8293
Generated clauses: 25755
Final active clauses: 1520
Final passive clauses: 6205
Input formulas: 48
Initial clauses: 95

Selected by SInE selection: 48
SInE iterations: 1

Duplicate literals: 1310
Fw subsumption resolutions: 1757

Simple tautologies: 814
Equational tautologies: 6
Forward subsumptions: 14189

Binary resolution: 14495
Factoring: 159
Forward superposition: 6066
Backward superposition: 1930

Unique components: 2

Memory used: 2814KB
Time elapsed: 1.401 s
------------------------------
% remaining time: 5981 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on theorems\conversions\codi_int_theorems_goal08.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 5131
Passive clauses: 12768
Generated clauses: 50030
Final active clauses: 4865
Final passive clauses: 7008
Input formulas: 48
Initial clauses: 93

Duplicate literals: 333
Bw subsumption resolutions: 192
Fw demodulations: 1765
Condensations: 32

Simple tautologies: 162
Equational tautologies: 374
Forward subsumptions: 33861

Binary resolution: 45520
Factoring: 17
Forward superposition: 176
Backward superposition: 990

Backtracking splits: 439
Backtracking splits refuted: 139
Backtracking splits refuted at zero level: 1

Memory used: 10618KB
Time elapsed: 3.310 s
------------------------------
% remaining time: 5947 next slice time: 95
dis+2_4_bs=off:ep=on:n=on:nicw=on:nwc=1.5:ptb=off:sac=on:sio=off:spl=backtracking:ssec=off_95 on theorems\conversions\codi_int_theorems_goal08.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 3204
Passive clauses: 59343
Generated clauses: 129941
Final active clauses: 1996
Final passive clauses: 42029
Input formulas: 48
Initial clauses: 95

Duplicate literals: 5569
Fw subsumption resolutions: 7488
Fw demodulations: 1745

Simple tautologies: 963
Equational tautologies: 78
Forward subsumptions: 57271

Binary resolution: 114753
Factoring: 49
Forward superposition: 722
Backward superposition: 16

Backtracking splits: 37
Backtracking splits refuted: 16
Backtracking splits refuted at zero level: 6

Memory used: 22899KB
Time elapsed: 9.614 s
------------------------------
% remaining time: 5851 next slice time: 4
lrs-1003_12_bms=on:bs=unit_only:bsr=on:cond=fast:drc=off:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=10:sfv=off:sgo=on:sio=off:sp=occurrence:spo=on:ssec=off:stl=81:updr=off_4 on theorems\conversions\codi_int_theorems_goal08.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 711
Passive clauses: 1526
Generated clauses: 3495
Final active clauses: 695
Final passive clauses: 773
Input formulas: 48
Initial clauses: 95

Duplicate literals: 24
Fw subsumption resolutions: 67
Bw subsumption resolutions: 3
Condensations: 26

Simple tautologies: 29
Equational tautologies: 26
Forward subsumptions: 1791
Backward subsumptions: 143

Binary resolution: 3114
Factoring: 116
Forward superposition: 12
Backward superposition: 37

Split clauses: 1
Split components: 2
Unique components: 787
Introduced splitting names: 1

Memory used: 1279KB
Time elapsed: 0.503 s
------------------------------
% remaining time: 5845 next slice time: 22
lrs+1010_50_bs=off:drc=off:flr=on:lcm=predicate:n=on:nwc=1.2:sfv=off:sgo=on:stl=81_22 on theorems\conversions\codi_int_theorems_goal08.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 842
Passive clauses: 2350
Generated clauses: 31816
Final active clauses: 842
Final passive clauses: 1508
Input formulas: 48
Initial clauses: 93

Duplicate literals: 3095
Fw subsumption resolutions: 4900
Fw literal rewrites: 1790

Simple tautologies: 759
Equational tautologies: 1
Forward subsumptions: 19098

Binary resolution: 22126
Factoring: 19

Unique components: 2

Memory used: 2302KB
Time elapsed: 2.302 s
------------------------------
% remaining time: 5822 next slice time: 41
dis-1010_5_bd=off:bs=off:cond=fast:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=1.3:ptb=off:sac=on:sgo=on:sio=off:sos=on:spl=backtracking:ssec=off:updr=off_41 on theorems\conversions\codi_int_theorems_goal08.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 3063
Passive clauses: 11018
Generated clauses: 43901
Final active clauses: 2996
Final passive clauses: 7210
Input formulas: 48
Initial clauses: 95

Duplicate literals: 427
Fw subsumption resolutions: 1219
Condensations: 126

Simple tautologies: 236
Forward subsumptions: 31633

Binary resolution: 41748
Factoring: 569

Backtracking splits: 17
Backtracking splits refuted: 6
Backtracking splits refuted at zero level: 6

Memory used: 4989KB
Time elapsed: 4.208 s
------------------------------
% remaining time: 5779 next slice time: 237
lrs+1_2:1_bs=off:ep=on:gsp=input_only:lcm=predicate:n=on:nicw=on:nwc=1.1:ptb=off:sac=on:sgo=on:sio=off:sos=on:sp=occurrence:spl=backtracking:spo=on:ssec=off:stl=81:updr=off_237 on theorems\conversions\codi_int_theorems_goal08.all
Refutation found. Thanks to Tanya!
1838. $false (6:0) [subsumption resolution 1831,1755]
1755. ~leq(sK0,intersection(sK0,sK1)) (8:5) [resolution 997,227]
227. ~lt(X0,X1) | ~leq(X1,X0) (0:6) [cnf transformation 117]
117. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 116]
116. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[ennf transformation 67]
67. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[flattening 44]
44. ! [X0] : ! [X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[input]
997. lt(intersection(sK0,sK1),sK0) (7:5) [subsumption resolution 990,174]
174. ~eqdim(intersection(sK0,sK1),sK0) (0:5) [cnf transformation 157]
157. po(sK0,sK1) & ~eqdim(intersection(sK0,sK1),sK0)[skolemisation 94]
94. ? [X0,X1] : (po(X0,X1) & ~eqdim(intersection(X0,X1),X0))[ennf transformation 49]
49. ~! [X0,X1] : (~po(X0,X1) | eqdim(intersection(X0,X1),X0))[flattening 1]
1. ~! [X0] : ! [X1] : (~po(X0,X1) | eqdim(intersection(X0,X1),X0))[negated conjecture]
990. lt(intersection(sK0,sK1),sK0) | eqdim(intersection(sK0,sK1),sK0) (7:10) [resolution 486,239]
239. ~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1) (0:9) [cnf transformation 125]
125. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[ennf transformation 77]
77. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 34]
34. ! [X0] : ! [X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[input]
486. cont(intersection(sK0,sK1),sK0) (6:5) [resolution 463,230]
230. zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0) (0:9) [cnf transformation 69]
69. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 7]
7. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
463. ~zex(intersection(sK0,sK1)) (5:4) [resolution 459,189]
189. ~c(X0,X1) | ~zex(intersection(X0,X1)) (0:7) [cnf transformation 101]
101. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[ennf transformation 52]
52. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[flattening 6]
6. ! [X0] : ! [X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[input]
459. c(sK0,sK1) (4:3) [resolution 334,292]
292. cont(sK10(sK1,sK0),sK0) (2:5) [resolution 278,220]
220. ~p(X0,X1) | cont(X0,X1) (0:6) [cnf transformation 113]
113. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 112]
112. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 65]
65. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 33]
33. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
278. p(sK10(sK1,sK0),sK0) (1:5) [resolution 173,217]
217. ~po(X0,X1) | p(sK10(X1,X0),X0) (0:8) [cnf transformation 168]
168. ! [X0,X1] : ((~po(X0,X1) | (p(sK10(X1,X0),X0) & p(sK10(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[skolemisation 111]
111. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[ennf transformation 64]
64. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[flattening 63]
63. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[rectify 10]
10. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X2] : (p(X2,X0) & p(X2,X1)) | po(X0,X1)))[input]
173. po(sK0,sK1) (0:3) [cnf transformation 157]
334. ~cont(sK10(sK1,sK0),X2) | c(X2,sK1) (3:8) [resolution 295,216]
216. ~cont(X3,X1) | c(X0,X1) | ~cont(X3,X0) (0:9) [cnf transformation 167]
167. ! [X0,X1] : ((~c(X0,X1) | (cont(sK9(X1,X0),X0) & cont(sK9(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 110]
110. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 62]
62. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 61]
61. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 32]
32. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
295. cont(sK10(sK1,sK0),sK1) (2:5) [resolution 279,220]
279. p(sK10(sK1,sK0),sK1) (1:5) [resolution 173,218]
218. ~po(X0,X1) | p(sK10(X1,X0),X1) (0:8) [cnf transformation 168]
1831. leq(sK0,intersection(sK0,sK1)) (6:5) [resolution 603,331]
331. leq(sK0,sK10(sK1,sK0)) (3:5) [resolution 293,224]
224. ~eqdim(X0,X1) | leq(X1,X0) (0:6) [cnf transformation 115]
115. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 114]
114. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[ennf transformation 66]
66. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[flattening 47]
47. ! [X0] : ! [X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[input]
293. eqdim(sK10(sK1,sK0),sK0) (2:5) [resolution 278,221]
221. ~p(X0,X1) | eqdim(X0,X1) (0:6) [cnf transformation 113]
603. ~leq(X0,sK10(sK1,sK0)) | leq(X0,intersection(sK0,sK1)) (5:10) [resolution 475,258]
258. ~leq(X1,X2) | leq(X0,X2) | ~leq(X0,X1) (0:9) [cnf transformation 147]
147. ! [X0,X1,X2] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2))[flattening 146]
146. ! [X0,X1,X2] : ((~leq(X0,X1) | ~leq(X1,X2)) | leq(X0,X2))[ennf transformation 90]
90. ! [X0,X1,X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[flattening 38]
38. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[input]
475. leq(sK10(sK1,sK0),intersection(sK0,sK1)) (4:7) [resolution 341,292]
341. ~cont(sK10(sK1,sK0),X6) | leq(sK10(sK1,sK0),intersection(X6,sK1)) (3:12) [resolution 295,261]
261. ~cont(X2,X1) | leq(X2,intersection(X0,X1)) | ~cont(X2,X0) (0:11) [cnf transformation 153]
153. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 152]
152. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 93]
93. ! [X0,X1,X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[flattening 8]
8. ! [X0] : ! [X1] : ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 342
Passive clauses: 835
Generated clauses: 1604
Final active clauses: 334
Final passive clauses: 549
Input formulas: 48
Initial clauses: 95

Duplicate literals: 8
Fw subsumption resolutions: 61

Simple tautologies: 8
Forward subsumptions: 661

Binary resolution: 1466

Backtracking splits: 16
Backtracking splits refuted: 6
Backtracking splits refuted at zero level: 6

Memory used: 639KB
Time elapsed: 0.121 s
------------------------------
% Success in time 22.236 s
