% Starting sequential portfolio solving with schedule "casc"
% spawned child 9968 with time: 4 (total remaining time 2999)
% WARNING: value z3 for option sas not known
% dis+1010_4_afp=10000:afq=1.2:anc=none:irw=on:lma=on:nm=64:nwc=10:sas=z3:sac=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6652
% Time elapsed: 0.500 s
% spawned child 9969 with time: 10 (total remaining time 2994)
% ott+11_8_amm=off:anc=none:bsr=on:cond=on:irw=on:nm=2:nwc=1.1:ss=axioms:st=2.0:sac=on_8 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit

% Memory used [KB]: 7036
% Time elapsed: 1.119 s
% ------------------------------
% ------------------------------
% spawned child 9970 with time: 4 (total remaining time 2984)
% dis+1011_10_aac=none:add=large:afp=10000:afq=1.1:amm=off:anc=none:fsr=off:gs=on:gsem=off:irw=on:nm=6:newcnf=on:nwc=2.5:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 11129
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 9971 with time: 6 (total remaining time 2979)
% dis+11_28_av=off:fsr=off:irw=on:lcm=predicate:nm=2:newcnf=on:nwc=5:sp=occurrence:urr=on:updr=off_4 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 2942
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% spawned child 9972 with time: 6 (total remaining time 2972)
% lrs+1011_5_add=large:afp=1000:afq=1.2:anc=none:fsr=off:irw=on:lma=on:nm=64:newcnf=on:nwc=1:stl=30:sd=3:ss=axioms:st=2.0:sos=on:sac=on:sp=reverse_arity:urr=on:updr=off_4 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7803
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% spawned child 9973 with time: 4 (total remaining time 2966)
% lrs+1011_3:1_add=off:afr=on:afp=10000:afq=1.1:amm=off:bce=on:cond=on:ep=R:fsr=off:nm=16:nwc=1:stl=30:sos=all:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8571
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 9974 with time: 23 (total remaining time 2960)
% lrs+1011_3:2_aac=none:afp=10000:afq=1.2:amm=off:anc=all:cond=on:fde=none:gs=on:gsem=on:irw=on:lma=on:nm=32:newcnf=on:nwc=3:nicw=on:stl=30:sac=on:sp=reverse_arity:urr=ec_only:updr=off:uhcvi=on_21 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 16886
% Time elapsed: 2.400 s
% ------------------------------
% ------------------------------
% spawned child 9975 with time: 4 (total remaining time 2937)
% lrs+10_6_aac=none:acc=model:add=large:afr=on:afp=10000:afq=1.2:amm=off:anc=none:bd=off:ccuc=small_ones:irw=on:lcm=reverse:nm=0:nwc=1:nicw=on:stl=30:sos=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7419
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 9976 with time: 7 (total remaining time 2931)
% dis+11_5_av=off:bd=off:bs=unit_only:bsr=on:cond=on:lcm=reverse:nm=0:nwc=1.2_5 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 2942
% Time elapsed: 0.800 s
% ------------------------------
% ------------------------------
% spawned child 9977 with time: 4 (total remaining time 2922)
% dis+11_32_av=off:ep=RST:fsr=off:lwlo=on:nm=6:nwc=1.1:sd=5:ss=axioms:st=5.0:sp=reverse_arity:uhcvi=on_2 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 2174
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 9978 with time: 7 (total remaining time 2916)
% dis+1003_64_add=off:afr=on:afp=100000:afq=1.1:anc=none:cond=fast:fde=none:irw=on:nm=6:newcnf=on:nwc=1.3:uhcvi=on_5 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7036
% Time elapsed: 0.800 s
% ------------------------------
% ------------------------------
% spawned child 9979 with time: 5 (total remaining time 2908)
% dis+11_4_afr=on:afp=1000:afq=1.4:amm=off:anc=none:lcm=reverse:lma=on:lwlo=on:nm=6:newcnf=on:nwc=1:sd=2:ss=axioms:st=2.0:sp=reverse_arity_3 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6780
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% spawned child 9980 with time: 87 (total remaining time 2901)
% ott+1002_128_av=off:bd=off:bs=on:bsr=on:cond=on:fsr=off:nm=6:newcnf=on:nwc=1:sp=reverse_arity:updr=off_82 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6652
% Time elapsed: 8.800 s
% ------------------------------
% ------------------------------
% spawned child 9982 with time: 8 (total remaining time 2821)
% lrs+1002_8_add=large:afp=40000:afq=1.0:amm=off:anc=none:cond=on:gs=on:irw=on:nm=16:newcnf=on:nwc=1:stl=30:sos=on:sp=reverse_arity:updr=off_6 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 9850
% Time elapsed: 1.020 s
% ------------------------------
% ------------------------------
% spawned child 9983 with time: 54 (total remaining time 2810)
% lrs+1011_3:1_av=off:cond=on:irw=on:lma=on:nm=16:nwc=1:stl=30:sos=all:updr=off_51 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8059
% Time elapsed: 5.500 s
% ------------------------------
% ------------------------------
% spawned child 9984 with time: 7 (total remaining time 2760)
% lrs+1_3:2_aac=none:add=large:anc=all_dependent:bce=on:cond=fast:ep=RST:fsr=off:lma=on:nm=2:nwc=1:stl=30:sos=on:sp=occurrence:urr=on:updr=off:uhcvi=on_5 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6652
% Time elapsed: 0.800 s
% ------------------------------
% ------------------------------
% spawned child 9985 with time: 11 (total remaining time 2752)
% ott+1_8_av=off:bd=off:bs=on:cond=on:gsp=input_only:gs=on:gsem=off:irw=on:lcm=predicate:lwlo=on:nwc=1:sos=on_10 on codi_down_sum_theorems_goal02.all
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 5245
% Time elapsed: 0.047 s
% ------------------------------
% ------------------------------
% spawned child 9986 with time: 7 (total remaining time 2751)
% dis+11_5_afp=40000:afq=1.4:anc=none:br=off:cond=on:fsr=off:irw=on:lma=on:nm=4:nwc=1:sd=3:ss=axioms:st=2.0:urr=on:updr=off_5 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8699
% Time elapsed: 0.800 s
% ------------------------------
% ------------------------------
% spawned child 9987 with time: 4 (total remaining time 2743)
% dis+1010_10_afr=on:afp=100000:afq=2.0:amm=sco:anc=none:ep=RS:fde=unused:gs=on:nwc=1:sos=on:sac=on:sp=occurrence_2 on codi_down_sum_theorems_goal02.all
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation not found, incomplete strategy

% Memory used [KB]: 9722
% Time elapsed: 0.114 s
% ------------------------------
% ------------------------------
% spawned child 9988 with time: 9 (total remaining time 2741)
% lrs+1011_8:1_afr=on:afp=1000:afq=2.0:br=off:gsp=input_only:gs=on:nm=16:nwc=1:stl=30:sos=all:sp=occurrence:urr=on_7 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 14711
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% spawned child 9989 with time: 9 (total remaining time 2731)
% lrs+1003_2_acc=on:afp=4000:afq=2.0:amm=sco:bd=off:ccuc=small_ones:fsr=off:fde=unused:gsp=input_only:nm=64:newcnf=on:nwc=2.5:nicw=on:stl=30:urr=ec_only_7 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8699
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% spawned child 9990 with time: 4 (total remaining time 2721)
% ott+4_40_av=off:bce=on:cond=fast:fde=none:nm=0:nwc=1:sos=all:updr=off_2 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 2558
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 9991 with time: 5 (total remaining time 2715)
% dis-3_5_av=off:cond=on:gs=on:gsem=off:lcm=predicate:nm=0:nwc=1:sos=on:sp=reverse_arity_3 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% spawned child 9992 with time: 56 (total remaining time 2709)
% lrs+1011_14_av=off:fsr=off:irw=on:nwc=1:stl=30:sos=on:sp=occurrence:urr=ec_only:updr=off_53 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 11769
% Time elapsed: 5.700 s
% ------------------------------
% ------------------------------
% spawned child 9993 with time: 5 (total remaining time 2655)
% dis+11_10_av=off:lma=on:nm=64:nwc=5:sp=reverse_arity_3 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 1407
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% spawned child 9994 with time: 13 (total remaining time 2648)
% lrs+1011_4:1_av=off:fsr=off:irw=on:nwc=1:stl=30:sd=4:ss=axioms:st=1.5:sp=reverse_arity_12 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 4093
% Time elapsed: 1.400 s
% ------------------------------
% ------------------------------
% spawned child 9995 with time: 245 (total remaining time 2636)
% lrs+10_2:3_afp=1000:afq=1.1:amm=sco:anc=none:er=known:lcm=reverse:lwlo=on:nm=64:newcnf=on:nwc=2:stl=30:sd=5:ss=axioms:sos=theory:sac=on:sp=occurrence_233 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 51811
% Time elapsed: 24.600 s
% ------------------------------
% ------------------------------
% spawned child 9998 with time: 7 (total remaining time 2431)
% lrs+1011_3:2_av=off:er=known:lma=on:nm=2:newcnf=on:nwc=2:stl=30:sd=2:ss=axioms:st=3.0:urr=on:updr=off_5 on codi_down_sum_theorems_goal02.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit

% Memory used [KB]: 767
% Time elapsed: 4.254 s
% ------------------------------
% ------------------------------
% spawned child 9999 with time: 38 (total remaining time 2388)
% lrs+1011_2:3_av=off:gs=on:gsem=off:nwc=1.5:stl=30:sos=theory:sp=occurrence:urr=ec_only:updr=off_36 on codi_down_sum_theorems_goal02.all
% Refutation found. Thanks to Tanya!
% SZS status Theorem for codi_down_sum_theorems_goal02.all
% SZS output start Proof for codi_down_sum_theorems_goal02.all
fof(f1,conjecture,(
  ! [X0] : (p(X0,sum(X0,X0)) | zex(X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',conjecture20)).
fof(f2,negated_conjecture,(
  ~! [X0] : (p(X0,sum(X0,X0)) | zex(X0))),
  inference(negated_conjecture,[],[f1])).
fof(f5,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (cont(X2,sum(X0,X1)) | ~(cont(X2,X1) & leq(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom30)).
fof(f6,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom40)).
fof(f8,axiom,(
  ! [X0] : ! [X1] : (difference(X0,X1) = X0 | ~lt(X1,X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom15160)).
fof(f12,axiom,(
  ! [X0] : ! [X1] : ((zex(difference(X0,X1)) | ~(cont(X0,X1) | zex(X0))) & (cont(X0,X1) | zex(X0) | ~zex(difference(X0,X1))))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom15200)).
fof(f21,axiom,(
  ! [X0] : ! [X1] : (p(difference(X0,X1),X0) | zex(difference(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom64980)).
fof(f32,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (p(X2,intersection(X0,difference(X0,X1))) | ~p(X2,difference(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom74390)).
fof(f66,axiom,(
  ! [X0] : ! [X1] : (cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom87200)).
fof(f67,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (leq(X2,intersection(X0,X1)) | ~(cont(X2,X1) & cont(X2,X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom87210)).
fof(f72,axiom,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom89380)).
fof(f84,axiom,(
  ! [X0] : ! [X1] : (lt(X0,X1) | ~(~zex(X1) & zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom89500)).
fof(f86,axiom,(
  ! [X0] : ! [X1] : (X0 = X1 | ~(cont(X1,X0) & cont(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom89520)).
fof(f88,axiom,(
  ! [X0] : ! [X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom89540)).
fof(f92,axiom,(
  ! [X0] : ! [X1] : ((p(X0,X1) | ~(eqdim(X0,X1) & cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom88840)).
fof(f95,axiom,(
  ! [X0] : leq(X0,X0)),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom89100)).
fof(f103,axiom,(
  ! [X0] : ! [X1] : ((lt(X0,X1) | ~(~leq(X1,X0) & leq(X0,X1))) & ((~leq(X1,X0) & leq(X0,X1)) | ~lt(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom89280)).
fof(f107,axiom,(
  ? [X0] : zex(X0)),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal02.all.tptp',axiom87750)).
fof(f150,plain,(
  ! [X0,X1] : ((lt(X0,X1) | ~(~leq(X1,X0) & leq(X0,X1))) & ((~leq(X1,X0) & leq(X0,X1)) | ~lt(X0,X1)))),
  inference(flattening,[],[f103])).
fof(f151,plain,(
  ! [X0,X1] : ((p(X0,X1) | ~(eqdim(X0,X1) & cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  inference(flattening,[],[f92])).
fof(f153,plain,(
  ! [X0,X1] : ((zex(difference(X0,X1)) | ~(cont(X0,X1) | zex(X0))) & (cont(X0,X1) | zex(X0) | ~zex(difference(X0,X1))))),
  inference(flattening,[],[f12])).
fof(f157,plain,(
  ! [X0,X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  inference(flattening,[],[f88])).
fof(f168,plain,(
  ! [X0,X1] : (difference(X0,X1) = X0 | ~lt(X1,X0))),
  inference(flattening,[],[f8])).
fof(f170,plain,(
  ! [X0,X1] : (eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1))),
  inference(flattening,[],[f72])).
fof(f177,plain,(
  ! [X0,X1] : (cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)))),
  inference(flattening,[],[f66])).
fof(f178,plain,(
  ! [X0,X1] : (p(difference(X0,X1),X0) | zex(difference(X0,X1)))),
  inference(flattening,[],[f21])).
fof(f189,plain,(
  ! [X0,X1] : (lt(X0,X1) | ~(~zex(X1) & zex(X0)))),
  inference(flattening,[],[f84])).
fof(f196,plain,(
  ! [X0,X1] : (X0 = X1 | ~(cont(X1,X0) & cont(X0,X1)))),
  inference(flattening,[],[f86])).
fof(f211,plain,(
  ! [X0,X1,X2] : (p(X2,intersection(X0,difference(X0,X1))) | ~p(X2,difference(X0,X1)))),
  inference(flattening,[],[f32])).
fof(f219,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | ~(cont(X2,X1) & leq(X0,X1)))),
  inference(flattening,[],[f5])).
fof(f227,plain,(
  ! [X0,X1,X2] : (leq(X2,intersection(X0,X1)) | ~(cont(X2,X1) & cont(X2,X0)))),
  inference(flattening,[],[f67])).
fof(f232,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  inference(flattening,[],[f6])).
fof(f235,plain,(
  ? [X0] : (~p(X0,sum(X0,X0)) & ~zex(X0))),
  inference(ennf_transformation,[],[f2])).
fof(f261,plain,(
  ! [X0,X1] : ((lt(X0,X1) | (leq(X1,X0) | ~leq(X0,X1))) & ((~leq(X1,X0) & leq(X0,X1)) | ~lt(X0,X1)))),
  inference(ennf_transformation,[],[f150])).
fof(f262,plain,(
  ! [X0,X1] : ((lt(X0,X1) | leq(X1,X0) | ~leq(X0,X1)) & ((~leq(X1,X0) & leq(X0,X1)) | ~lt(X0,X1)))),
  inference(flattening,[],[f261])).
fof(f263,plain,(
  ! [X0,X1] : ((p(X0,X1) | (~eqdim(X0,X1) | ~cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  inference(ennf_transformation,[],[f151])).
fof(f264,plain,(
  ! [X0,X1] : ((p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1)) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  inference(flattening,[],[f263])).
fof(f265,plain,(
  ! [X0,X1] : ((zex(difference(X0,X1)) | (~cont(X0,X1) & ~zex(X0))) & (cont(X0,X1) | zex(X0) | ~zex(difference(X0,X1))))),
  inference(ennf_transformation,[],[f153])).
fof(f281,plain,(
  ! [X0,X1] : (lt(X0,X1) | (zex(X1) | ~zex(X0)))),
  inference(ennf_transformation,[],[f189])).
fof(f282,plain,(
  ! [X0,X1] : (lt(X0,X1) | zex(X1) | ~zex(X0))),
  inference(flattening,[],[f281])).
fof(f295,plain,(
  ! [X0,X1] : (X0 = X1 | (~cont(X1,X0) | ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f196])).
fof(f296,plain,(
  ! [X0,X1] : (X0 = X1 | ~cont(X1,X0) | ~cont(X0,X1))),
  inference(flattening,[],[f295])).
fof(f317,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | (~cont(X2,X1) | ~leq(X0,X1)))),
  inference(ennf_transformation,[],[f219])).
fof(f318,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | ~cont(X2,X1) | ~leq(X0,X1))),
  inference(flattening,[],[f317])).
fof(f333,plain,(
  ! [X0,X1,X2] : (leq(X2,intersection(X0,X1)) | (~cont(X2,X1) | ~cont(X2,X0)))),
  inference(ennf_transformation,[],[f227])).
fof(f334,plain,(
  ! [X0,X1,X2] : (leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0))),
  inference(flattening,[],[f333])).
fof(f343,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | (cont(X2,X0) | ~cont(X2,sum(X0,X1))))),
  inference(ennf_transformation,[],[f232])).
fof(f344,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f343])).
fof(f351,plain,(
  ? [X0] : (~p(X0,sum(X0,X0)) & ~zex(X0)) => (~p(sK1,sum(sK1,sK1)) & ~zex(sK1))),
  introduced(choice_axiom,[])).
fof(f352,plain,(
  ~p(sK1,sum(sK1,sK1)) & ~zex(sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f235,f351])).
fof(f402,plain,(
  ? [X0] : zex(X0) => zex(sK23)),
  introduced(choice_axiom,[])).
fof(f403,plain,(
  zex(sK23)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23])],[f107,f402])).
fof(f406,plain,(
  ~zex(sK1)),
  inference(cnf_transformation,[],[f352])).
fof(f407,plain,(
  ~p(sK1,sum(sK1,sK1))),
  inference(cnf_transformation,[],[f352])).
fof(f410,plain,(
  ( ! [X0] : (leq(X0,X0)) )),
  inference(cnf_transformation,[],[f95])).
fof(f475,plain,(
  ( ! [X0,X1] : (~lt(X0,X1) | ~leq(X1,X0)) )),
  inference(cnf_transformation,[],[f262])).
fof(f477,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f479,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f482,plain,(
  ( ! [X0,X1] : (~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f490,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f157])).
fof(f503,plain,(
  ( ! [X0,X1] : (~lt(X1,X0) | difference(X0,X1) = X0) )),
  inference(cnf_transformation,[],[f168])).
fof(f505,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f170])).
fof(f512,plain,(
  ( ! [X0,X1] : (cont(intersection(X0,X1),X0) | zex(intersection(X0,X1))) )),
  inference(cnf_transformation,[],[f177])).
fof(f513,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0)) )),
  inference(cnf_transformation,[],[f178])).
fof(f530,plain,(
  ( ! [X0,X1] : (zex(X1) | lt(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f282])).
fof(f537,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | X0 = X1 | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f296])).
fof(f559,plain,(
  ( ! [X2,X0,X1] : (~p(X2,difference(X0,X1)) | p(X2,intersection(X0,difference(X0,X1)))) )),
  inference(cnf_transformation,[],[f211])).
fof(f567,plain,(
  ( ! [X2,X0,X1] : (~cont(X2,X1) | cont(X2,sum(X0,X1)) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f318])).
fof(f575,plain,(
  ( ! [X2,X0,X1] : (leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0)) )),
  inference(cnf_transformation,[],[f334])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (~cont(X2,sum(X0,X1)) | cont(X2,X0) | cont(difference(X2,X0),X1)) )),
  inference(cnf_transformation,[],[f344])).
fof(f584,plain,(
  zex(sK23)),
  inference(cnf_transformation,[],[f403])).
fof(f606,plain,(
  ( ! [X2] : (lt(X2,sK1) | ~zex(X2)) )),
  inference(resolution,[],[f530,f406])).
fof(f730,plain,(
  ( ! [X0] : (~zex(X0) | difference(sK1,X0) = sK1) )),
  inference(resolution,[],[f503,f606])).
fof(f741,plain,(
  difference(sK1,sK23) = sK1),
  inference(resolution,[],[f730,f584])).
fof(f891,plain,(
  ( ! [X0,X1] : (cont(X0,X1) | zex(X0) | p(difference(X0,X1),X0)) )),
  inference(resolution,[],[f513,f482])).
fof(f900,plain,(
  zex(sK1) | p(difference(sK1,sK23),sK1)),
  inference(superposition,[],[f513,f741])).
fof(f901,plain,(
  p(difference(sK1,sK23),sK1)),
  inference(subsumption_resolution,[],[f900,f406])).
fof(f902,plain,(
  p(sK1,sK1)),
  inference(forward_demodulation,[],[f901,f741])).
fof(f907,plain,(
  cont(sK1,sK1)),
  inference(resolution,[],[f902,f477])).
fof(f915,plain,(
  ( ! [X2,X1] : (intersection(X1,X2) = X1 | ~cont(X1,intersection(X1,X2)) | zex(intersection(X1,X2))) )),
  inference(resolution,[],[f537,f512])).
fof(f917,plain,(
  ( ! [X2,X1] : (~cont(X1,intersection(X1,X2)) | intersection(X1,X2) = X1) )),
  inference(subsumption_resolution,[],[f915,f490])).
fof(f1165,plain,(
  ( ! [X5] : (~leq(X5,sK1) | cont(sK1,sum(X5,sK1))) )),
  inference(resolution,[],[f567,f907])).
fof(f1169,plain,(
  cont(sK1,sum(sK1,sK1))),
  inference(resolution,[],[f1165,f410])).
fof(f1183,plain,(
  ~eqdim(sK1,sum(sK1,sK1)) | p(sK1,sum(sK1,sK1))),
  inference(resolution,[],[f1169,f479])).
fof(f1184,plain,(
  ~zex(sum(sK1,sK1))),
  inference(resolution,[],[f1169,f490])).
fof(f1186,plain,(
  lt(sK1,sum(sK1,sK1)) | eqdim(sK1,sum(sK1,sK1))),
  inference(resolution,[],[f1169,f505])).
fof(f1191,plain,(
  ~eqdim(sK1,sum(sK1,sK1))),
  inference(subsumption_resolution,[],[f1183,f407])).
fof(f1192,plain,(
  lt(sK1,sum(sK1,sK1))),
  inference(subsumption_resolution,[],[f1186,f1191])).
fof(f1214,plain,(
  ~leq(sum(sK1,sK1),sK1)),
  inference(resolution,[],[f1192,f475])).
fof(f1218,plain,(
  sum(sK1,sK1) = difference(sum(sK1,sK1),sK1)),
  inference(resolution,[],[f1192,f503])).
fof(f1317,plain,(
  ( ! [X2] : (~p(X2,sK1) | p(X2,intersection(sK1,difference(sK1,sK23)))) )),
  inference(superposition,[],[f559,f741])).
fof(f1320,plain,(
  ( ! [X2] : (~p(X2,sK1) | p(X2,intersection(sK1,sK1))) )),
  inference(forward_demodulation,[],[f1317,f741])).
fof(f1482,plain,(
  p(sK1,intersection(sK1,sK1))),
  inference(resolution,[],[f1320,f902])).
fof(f1495,plain,(
  cont(sK1,intersection(sK1,sK1))),
  inference(resolution,[],[f1482,f477])).
fof(f2420,plain,(
  intersection(sK1,sK1) = sK1),
  inference(resolution,[],[f917,f1495])).
fof(f2649,plain,(
  ( ! [X7] : (leq(X7,sK1) | ~cont(X7,sK1) | ~cont(X7,sK1)) )),
  inference(superposition,[],[f575,f2420])).
fof(f2654,plain,(
  ( ! [X7] : (leq(X7,sK1) | ~cont(X7,sK1)) )),
  inference(duplicate_literal_removal,[],[f2649])).
fof(f2738,plain,(
  ~cont(sum(sK1,sK1),sK1)),
  inference(resolution,[],[f2654,f1214])).
fof(f2776,plain,(
  zex(sum(sK1,sK1)) | p(difference(sum(sK1,sK1),sK1),sum(sK1,sK1))),
  inference(resolution,[],[f2738,f891])).
fof(f2778,plain,(
  p(difference(sum(sK1,sK1),sK1),sum(sK1,sK1))),
  inference(subsumption_resolution,[],[f2776,f1184])).
fof(f2779,plain,(
  p(sum(sK1,sK1),sum(sK1,sK1))),
  inference(forward_demodulation,[],[f2778,f1218])).
fof(f3829,plain,(
  cont(sum(sK1,sK1),sum(sK1,sK1))),
  inference(resolution,[],[f2779,f477])).
fof(f4108,plain,(
  cont(sum(sK1,sK1),sK1) | cont(difference(sum(sK1,sK1),sK1),sK1)),
  inference(resolution,[],[f3829,f580])).
fof(f4119,plain,(
  cont(difference(sum(sK1,sK1),sK1),sK1)),
  inference(subsumption_resolution,[],[f4108,f2738])).
fof(f4120,plain,(
  cont(sum(sK1,sK1),sK1)),
  inference(forward_demodulation,[],[f4119,f1218])).
fof(f4121,plain,(
  $false),
  inference(subsumption_resolution,[],[f4120,f2738])).
% SZS output end Proof for codi_down_sum_theorems_goal02.all
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 9338
% Time elapsed: 0.792 s
% ------------------------------
% ------------------------------
% Success in time 61.84 s
