Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_down_theoremsT6-T10_goal07.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1464
Passive clauses: 8971
Generated clauses: 16585
Final active clauses: 1447
Final passive clauses: 7087
Input formulas: 67
Initial clauses: 108

Selected by SInE selection: 67
SInE iterations: 1

Duplicate literals: 114
Fw subsumption resolutions: 555

Simple tautologies: 70
Forward subsumptions: 6911

Binary resolution: 15509
Factoring: 25

Backtracking splits: 243
Backtracking splits refuted: 21
Backtracking splits refuted at zero level: 8

Memory used: 3198KB
Time elapsed: 0.404 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_down_theoremsT6-T10_goal07.all
Refutation found. Thanks to Tanya!
384. $false (2:0) [subsumption resolution 382,369]
369. ~zex(sK0) (1:2) [resolution 296,351]
351. ~zex(difference(sK0,sK0)) (0:4) [forward demodulation 240,239]
239. sK0 = sK1 (0:3) [cnf transformation 221]
221. sK0 = sK1 & ~zex(difference(sK0,sK1))[skolemisation 137]
137. ? [X0,X1] : (X0 = X1 & ~zex(difference(X0,X1)))[ennf transformation 68]
68. ~! [X0,X1] : (~X0 = X1 | zex(difference(X0,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : (~X0 = X1 | zex(difference(X0,X1)))[negated conjecture]
240. ~zex(difference(sK0,sK1)) (0:4) [cnf transformation 221]
296. zex(difference(X0,X1)) | ~zex(X0) (0:6) [cnf transformation 168]
168. ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))[ennf transformation 96]
96. ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & (~(zex(X0) | cont(X0,X1)) | zex(difference(X0,X1))))[flattening 12]
12. ! [X0] : ! [X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & (~(zex(X0) | cont(X0,X1)) | zex(difference(X0,X1))))[input]
382. zex(sK0) (2:2) [resolution 381,242]
242. cont(X0,X0) | zex(X0) (0:5) [cnf transformation 138]
138. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[ennf transformation 19]
19. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[input]
381. ~cont(sK0,sK0) (1:3) [resolution 297,351]
297. zex(difference(X0,X1)) | ~cont(X0,X1) (0:7) [cnf transformation 168]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 50
Passive clauses: 128
Generated clauses: 141
Final active clauses: 50
Final passive clauses: 77
Input formulas: 67
Initial clauses: 107

Selected by SInE selection: 67
SInE iterations: 1

Fw subsumption resolutions: 1
Fw demodulations: 1

Simple tautologies: 6
Equational tautologies: 1
Forward subsumptions: 3

Binary resolution: 31
Factoring: 1

Unique components: 2

Memory used: 127KB
Time elapsed: 0.005 s
------------------------------
% Success in time 0.453 s
