Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal09.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 603
Passive clauses: 2531
Generated clauses: 5088
Final active clauses: 592
Final passive clauses: 1793
Input formulas: 43
Initial clauses: 90

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 54
Fw subsumption resolutions: 169

Simple tautologies: 60
Forward subsumptions: 2300

Binary resolution: 4630
Factoring: 27

Backtracking splits: 76
Backtracking splits refuted: 15

Memory used: 1151KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 5994 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal09.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1039
Passive clauses: 4518
Generated clauses: 11152
Final active clauses: 1039
Final passive clauses: 3211
Input formulas: 43
Initial clauses: 85

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 204
Fw subsumption resolutions: 583

Simple tautologies: 164
Equational tautologies: 5
Forward subsumptions: 5962

Binary resolution: 7365
Factoring: 125
Forward superposition: 2748
Backward superposition: 62

Unique components: 3

Memory used: 1791KB
Time elapsed: 1.403 s
------------------------------
% remaining time: 5978 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal09.all
Refutation found. Thanks to Tanya!
12458. $false (5:0) [subsumption resolution 12389,12451]
12451. ~lt(sK0,sK10(sK0,sK1)) (3:5) [resolution 12391,236]
236. ~lt(X0,X1) | ~lt(X1,X0) (0:6) [cnf transformation 134]
134. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[ennf transformation 64]
64. ! [X0,X1] : (lt(X0,X1) => ~lt(X1,X0))[flattening 63]
63. ! [X0] : ! [X1] : (lt(X0,X1) => ~lt(X1,X0))[rectify 30]
30. ! [X62] : ! [X63] : (lt(X62,X63) => ~lt(X63,X62))[input]
12391. lt(sK10(sK0,sK1),sK0) (2:5) [resolution 12375,265]
265. ~sP0(X0,X1) | lt(sK10(X1,X0),X1) (0:8) [cnf transformation 199]
199. ! [X0,X1] : ((~sP0(X0,X1) | (lt(sK10(X1,X0),X1) & cont(sK10(X1,X0),X1) & p(sK10(X1,X0),X0))) & (! [X3] : (~lt(X3,X1) | ~cont(X3,X1) | ~p(X3,X0)) | sP0(X0,X1)))[skolemisation 198]
198. ! [X0,X1] : ((~sP0(X0,X1) | ? [X2] : (lt(X2,X1) & cont(X2,X1) & p(X2,X0))) & (! [X3] : (~lt(X3,X1) | ~cont(X3,X1) | ~p(X3,X0)) | sP0(X0,X1)))[rectify 197]
197. ! [X1,X0] : ((~sP0(X1,X0) | ? [X2] : (lt(X2,X0) & cont(X2,X0) & p(X2,X1))) & (! [X2] : (~lt(X2,X0) | ~cont(X2,X0) | ~p(X2,X1)) | sP0(X1,X0)))[nnf transformation 167]
167. ! [X1,X0] : (sP0(X1,X0) <=> ? [X2] : (lt(X2,X0) & cont(X2,X0) & p(X2,X1)))[predicate definition introduction]
12375. sP0(sK1,sK0) (1:3) [backtracking split refutation 298,301,12372]
12372. $false {2} (5:0) [subsumption resolution 12169,12365]
12365. ~lt(sK0,sK10(sK0,sK1)) {2} (3:5) [resolution 12171,236]
12171. lt(sK10(sK0,sK1),sK0) {2} (2:5) [resolution 12129,265]
12129. sP0(sK1,sK0) {2} (1:3) [backtracking split refutation 299,302,12126]
12126. $false {2, 3} (1:0) [subsumption resolution 302,11631]
11631. ~p(sK11(sK1,sK0),sK0) {2, 3} (4:5) [resolution 11107,263]
263. eqdim(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 196]
196. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 195]
195. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[nnf transformation 92]
92. ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[flattening 91]
91. ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[rectify 7]
7. ! [X14] : ! [X15] : (p(X14,X15) <=> (cont(X14,X15) & eqdim(X14,X15)))[input]
11107. ~eqdim(sK11(sK1,sK0),sK0) {2, 3} (3:5) [backtracking split refutation 453,459,11103]
11103. $false {2, 3, 14} (5:0) [subsumption resolution 11034,11099]
11099. lt(sK11(sK1,sK0),intersection(sK0,sK1)) {2, 3, 14} (5:7) [backtracking split refutation 11015,11019,11098]
11098. $false {2, 3, 302} (1:0) [subsumption resolution 301,11096]
11096. ~cont(sK11(sK1,sK0),sK1) {3, 302} (6:5) [backtracking split refutation 11087,11094,11095]
11095. $false {3, 306} (2:0) [subsumption resolution 399,11094]
399. cont(sK11(sK1,sK0),sK0) {3} (2:5) [resolution 302,262]
262. ~p(X0,X1) | cont(X0,X1) (0:6) [cnf transformation 196]
11094. ~cont(sK11(sK1,sK0),sK0) {306} (6:5) [backtracking splitting component]
11087. ~cont(sK11(sK1,sK0),sK1) | ~cont(sK11(sK1,sK0),sK0) {302} (6:10) [resolution 11019,285]
285. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 159]
159. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 158]
158. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 108]
108. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[flattening 107]
107. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[rectify 26]
26. ! [X53] : ! [X54] : ! [X55] : ((cont(X55,X53) & cont(X55,X54)) => leq(X55,intersection(X53,X54)))[input]
11019. ~leq(sK11(sK1,sK0),intersection(sK0,sK1)) {302} (5:7) [backtracking splitting component]
11015. lt(sK11(sK1,sK0),intersection(sK0,sK1)) | ~leq(sK11(sK1,sK0),intersection(sK0,sK1)) {14} (5:14) [resolution 1942,279]
279. eqdim(X0,X1) | lt(X0,X1) | ~leq(X0,X1) (0:9) [cnf transformation 207]
207. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | lt(X0,X1)))[flattening 206]
206. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | lt(X0,X1)))[nnf transformation 150]
150. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[ennf transformation 98]
98. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[flattening 97]
97. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[rectify 6]
6. ! [X12] : ! [X13] : (lt(X12,X13) <=> (leq(X12,X13) & ~eqdim(X12,X13)))[input]
1942. ~eqdim(sK11(sK1,sK0),intersection(sK0,sK1)) {14} (4:7) [resolution 728,459]
728. ~eqdim(X5,sK0) | ~eqdim(X5,intersection(sK0,sK1)) (2:8) [resolution 360,234]
234. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 132]
132. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 60]
60. ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))[flattening 59]
59. ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))[rectify 33]
33. ! [X68] : ! [X69] : (eqdim(X68,X69) => eqdim(X69,X68))[input]
360. ~eqdim(intersection(sK0,sK1),X0) | ~eqdim(X0,sK0) (1:8) [resolution 216,287]
287. eqdim(X0,X2) | ~eqdim(X1,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 163]
163. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~eqdim(X1,X2) | eqdim(X0,X2))[flattening 162]
162. ! [X0,X1,X2] : ((~eqdim(X0,X1) | ~eqdim(X1,X2)) | eqdim(X0,X2))[ennf transformation 112]
112. ! [X0,X1,X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))[flattening 111]
111. ! [X0] : ! [X1] : ! [X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))[rectify 34]
34. ! [X70] : ! [X71] : ! [X72] : ((eqdim(X70,X71) & eqdim(X71,X72)) => eqdim(X70,X72))[input]
216. ~eqdim(intersection(sK0,sK1),sK0) (0:5) [cnf transformation 169]
169. inc(sK0,sK1) & lt(sK0,sK1) & ~eqdim(intersection(sK0,sK1),sK0)[skolemisation 121]
121. ? [X0,X1] : (inc(X0,X1) & lt(X0,X1) & ~eqdim(intersection(X0,X1),X0))[flattening 120]
120. ? [X0,X1] : ((inc(X0,X1) & lt(X0,X1)) & ~eqdim(intersection(X0,X1),X0))[ennf transformation 44]
44. ~! [X0,X1] : ((inc(X0,X1) & lt(X0,X1)) => eqdim(intersection(X0,X1),X0))[flattening 1]
1. ~! [X0] : ! [X1] : ((inc(X0,X1) & lt(X0,X1)) => eqdim(intersection(X0,X1),X0))[negated conjecture]
11034. ~lt(sK11(sK1,sK0),intersection(sK0,sK1)) {14} (5:7) [resolution 10755,236]
10755. lt(intersection(sK0,sK1),sK11(sK1,sK0)) {14} (4:7) [resolution 459,1746]
1746. ~eqdim(X5,sK0) | lt(intersection(sK0,sK1),X5) (3:8) [resolution 585,234]
585. ~eqdim(sK0,X2) | lt(intersection(sK0,sK1),X2) (2:8) [resolution 569,282]
282. ~lt(X2,X0) | lt(X2,X1) | ~eqdim(X0,X1) (0:9) [cnf transformation 153]
153. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[flattening 152]
152. ! [X0,X1,X2] : (~eqdim(X0,X1) | (~lt(X2,X0) | lt(X2,X1)))[ennf transformation 102]
102. ! [X0,X1,X2] : (eqdim(X0,X1) => (lt(X2,X0) => lt(X2,X1)))[flattening 101]
101. ! [X0] : ! [X1] : ! [X2] : (eqdim(X0,X1) => (lt(X2,X0) => lt(X2,X1)))[rectify 36]
36. ! [X75] : ! [X76] : ! [X77] : (eqdim(X75,X76) => (lt(X77,X75) => lt(X77,X76)))[input]
569. lt(intersection(sK0,sK1),sK0) (1:5) [backtracking split refutation 361,365,566]
566. $false {8} (2:0) [subsumption resolution 434,547]
547. lt(intersection(sK0,sK1),sK0) {8} (4:5) [backtracking split refutation 535,540,541]
541. $false {20} (0:0) [subsumption resolution 216,540]
540. eqdim(intersection(sK0,sK1),sK0) {20} (4:5) [backtracking splitting component]
535. lt(intersection(sK0,sK1),sK0) | eqdim(intersection(sK0,sK1),sK0) {8} (4:10) [resolution 440,232]
232. ~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1) (0:9) [cnf transformation 129]
129. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 128]
128. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 56]
56. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 55]
55. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 28]
28. ! [X59] : ! [X60] : (cont(X59,X60) => (lt(X59,X60) | eqdim(X59,X60)))[input]
440. cont(intersection(sK0,sK1),sK0) {8} (3:5) [resolution 437,242]
242. zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0) (0:9) [cnf transformation 145]
145. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 76]
76. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 75]
75. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 25]
25. ! [X51] : ! [X52] : (~zex(intersection(X51,X52)) => cont(intersection(X51,X52),X51))[input]
437. ~zex(intersection(sK0,sK1)) {8} (2:4) [resolution 365,229]
229. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 126]
126. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 52]
52. ! [X0,X1] : (zex(X0) => leq(X0,X1))[flattening 51]
51. ! [X0] : ! [X1] : (zex(X0) => leq(X0,X1))[rectify 16]
16. ! [X32] : ! [X33] : (zex(X32) => leq(X32,X33))[input]
434. ~lt(intersection(sK0,sK1),sK0) {8} (2:5) [resolution 365,277]
277. leq(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 207]
365. ~leq(intersection(sK0,sK1),sK0) {8} (1:5) [backtracking splitting component]
361. lt(intersection(sK0,sK1),sK0) | ~leq(intersection(sK0,sK1),sK0) (1:10) [resolution 216,279]
459. eqdim(sK11(sK1,sK0),sK0) {14} (3:5) [backtracking splitting component]
453. lt(sK11(sK1,sK0),sK0) | eqdim(sK11(sK1,sK0),sK0) {3} (3:10) [resolution 399,232]
302. p(sK11(sK1,sK0),sK0) {3} (1:5) [backtracking splitting component]
299. sP0(sK1,sK0) | p(sK11(sK1,sK0),sK0) (1:8) [resolution 214,271]
271. ~inc(X0,X1) | sP0(X1,X0) | p(sK11(X1,X0),X0) (0:11) [cnf transformation 203]
203. ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | (lt(sK11(X1,X0),X1) & cont(sK11(X1,X0),X1) & p(sK11(X1,X0),X0))) & ((~sP0(X1,X0) & ! [X3] : (~lt(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))[skolemisation 202]
202. ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | ? [X2] : (lt(X2,X1) & cont(X2,X1) & p(X2,X0))) & ((~sP0(X1,X0) & ! [X3] : (~lt(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))[rectify 201]
201. ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | ? [X3] : (lt(X3,X1) & cont(X3,X1) & p(X3,X0))) & ((~sP0(X1,X0) & ! [X3] : (~lt(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))[flattening 200]
200. ! [X0,X1] : ((~inc(X0,X1) | (sP0(X1,X0) | ? [X3] : (lt(X3,X1) & cont(X3,X1) & p(X3,X0)))) & ((~sP0(X1,X0) & ! [X3] : (~lt(X3,X1) | ~cont(X3,X1) | ~p(X3,X0))) | inc(X0,X1)))[nnf transformation 168]
168. ! [X0,X1] : (inc(X0,X1) <=> (sP0(X1,X0) | ? [X3] : (lt(X3,X1) & cont(X3,X1) & p(X3,X0))))[definition folding 167]
214. inc(sK0,sK1) (0:3) [cnf transformation 169]
12169. lt(sK0,sK10(sK0,sK1)) {2} (5:5) [resolution 12129,779]
779. ~sP0(sK1,X3) | lt(sK0,sK10(X3,sK1)) (4:8) [resolution 748,267]
267. p(sK10(X1,X0),X0) | ~sP0(X0,X1) (0:8) [cnf transformation 199]
748. ~p(X4,sK1) | lt(sK0,X4) (3:6) [resolution 509,263]
509. ~eqdim(X5,sK1) | lt(sK0,X5) (2:6) [resolution 308,234]
308. ~eqdim(sK1,X2) | lt(sK0,X2) (1:6) [resolution 215,282]
215. lt(sK0,sK1) (0:3) [cnf transformation 169]
301. cont(sK11(sK1,sK0),sK1) {2} (1:5) [backtracking splitting component]
298. sP0(sK1,sK0) | cont(sK11(sK1,sK0),sK1) (1:8) [resolution 214,270]
270. ~inc(X0,X1) | sP0(X1,X0) | cont(sK11(X1,X0),X1) (0:11) [cnf transformation 203]
12389. lt(sK0,sK10(sK0,sK1)) (5:5) [resolution 12375,779]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1993
Passive clauses: 4531
Generated clauses: 13478
Final active clauses: 526
Final passive clauses: 549
Input formulas: 43
Initial clauses: 83

Unused predicate definitions: 2

Duplicate literals: 26
Trivial inequalities: 4
Bw subsumption resolutions: 261
Fw demodulations: 1136

Simple tautologies: 69
Equational tautologies: 136
Forward subsumptions: 6927

Binary resolution: 9516
Factoring: 2
Forward superposition: 27
Backward superposition: 442

Backtracking splits: 370
Backtracking splits refuted: 189
Backtracking splits refuted at zero level: 5

Memory used: 2430KB
Time elapsed: 2.125 s
------------------------------
% Success in time 4.522 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 23:33:39 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal09.all.tptp"
Input read from 
============================ end of footer ===========================
