Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal14.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 592
Passive clauses: 2345
Generated clauses: 4959
Final active clauses: 592
Final passive clauses: 1671
Input formulas: 43
Initial clauses: 90

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 32
Fw subsumption resolutions: 42

Simple tautologies: 30
Forward subsumptions: 2568

Binary resolution: 4854
Factoring: 41

Backtracking splits: 26

Memory used: 1023KB
Time elapsed: 0.404 s
------------------------------
% remaining time: 5994 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal14.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 927
Passive clauses: 4657
Generated clauses: 10817
Final active clauses: 927
Final passive clauses: 3495
Input formulas: 43
Initial clauses: 85

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 201
Fw subsumption resolutions: 710

Simple tautologies: 147
Equational tautologies: 6
Forward subsumptions: 5262

Binary resolution: 7691
Factoring: 119
Forward superposition: 2002
Backward superposition: 30

Unique components: 3

Memory used: 1918KB
Time elapsed: 1.404 s
------------------------------
% remaining time: 5978 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal14.all
Refutation found. Thanks to Tanya!
384. $false (3:0) [subsumption resolution 349,376]
376. ~lt(intersection(sK0,sK0),sK0) (3:5) [resolution 369,235]
235. ~lt(X0,X1) | ~lt(X1,X0) (0:6) [cnf transformation 133]
133. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[ennf transformation 63]
63. ! [X0,X1] : (lt(X0,X1) => ~lt(X1,X0))[flattening 62]
62. ! [X0] : ! [X1] : (lt(X0,X1) => ~lt(X1,X0))[rectify 30]
30. ! [X61] : ! [X62] : (lt(X61,X62) => ~lt(X62,X61))[input]
369. lt(sK0,intersection(sK0,sK0)) (2:5) [backtracking split refutation 330,367,368]
368. $false {8} (2:0) [subsumption resolution 367,366]
366. leq(sK0,intersection(sK0,sK0)) (2:5) [backtracking split refutation 329,332,364]
364. $false {3} (1:0) [subsumption resolution 296,363]
363. ~cont(sK0,sK0) {3} (3:3) [duplicate literal removal 356]
356. ~cont(sK0,sK0) | ~cont(sK0,sK0) {3} (3:6) [resolution 332,284]
284. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 158]
158. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 157]
157. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 107]
107. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[flattening 106]
106. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[rectify 26]
26. ! [X52] : ! [X53] : ! [X54] : ((cont(X54,X52) & cont(X54,X53)) => leq(X54,intersection(X52,X53)))[input]
296. cont(sK0,sK0) (1:3) [resolution 213,225]
225. zex(X0) | cont(X0,X0) (0:5) [cnf transformation 177]
177. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[nnf transformation 123]
123. ! [X0] : (~zex(X0) <=> cont(X0,X0))[ennf transformation 48]
48. ! [X0] : (~zex(X0) <=> cont(X0,X0))[rectify 17]
17. ! [X33] : (~zex(X33) <=> cont(X33,X33))[input]
213. ~zex(sK0) (0:2) [cnf transformation 168]
168. ~zex(sK0) & ~zex(intersection(sK0,sK0)) & ~p(sK0,intersection(sK0,sK0))[skolemisation 120]
120. ? [X0] : (~zex(X0) & ~zex(intersection(X0,X0)) & ~p(X0,intersection(X0,X0)))[flattening 119]
119. ? [X0] : ((~zex(X0) & ~zex(intersection(X0,X0))) & ~p(X0,intersection(X0,X0)))[ennf transformation 1]
1. ~! [X0] : ((~zex(X0) & ~zex(intersection(X0,X0))) => p(X0,intersection(X0,X0)))[negated conjecture]
332. ~leq(sK0,intersection(sK0,sK0)) {3} (2:5) [backtracking splitting component]
329. ~leq(intersection(sK0,sK0),sK0) | ~leq(sK0,intersection(sK0,sK0)) (2:10) [resolution 313,275]
275. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 204]
204. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 203]
203. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[nnf transformation 95]
95. ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[flattening 94]
94. ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[rectify 41]
41. ! [X86] : ! [X87] : (eqdim(X86,X87) <=> (leq(X86,X87) & leq(X87,X86)))[input]
313. ~eqdim(sK0,intersection(sK0,sK0)) (1:5) [backtracking split refutation 309,310,311]
311. $false {1} (1:0) [subsumption resolution 296,310]
310. ~cont(sK0,sK0) {1} (1:3) [backtracking splitting component]
309. ~eqdim(sK0,intersection(sK0,sK0)) | ~cont(sK0,sK0) (1:8) [duplicate literal removal 307]
307. ~eqdim(sK0,intersection(sK0,sK0)) | ~cont(sK0,sK0) | ~cont(sK0,sK0) (1:11) [resolution 215,288]
288. p(X2,intersection(X0,X1)) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:16) [cnf transformation 209]
209. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 208]
208. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[nnf transformation 115]
115. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[flattening 114]
114. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[rectify 27]
27. ! [X55] : ! [X56] : ! [X57] : ((cont(X57,X55) & cont(X57,X56) & eqdim(X57,intersection(X55,X56))) <=> p(X57,intersection(X55,X56)))[input]
215. ~p(sK0,intersection(sK0,sK0)) (0:5) [cnf transformation 168]
367. ~leq(sK0,intersection(sK0,sK0)) {8} (2:5) [backtracking splitting component]
330. lt(sK0,intersection(sK0,sK0)) | ~leq(sK0,intersection(sK0,sK0)) (2:10) [resolution 313,278]
278. eqdim(X0,X1) | lt(X0,X1) | ~leq(X0,X1) (0:9) [cnf transformation 206]
206. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | lt(X0,X1)))[flattening 205]
205. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | lt(X0,X1)))[nnf transformation 149]
149. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[ennf transformation 97]
97. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[flattening 96]
96. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[rectify 6]
6. ! [X11] : ! [X12] : (lt(X11,X12) <=> (leq(X11,X12) & ~eqdim(X11,X12)))[input]
349. lt(intersection(sK0,sK0),sK0) (3:5) [backtracking split refutation 339,346,347]
347. $false {5} (1:0) [subsumption resolution 298,346]
298. cont(intersection(sK0,sK0),sK0) (1:5) [resolution 214,241]
241. zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0) (0:9) [cnf transformation 144]
144. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 75]
75. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 74]
74. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 25]
25. ! [X50] : ! [X51] : (~zex(intersection(X50,X51)) => cont(intersection(X50,X51),X50))[input]
214. ~zex(intersection(sK0,sK0)) (0:4) [cnf transformation 168]
346. ~cont(intersection(sK0,sK0),sK0) {5} (3:5) [backtracking splitting component]
339. lt(intersection(sK0,sK0),sK0) | ~cont(intersection(sK0,sK0),sK0) (3:10) [resolution 327,231]
231. eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 128]
128. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 127]
127. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 55]
55. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 54]
54. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 28]
28. ! [X58] : ! [X59] : (cont(X58,X59) => (lt(X58,X59) | eqdim(X58,X59)))[input]
327. ~eqdim(intersection(sK0,sK0),sK0) (2:5) [resolution 313,233]
233. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 131]
131. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 59]
59. ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))[flattening 58]
58. ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))[rectify 33]
33. ! [X67] : ! [X68] : (eqdim(X67,X68) => eqdim(X68,X67))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 95
Passive clauses: 35
Generated clauses: 100
Final active clauses: 91
Final passive clauses: 10
Input formulas: 43
Initial clauses: 83

Unused predicate definitions: 2

Duplicate literals: 2
Bw subsumption resolutions: 7
Fw demodulations: 1

Simple tautologies: 3
Equational tautologies: 3
Forward subsumptions: 34

Binary resolution: 61

Backtracking splits: 8
Backtracking splits refuted: 5
Backtracking splits refuted at zero level: 5

Memory used: 255KB
Time elapsed: 0.015 s
------------------------------
% Success in time 2.352 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 23:34:20 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal14.all.tptp"
Input read from 
============================ end of footer ===========================
