Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal12.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 534
Passive clauses: 2176
Generated clauses: 3808
Final active clauses: 532
Final passive clauses: 1582
Input formulas: 43
Initial clauses: 91

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 24
Fw subsumption resolutions: 86

Simple tautologies: 29
Forward subsumptions: 1480

Binary resolution: 3529
Factoring: 9

Backtracking splits: 59
Backtracking splits refuted: 8
Backtracking splits refuted at zero level: 2

Memory used: 895KB
Time elapsed: 0.403 s
------------------------------
% remaining time: 5994 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal12.all
Refutation found. Thanks to Tanya!
3217. $false (7:0) [subsumption resolution 3212,214]
214. p(sK2,sK0) (0:3) [cnf transformation 168]
168. po(sK0,sK1) & p(sK2,sK0) & p(sK2,sK1) & ~p(sK2,intersection(sK0,sK1))[skolemisation 120]
120. ? [X0,X1,X2] : (po(X0,X1) & p(X2,X0) & p(X2,X1) & ~p(X2,intersection(X0,X1)))[flattening 119]
119. ? [X0,X1,X2] : ((po(X0,X1) & p(X2,X0) & p(X2,X1)) & ~p(X2,intersection(X0,X1)))[ennf transformation 44]
44. ~! [X0,X1,X2] : ((po(X0,X1) & p(X2,X0) & p(X2,X1)) => p(X2,intersection(X0,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : ((po(X0,X1) & p(X2,X0) & p(X2,X1)) => p(X2,intersection(X0,X1)))[negated conjecture]
3212. ~p(sK2,sK0) (7:3) [resolution 3210,271]
271. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 202]
202. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 201]
201. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[nnf transformation 93]
93. ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[flattening 92]
92. ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[rectify 19]
19. ! [X37] : ! [X38] : (p(X37,X38) <=> (cont(X37,X38) & eqdim(X37,X38)))[input]
3210. ~cont(sK2,sK0) (6:3) [subsumption resolution 3209,215]
215. p(sK2,sK1) (0:3) [cnf transformation 168]
3209. ~cont(sK2,sK0) | ~p(sK2,sK1) (6:6) [subsumption resolution 3204,272]
272. eqdim(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 202]
3204. ~cont(sK2,sK0) | ~eqdim(sK2,sK1) | ~p(sK2,sK1) (6:9) [resolution 3185,271]
3185. ~cont(sK2,sK1) | ~cont(sK2,sK0) | ~eqdim(sK2,sK1) (5:9) [resolution 2780,1014]
1014. lt(sK2,intersection(sK0,sK1)) | ~cont(sK2,sK0) | ~cont(sK2,sK1) (2:11) [subsumption resolution 1001,285]
285. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 158]
158. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 157]
157. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 107]
107. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[flattening 106]
106. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[rectify 31]
31. ! [X62] : ! [X63] : ! [X64] : ((cont(X64,X62) & cont(X64,X63)) => leq(X64,intersection(X62,X63)))[input]
1001. ~cont(sK2,sK1) | ~cont(sK2,sK0) | lt(sK2,intersection(sK0,sK1)) | ~leq(sK2,intersection(sK0,sK1)) (2:16) [resolution 992,279]
279. eqdim(X0,X1) | lt(X0,X1) | ~leq(X0,X1) (0:9) [cnf transformation 206]
206. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | lt(X0,X1)))[flattening 205]
205. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | lt(X0,X1)))[nnf transformation 149]
149. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[ennf transformation 97]
97. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[flattening 96]
96. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[rectify 18]
18. ! [X35] : ! [X36] : (lt(X35,X36) <=> (leq(X35,X36) & ~eqdim(X35,X36)))[input]
992. ~eqdim(sK2,intersection(sK0,sK1)) | ~cont(sK2,sK1) | ~cont(sK2,sK0) (1:11) [resolution 289,216]
216. ~p(sK2,intersection(sK0,sK1)) (0:5) [cnf transformation 168]
289. p(X2,intersection(X0,X1)) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:16) [cnf transformation 209]
209. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 208]
208. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[nnf transformation 115]
115. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[flattening 114]
114. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[rectify 32]
32. ! [X65] : ! [X66] : ! [X67] : ((cont(X67,X65) & cont(X67,X66) & eqdim(X67,intersection(X65,X66))) <=> p(X67,intersection(X65,X66)))[input]
2780. ~lt(X10,intersection(X11,X12)) | ~eqdim(X10,X12) (4:8) [resolution 2766,282]
282. lt(X1,X2) | ~lt(X0,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 152]
152. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X0,X2) | lt(X1,X2))[flattening 151]
151. ! [X0,X1,X2] : (~eqdim(X0,X1) | (~lt(X0,X2) | lt(X1,X2)))[ennf transformation 101]
101. ! [X0,X1,X2] : (eqdim(X0,X1) => (lt(X0,X2) => lt(X1,X2)))[flattening 100]
100. ! [X0] : ! [X1] : ! [X2] : (eqdim(X0,X1) => (lt(X0,X2) => lt(X1,X2)))[rectify 10]
10. ! [X20] : ! [X21] : ! [X22] : (eqdim(X20,X21) => (lt(X20,X22) => lt(X21,X22)))[input]
2766. ~lt(X5,intersection(X4,X5)) (3:5) [subsumption resolution 2749,615]
615. ~lt(X9,X10) | ~zex(X10) (2:5) [resolution 609,229]
229. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 125]
125. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 52]
52. ! [X0,X1] : (zex(X0) => leq(X0,X1))[flattening 51]
51. ! [X0] : ! [X1] : (zex(X0) => leq(X0,X1))[rectify 40]
40. ! [X84] : ! [X85] : (zex(X84) => leq(X84,X85))[input]
609. ~leq(X2,X3) | ~lt(X3,X2) (1:6) [subsumption resolution 603,277]
277. leq(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 206]
603. ~leq(X2,X3) | ~leq(X3,X2) | ~lt(X3,X2) (1:9) [resolution 276,235]
235. ~eqdim(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 132]
132. ! [X0,X1] : (~lt(X0,X1) | ~eqdim(X0,X1))[ennf transformation 62]
62. ! [X0,X1] : (lt(X0,X1) => ~eqdim(X0,X1))[flattening 61]
61. ! [X0] : ! [X1] : (lt(X0,X1) => ~eqdim(X0,X1))[rectify 8]
8. ! [X15] : ! [X16] : (lt(X15,X16) => ~eqdim(X15,X16))[input]
276. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 204]
204. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 203]
203. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[nnf transformation 95]
95. ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[flattening 94]
94. ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[rectify 28]
28. ! [X56] : ! [X57] : (eqdim(X56,X57) <=> (leq(X56,X57) & leq(X57,X56)))[input]
2749. zex(intersection(X4,X5)) | ~lt(X5,intersection(X4,X5)) (3:9) [resolution 1570,530]
530. ~cont(X0,X1) | ~lt(X1,X0) (2:6) [subsumption resolution 527,236]
236. ~lt(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 133]
133. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[ennf transformation 64]
64. ! [X0,X1] : (lt(X0,X1) => ~lt(X1,X0))[flattening 63]
63. ! [X0] : ! [X1] : (lt(X0,X1) => ~lt(X1,X0))[rectify 3]
3. ! [X4] : ! [X5] : (lt(X4,X5) => ~lt(X5,X4))[input]
527. lt(X0,X1) | ~cont(X0,X1) | ~lt(X1,X0) (2:9) [resolution 232,310]
310. ~eqdim(X2,X1) | ~lt(X1,X2) (1:6) [resolution 235,234]
234. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 131]
131. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 60]
60. ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))[flattening 59]
59. ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))[rectify 6]
6. ! [X10] : ! [X11] : (eqdim(X10,X11) => eqdim(X11,X10))[input]
232. eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 128]
128. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 127]
127. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 56]
56. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 55]
55. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 36]
36. ! [X76] : ! [X77] : (cont(X76,X77) => (lt(X76,X77) | eqdim(X76,X77)))[input]
1570. cont(intersection(X0,X1),X1) | zex(intersection(X0,X1)) (2:9) [subsumption resolution 1548,218]
218. eqdim(X0,X0) (0:3) [cnf transformation 46]
46. ! [X0] : eqdim(X0,X0)[rectify 5]
5. ! [X9] : eqdim(X9,X9)[input]
1548. ~eqdim(intersection(X0,X1),intersection(X0,X1)) | cont(intersection(X0,X1),X1) | zex(intersection(X0,X1)) (2:16) [resolution 593,226]
226. cont(X0,X0) | zex(X0) (0:5) [cnf transformation 177]
177. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[nnf transformation 123]
123. ! [X0] : (~zex(X0) <=> cont(X0,X0))[ennf transformation 49]
49. ! [X0] : (~zex(X0) <=> cont(X0,X0))[rectify 23]
23. ! [X46] : (~zex(X46) <=> cont(X46,X46))[input]
593. ~cont(X5,intersection(X6,X7)) | ~eqdim(X5,intersection(X6,X7)) | cont(X5,X7) (1:13) [resolution 273,291]
291. ~p(X2,intersection(X0,X1)) | cont(X2,X1) (0:8) [cnf transformation 209]
273. p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 202]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 406
Passive clauses: 1446
Generated clauses: 2999
Final active clauses: 406
Final passive clauses: 944
Input formulas: 43
Initial clauses: 86

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 60
Fw subsumption resolutions: 247

Simple tautologies: 66
Equational tautologies: 4
Forward subsumptions: 1270

Binary resolution: 1802
Factoring: 33
Forward superposition: 730
Backward superposition: 44

Unique components: 4

Memory used: 767KB
Time elapsed: 0.384 s
------------------------------
% Success in time 1.13 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 00:58:32 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal12.all.tptp"
Input read from 
============================ end of footer ===========================
