Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_down_theoremsT6-T10_goal03.all
Refutation found. Thanks to Tanya!
2464. $false (1:0) [subsumption resolution 2463,644]
644. p(sK2,sK0) (1:3) [resolution 331,239]
239. p(sK2,difference(sK0,sK1)) (0:5) [cnf transformation 221]
221. p(sK2,difference(sK0,sK1)) & ~p(sK2,intersection(sK0,difference(sK0,sK1)))[skolemisation 137]
137. ? [X0,X1,X2] : (p(X2,difference(X0,X1)) & ~p(X2,intersection(X0,difference(X0,X1))))[ennf transformation 68]
68. ~! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,intersection(X0,difference(X0,X1))))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : (~p(X2,difference(X0,X1)) | p(X2,intersection(X0,difference(X0,X1))))[negated conjecture]
331. ~p(X2,difference(X0,X1)) | p(X2,X0) (0:8) [cnf transformation 199]
199. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))[ennf transformation 126]
126. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))[flattening 65]
65. ! [X0] : ! [X1] : ! [X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))[input]
2463. ~p(sK2,sK0) (1:3) [subsumption resolution 2452,239]
2452. ~p(sK2,difference(sK0,sK1)) | ~p(sK2,sK0) (1:8) [resolution 2415,240]
240. ~p(sK2,intersection(sK0,difference(sK0,sK1))) (0:7) [cnf transformation 221]
2415. p(X2,intersection(X0,X1)) | ~p(X2,X1) | ~p(X2,X0) (0:11) [subsumption resolution 333,279]
279. po(X0,X1) | ~p(X3,X1) | ~p(X3,X0) (0:9) [cnf transformation 235]
235. ! [X0,X1] : ((~po(X0,X1) | (p(sK8(X1,X0),X0) & p(sK8(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[skolemisation 158]
158. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[ennf transformation 89]
89. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[flattening 88]
88. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[rectify 4]
4. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X2] : (p(X2,X0) & p(X2,X1)) | po(X0,X1)))[input]
333. p(X2,intersection(X0,X1)) | ~p(X2,X1) | ~p(X2,X0) | ~po(X0,X1) (0:14) [cnf transformation 202]
202. ! [X0,X1,X2] : (~po(X0,X1) | ~p(X2,X0) | ~p(X2,X1) | p(X2,intersection(X0,X1)))[flattening 201]
201. ! [X0,X1,X2] : ((~po(X0,X1) | ~p(X2,X0) | ~p(X2,X1)) | p(X2,intersection(X0,X1)))[ennf transformation 128]
128. ! [X0,X1,X2] : (~(po(X0,X1) & p(X2,X0) & p(X2,X1)) | p(X2,intersection(X0,X1)))[flattening 46]
46. ! [X0] : ! [X1] : ! [X2] : (~(po(X0,X1) & p(X2,X0) & p(X2,X1)) | p(X2,intersection(X0,X1)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 442
Passive clauses: 1305
Generated clauses: 2220
Final active clauses: 442
Final passive clauses: 829
Input formulas: 67
Initial clauses: 108

Selected by SInE selection: 67
SInE iterations: 1

Duplicate literals: 28
Fw subsumption resolutions: 78

Simple tautologies: 24
Forward subsumptions: 791

Binary resolution: 1987
Factoring: 2

Backtracking splits: 17

Memory used: 639KB
Time elapsed: 0.049 s
------------------------------
% Success in time 0.075 s
