Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal16.all
Refutation found. Thanks to Tanya!
458. $false (4:0) [subsumption resolution 457,161]
161. zex(intersection(sK0,sK0)) (0:4) [cnf transformation 145]
145. zex(intersection(sK0,sK0)) & ~zex(sK0)[skolemisation 86]
86. ? [X0] : (zex(intersection(X0,X0)) & ~zex(X0))[ennf transformation 1]
1. ~! [X0] : (~zex(intersection(X0,X0)) | zex(X0))[negated conjecture]
457. ~zex(intersection(sK0,sK0)) (4:4) [resolution 453,176]
176. ~c(X0,X1) | ~zex(intersection(X0,X1)) (0:7) [cnf transformation 93]
93. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[ennf transformation 47]
47. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[flattening 2]
2. ! [X0] : ! [X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[input]
453. c(sK0,sK0) (3:3) [resolution 430,262]
262. cont(sK0,sK0) (1:3) [resolution 165,162]
162. ~zex(sK0) (0:2) [cnf transformation 145]
165. zex(X0) | cont(X0,X0) (0:5) [cnf transformation 87]
87. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[ennf transformation 22]
22. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[input]
430. ~cont(sK0,X0) | c(X0,sK0) (2:6) [resolution 206,262]
206. ~cont(X3,X1) | ~cont(X3,X0) | c(X0,X1) (0:9) [cnf transformation 156]
156. ! [X0,X1] : ((~c(X0,X1) | (cont(sK9(X1,X0),X0) & cont(sK9(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 103]
103. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 59]
59. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 58]
58. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 28]
28. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 106
Passive clauses: 197
Generated clauses: 290
Final active clauses: 106
Final passive clauses: 85
Input formulas: 44
Initial clauses: 92

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 3
Fw subsumption resolutions: 2

Simple tautologies: 8
Forward subsumptions: 84

Binary resolution: 189
Factoring: 3

Backtracking splits: 1

Memory used: 127KB
Time elapsed: 0.015 s
------------------------------
% Success in time 0.058 s
