% Starting sequential portfolio solving with schedule "casc"
% spawned child 10024 with time: 4 (total remaining time 2999)
% WARNING: value z3 for option sas not known
% dis+1010_4_afp=10000:afq=1.2:anc=none:irw=on:lma=on:nm=64:nwc=10:sas=z3:sac=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal05.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6780
% Time elapsed: 0.500 s
% spawned child 10025 with time: 10 (total remaining time 2994)
% ott+11_8_amm=off:anc=none:bsr=on:cond=on:irw=on:nm=2:nwc=1.1:ss=axioms:st=2.0:sac=on_8 on codi_down_sum_theorems_goal05.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7419
% Time elapsed: 1.100 s
% ------------------------------
% ------------------------------
% spawned child 10026 with time: 4 (total remaining time 2983)
% dis+1011_10_aac=none:add=large:afp=10000:afq=1.1:amm=off:anc=none:fsr=off:gs=on:gsem=off:irw=on:nm=6:newcnf=on:nwc=2.5:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal05.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 11769
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 10027 with time: 6 (total remaining time 2978)
% dis+11_28_av=off:fsr=off:irw=on:lcm=predicate:nm=2:newcnf=on:nwc=5:sp=occurrence:urr=on:updr=off_4 on codi_down_sum_theorems_goal05.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 3582
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% spawned child 10028 with time: 6 (total remaining time 2970)
% lrs+1011_5_add=large:afp=1000:afq=1.2:anc=none:fsr=off:irw=on:lma=on:nm=64:newcnf=on:nwc=1:stl=30:sd=3:ss=axioms:st=2.0:sos=on:sac=on:sp=reverse_arity:urr=on:updr=off_4 on codi_down_sum_theorems_goal05.all
% Refutation found. Thanks to Tanya!
% SZS status Theorem for codi_down_sum_theorems_goal05.all
% SZS output start Proof for codi_down_sum_theorems_goal05.all
fof(f1,conjecture,(
  ! [X0] : ! [X1] : (zex(X0) | ~zex(sum(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal05.all.tptp',conjecture50)).
fof(f2,negated_conjecture,(
  ~! [X0] : ! [X1] : (zex(X0) | ~zex(sum(X0,X1)))),
  inference(negated_conjecture,[],[f1])).
fof(f3,axiom,(
  ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal05.all.tptp',axiom10)).
fof(f5,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (cont(X2,sum(X0,X1)) | ~(cont(X2,X1) & leq(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal05.all.tptp',axiom30)).
fof(f84,axiom,(
  ! [X0] : ! [X1] : (lt(X0,X1) | ~(~zex(X1) & zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal05.all.tptp',axiom89500)).
fof(f85,axiom,(
  ! [X0] : ((~zex(X0) | ~cont(X0,X0)) & (cont(X0,X0) | zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal05.all.tptp',axiom89510)).
fof(f88,axiom,(
  ! [X0] : ! [X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal05.all.tptp',axiom89540)).
fof(f96,axiom,(
  ! [X0] : ! [X1] : (leq(X1,X0) | leq(X0,X1))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal05.all.tptp',axiom89110)).
fof(f97,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (leq(X0,X2) | ~(leq(X1,X2) & leq(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal05.all.tptp',axiom89120)).
fof(f98,axiom,(
  ! [X0] : ! [X1] : (X0 = X1 | ~(zex(X1) & zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal05.all.tptp',axiom89130)).
fof(f99,axiom,(
  ! [X0] : ! [X1] : (leq(X0,X1) | ~zex(X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal05.all.tptp',axiom89140)).
fof(f103,axiom,(
  ! [X0] : ! [X1] : ((lt(X0,X1) | ~(~leq(X1,X0) & leq(X0,X1))) & ((~leq(X1,X0) & leq(X0,X1)) | ~lt(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal05.all.tptp',axiom89280)).
fof(f107,axiom,(
  ? [X0] : zex(X0)),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal05.all.tptp',axiom87750)).
fof(f118,plain,(
  ~! [X0,X1] : (zex(X0) | ~zex(sum(X0,X1)))),
  inference(flattening,[],[f2])).
fof(f144,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | ~(cont(X2,X1) & leq(X0,X1)))),
  inference(flattening,[],[f5])).
fof(f146,plain,(
  ! [X0,X1] : sum(X0,X1) = sum(X1,X0)),
  inference(flattening,[],[f3])).
fof(f174,plain,(
  ! [X0,X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  inference(flattening,[],[f88])).
fof(f175,plain,(
  ! [X0,X1,X2] : (leq(X0,X2) | ~(leq(X1,X2) & leq(X0,X1)))),
  inference(flattening,[],[f97])).
fof(f176,plain,(
  ! [X0,X1] : (leq(X1,X0) | leq(X0,X1))),
  inference(flattening,[],[f96])).
fof(f177,plain,(
  ! [X0,X1] : (leq(X0,X1) | ~zex(X0))),
  inference(flattening,[],[f99])).
fof(f178,plain,(
  ! [X0,X1] : ((lt(X0,X1) | ~(~leq(X1,X0) & leq(X0,X1))) & ((~leq(X1,X0) & leq(X0,X1)) | ~lt(X0,X1)))),
  inference(flattening,[],[f103])).
fof(f183,plain,(
  ! [X0,X1] : (X0 = X1 | ~(zex(X1) & zex(X0)))),
  inference(flattening,[],[f98])).
fof(f191,plain,(
  ! [X0,X1] : (lt(X0,X1) | ~(~zex(X1) & zex(X0)))),
  inference(flattening,[],[f84])).
fof(f212,plain,(
  ? [X0,X1] : (~zex(X0) & zex(sum(X0,X1)))),
  inference(ennf_transformation,[],[f118])).
fof(f229,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | (~cont(X2,X1) | ~leq(X0,X1)))),
  inference(ennf_transformation,[],[f144])).
fof(f230,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | ~cont(X2,X1) | ~leq(X0,X1))),
  inference(flattening,[],[f229])).
fof(f264,plain,(
  ! [X0,X1,X2] : (leq(X0,X2) | (~leq(X1,X2) | ~leq(X0,X1)))),
  inference(ennf_transformation,[],[f175])).
fof(f265,plain,(
  ! [X0,X1,X2] : (leq(X0,X2) | ~leq(X1,X2) | ~leq(X0,X1))),
  inference(flattening,[],[f264])).
fof(f266,plain,(
  ! [X0,X1] : ((lt(X0,X1) | (leq(X1,X0) | ~leq(X0,X1))) & ((~leq(X1,X0) & leq(X0,X1)) | ~lt(X0,X1)))),
  inference(ennf_transformation,[],[f178])).
fof(f267,plain,(
  ! [X0,X1] : ((lt(X0,X1) | leq(X1,X0) | ~leq(X0,X1)) & ((~leq(X1,X0) & leq(X0,X1)) | ~lt(X0,X1)))),
  inference(flattening,[],[f266])).
fof(f275,plain,(
  ! [X0,X1] : (X0 = X1 | (~zex(X1) | ~zex(X0)))),
  inference(ennf_transformation,[],[f183])).
fof(f276,plain,(
  ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0))),
  inference(flattening,[],[f275])).
fof(f281,plain,(
  ! [X0,X1] : (lt(X0,X1) | (zex(X1) | ~zex(X0)))),
  inference(ennf_transformation,[],[f191])).
fof(f282,plain,(
  ! [X0,X1] : (lt(X0,X1) | zex(X1) | ~zex(X0))),
  inference(flattening,[],[f281])).
fof(f304,plain,(
  zex(sum(sK0,sK1))),
  inference(cnf_transformation,[],[f212])).
fof(f305,plain,(
  ~zex(sK0)),
  inference(cnf_transformation,[],[f212])).
fof(f306,plain,(
  zex(sK2)),
  inference(cnf_transformation,[],[f107])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (cont(X2,sum(X0,X1)) | ~cont(X2,X1) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f332,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = sum(X1,X0)) )),
  inference(cnf_transformation,[],[f146])).
fof(f371,plain,(
  ( ! [X0,X1] : (~zex(X0) | ~cont(X1,X0)) )),
  inference(cnf_transformation,[],[f174])).
fof(f374,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f85])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2)) )),
  inference(cnf_transformation,[],[f265])).
fof(f376,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | leq(X1,X0)) )),
  inference(cnf_transformation,[],[f176])).
fof(f377,plain,(
  ( ! [X0,X1] : (~zex(X0) | leq(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f379,plain,(
  ( ! [X0,X1] : (~lt(X0,X1) | ~leq(X1,X0)) )),
  inference(cnf_transformation,[],[f267])).
fof(f389,plain,(
  ( ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1) )),
  inference(cnf_transformation,[],[f276])).
fof(f401,plain,(
  ( ! [X0,X1] : (~zex(X0) | zex(X1) | lt(X0,X1)) )),
  inference(cnf_transformation,[],[f282])).
fof(f440,plain,(
  cont(sK0,sK0)),
  inference(unit_resulting_resolution,[],[f305,f374])).
fof(f442,plain,(
  lt(sK2,sK0)),
  inference(unit_resulting_resolution,[],[f306,f305,f401])).
fof(f468,plain,(
  sum(sK0,sK1) = sK2),
  inference(unit_resulting_resolution,[],[f306,f304,f389])).
fof(f551,plain,(
  ~leq(sK0,sK2)),
  inference(unit_resulting_resolution,[],[f442,f379])).
fof(f4065,plain,(
  spl14_8 <=> leq(sK0,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_8])])).
fof(f4066,plain,(
  leq(sK0,sK1) | ~spl14_8),
  inference(avatar_component_clause,[],[f4065])).
fof(f4068,plain,(
  spl14_9 <=> ~leq(sK0,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_9])])).
fof(f4069,plain,(
  ~leq(sK0,sK1) | ~spl14_9),
  inference(avatar_component_clause,[],[f4068])).
fof(f4080,plain,(
  leq(sK1,sK0) | ~spl14_9),
  inference(unit_resulting_resolution,[],[f4069,f376])).
fof(f4088,plain,(
  cont(sK0,sum(sK1,sK0)) | ~spl14_9),
  inference(unit_resulting_resolution,[],[f440,f4080,f330])).
fof(f4111,plain,(
  cont(sK0,sum(sK0,sK1)) | ~spl14_9),
  inference(forward_demodulation,[],[f4088,f332])).
fof(f4112,plain,(
  cont(sK0,sK2) | ~spl14_9),
  inference(forward_demodulation,[],[f4111,f468])).
fof(f4211,plain,(
  $false | ~spl14_9),
  inference(unit_resulting_resolution,[],[f306,f4112,f371])).
fof(f4222,plain,(
  spl14_9),
  inference(avatar_contradiction_clause,[],[f4211])).
fof(f5114,plain,(
  ~leq(sK1,sK2) | ~spl14_8),
  inference(unit_resulting_resolution,[],[f551,f4066,f375])).
fof(f5150,plain,(
  ~zex(sK1) | ~spl14_8),
  inference(unit_resulting_resolution,[],[f5114,f377])).
fof(f5162,plain,(
  cont(sK1,sK1) | ~spl14_8),
  inference(unit_resulting_resolution,[],[f5150,f374])).
fof(f5249,plain,(
  cont(sK1,sum(sK0,sK1)) | ~spl14_8),
  inference(unit_resulting_resolution,[],[f4066,f5162,f330])).
fof(f5273,plain,(
  cont(sK1,sK2) | ~spl14_8),
  inference(forward_demodulation,[],[f5249,f468])).
fof(f5334,plain,(
  $false | ~spl14_8),
  inference(unit_resulting_resolution,[],[f306,f5273,f371])).
fof(f5343,plain,(
  ~spl14_8),
  inference(avatar_contradiction_clause,[],[f5334])).
fof(f5352,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f4222,f5343])).
% SZS output end Proof for codi_down_sum_theorems_goal05.all
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 6396
% Time elapsed: 0.371 s
% ------------------------------
% ------------------------------
% Success in time 3.301 s
