Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 2999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal07.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 833
Passive clauses: 4450
Generated clauses: 14592
Final active clauses: 807
Final passive clauses: 3285
Input formulas: 44
Initial clauses: 95

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 182
Fw subsumption resolutions: 349

Simple tautologies: 62
Forward subsumptions: 9557

Binary resolution: 13755
Factoring: 41

Backtracking splits: 97
Backtracking splits refuted: 25
Backtracking splits refuted at zero level: 1

Memory used: 2046KB
Time elapsed: 0.404 s
------------------------------
% remaining time: 2995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal07.all
Refutation found. Thanks to Tanya!
2199. $false (8:0) [subsumption resolution 2188,165]
165. eqdim(sK2,sK0) (0:3) [cnf transformation 147]
147. cont(sK2,sK0) & cont(sK2,sK1) & eqdim(sK2,sK0) & eqdim(sK2,sK1) & ~cont(sK2,intersection(sK0,sK1))[skolemisation 88]
88. ? [X0,X1,X2] : (cont(X2,X0) & cont(X2,X1) & eqdim(X2,X0) & eqdim(X2,X1) & ~cont(X2,intersection(X0,X1)))[flattening 87]
87. ? [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,X0) & eqdim(X2,X1)) & ~cont(X2,intersection(X0,X1)))[ennf transformation 45]
45. ~! [X0,X1,X2] : (~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,X0) & eqdim(X2,X1)) | cont(X2,intersection(X0,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : (~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,X0) & eqdim(X2,X1)) | cont(X2,intersection(X0,X1)))[negated conjecture]
2188. ~eqdim(sK2,sK0) (8:3) [resolution 2182,216]
216. leq(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 109]
109. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 108]
108. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[ennf transformation 62]
62. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[flattening 43]
43. ! [X0] : ! [X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[input]
2182. ~leq(sK0,sK2) (7:3) [subsumption resolution 2172,1473]
1473. ~zex(intersection(sK0,sK1)) (5:4) [resolution 1444,223]
223. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 112]
112. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 66]
66. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[flattening 36]
36. ! [X0] : ! [X1] : (~zex(X0) | leq(X0,X1))[input]
1444. ~leq(intersection(sK0,sK1),sK2) (4:5) [subsumption resolution 1443,163]
163. cont(sK2,sK0) (0:3) [cnf transformation 147]
1443. ~leq(intersection(sK0,sK1),sK2) | ~cont(sK2,sK0) (4:8) [subsumption resolution 1423,164]
164. cont(sK2,sK1) (0:3) [cnf transformation 147]
1423. ~leq(intersection(sK0,sK1),sK2) | ~cont(sK2,sK1) | ~cont(sK2,sK0) (4:11) [resolution 1122,250]
250. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 143]
143. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 142]
142. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 86]
86. ! [X0,X1,X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[flattening 4]
4. ! [X0] : ! [X1] : ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[input]
1122. ~leq(sK2,intersection(sK0,sK1)) | ~leq(intersection(sK0,sK1),sK2) (3:10) [resolution 1118,217]
217. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 109]
1118. ~eqdim(sK2,intersection(sK0,sK1)) (2:5) [subsumption resolution 1117,163]
1117. ~eqdim(sK2,intersection(sK0,sK1)) | ~cont(sK2,sK0) (2:8) [subsumption resolution 1109,164]
1109. ~eqdim(sK2,intersection(sK0,sK1)) | ~cont(sK2,sK1) | ~cont(sK2,sK0) (2:11) [resolution 240,275]
275. ~p(sK2,intersection(sK0,sK1)) (1:5) [resolution 212,167]
167. ~cont(sK2,intersection(sK0,sK1)) (0:5) [cnf transformation 147]
212. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 107]
107. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 106]
106. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 61]
61. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 29]
29. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
240. p(X2,intersection(X0,X1)) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:16) [cnf transformation 131]
131. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 130]
130. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[ennf transformation 79]
79. ! [X0,X1,X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 5]
5. ! [X0] : ! [X1] : ! [X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[input]
2172. ~leq(sK0,sK2) | zex(intersection(sK0,sK1)) (7:7) [resolution 1657,222]
222. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 65]
65. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 3]
3. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
1657. ~cont(intersection(sK0,sK1),X11) | ~leq(X11,sK2) (6:8) [resolution 1474,1215]
1215. leq(X21,X22) | ~cont(X21,X22) (2:6) [resolution 1182,570]
570. lt(X0,X1) | leq(X1,X0) (0:6) [subsumption resolution 220,221]
221. leq(X1,X0) | leq(X0,X1) (0:6) [cnf transformation 64]
64. ! [X0,X1] : (leq(X0,X1) | leq(X1,X0))[flattening 33]
33. ! [X0] : ! [X1] : (leq(X0,X1) | leq(X1,X0))[input]
220. lt(X0,X1) | leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 111]
111. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 110]
110. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[ennf transformation 63]
63. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[flattening 40]
40. ! [X0] : ! [X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[input]
1182. ~lt(X5,X4) | ~cont(X4,X5) (1:6) [subsumption resolution 598,638]
638. ~lt(X0,X1) | ~eqdim(X1,X0) (1:6) [resolution 245,179]
179. ~lt(X0,X0) (0:3) [cnf transformation 94]
94. ! [X0] : ~lt(X0,X0)[ennf transformation 10]
10. ! [X0] : ~lt(X0,X0)[input]
245. lt(X2,X1) | ~lt(X2,X0) | ~eqdim(X0,X1) (0:9) [cnf transformation 133]
133. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[ennf transformation 81]
81. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[flattening 17]
17. ! [X0] : ! [X1] : ! [X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[input]
598. eqdim(X4,X5) | ~cont(X4,X5) | ~lt(X5,X4) (1:9) [resolution 229,228]
228. ~lt(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 116]
116. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[ennf transformation 70]
70. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[flattening 11]
11. ! [X0] : ! [X1] : (~lt(X0,X1) | ~lt(X1,X0))[input]
229. lt(X0,X1) | eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 117]
117. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[ennf transformation 71]
71. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 30]
30. ! [X0] : ! [X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[input]
1474. ~leq(intersection(sK0,sK1),X0) | ~leq(X0,sK2) (5:8) [resolution 1444,248]
248. leq(X0,X2) | ~leq(X1,X2) | ~leq(X0,X1) (0:9) [cnf transformation 139]
139. ! [X0,X1,X2] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2))[flattening 138]
138. ! [X0,X1,X2] : ((~leq(X0,X1) | ~leq(X1,X2)) | leq(X0,X2))[ennf transformation 84]
84. ! [X0,X1,X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[flattening 34]
34. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 291
Passive clauses: 934
Generated clauses: 2029
Final active clauses: 291
Final passive clauses: 612
Input formulas: 44
Initial clauses: 94

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 49
Fw subsumption resolutions: 123

Simple tautologies: 44
Equational tautologies: 3
Forward subsumptions: 900

Binary resolution: 1012
Factoring: 14
Forward superposition: 678
Backward superposition: 62

Unique components: 5

Memory used: 511KB
Time elapsed: 0.100 s
------------------------------
% Success in time 0.593 s
