Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal13.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 960
Passive clauses: 4405
Generated clauses: 12623
Final active clauses: 957
Final passive clauses: 3254
Input formulas: 44
Initial clauses: 93

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 63
Fw subsumption resolutions: 164

Simple tautologies: 34
Forward subsumptions: 8042

Binary resolution: 12164
Factoring: 43

Backtracking splits: 77
Backtracking splits refuted: 8

Memory used: 2046KB
Time elapsed: 0.405 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal13.all
Refutation found. Thanks to Tanya!
1081. $false (5:0) [subsumption resolution 1076,162]
162. ~zex(sK0) (0:2) [cnf transformation 146]
146. ~zex(sK0) & ~zex(intersection(sK0,sK0)) & ~p(intersection(sK0,sK0),sK0)[skolemisation 87]
87. ? [X0] : (~zex(X0) & ~zex(intersection(X0,X0)) & ~p(intersection(X0,X0),X0))[flattening 86]
86. ? [X0] : ((~zex(X0) & ~zex(intersection(X0,X0))) & ~p(intersection(X0,X0),X0))[ennf transformation 1]
1. ~! [X0] : (~(~zex(X0) & ~zex(intersection(X0,X0))) | p(intersection(X0,X0),X0))[negated conjecture]
1076. zex(sK0) (5:2) [resolution 1062,167]
167. cont(X0,X0) | zex(X0) (0:5) [cnf transformation 88]
88. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[ennf transformation 22]
22. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[input]
1062. ~cont(sK0,sK0) (4:3) [subsumption resolution 1055,163]
163. ~zex(intersection(sK0,sK0)) (0:4) [cnf transformation 146]
1055. ~cont(sK0,sK0) | zex(intersection(sK0,sK0)) (4:7) [resolution 1036,219]
219. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 64]
64. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 3]
3. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
1036. ~cont(intersection(sK0,sK0),sK0) | ~cont(sK0,sK0) (3:8) [duplicate literal removal 1020]
1020. ~cont(intersection(sK0,sK0),sK0) | ~cont(sK0,sK0) | ~cont(sK0,sK0) (3:11) [resolution 1013,247]
247. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 142]
142. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 141]
141. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 85]
85. ! [X0,X1,X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[flattening 4]
4. ! [X0] : ! [X1] : ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[input]
1013. ~leq(sK0,intersection(sK0,sK0)) | ~cont(intersection(sK0,sK0),sK0) (2:10) [duplicate literal removal 1004]
1004. ~cont(intersection(sK0,sK0),sK0) | ~leq(sK0,intersection(sK0,sK0)) | ~cont(intersection(sK0,sK0),sK0) (2:15) [resolution 565,500]
500. ~eqdim(intersection(sK0,sK0),sK0) | ~cont(intersection(sK0,sK0),sK0) (1:10) [resolution 211,164]
164. ~p(intersection(sK0,sK0),sK0) (0:5) [cnf transformation 146]
211. p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 106]
106. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 105]
105. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 60]
60. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 29]
29. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
565. eqdim(X8,X9) | ~cont(X8,X9) | ~leq(X9,X8) (1:9) [resolution 226,216]
216. ~lt(X0,X1) | ~leq(X1,X0) (0:6) [cnf transformation 110]
110. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 109]
109. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[ennf transformation 62]
62. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[flattening 40]
40. ! [X0] : ! [X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[input]
226. lt(X0,X1) | eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 116]
116. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[ennf transformation 70]
70. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 30]
30. ! [X0] : ! [X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 189
Passive clauses: 455
Generated clauses: 912
Final active clauses: 189
Final passive clauses: 248
Input formulas: 44
Initial clauses: 92

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 34
Fw subsumption resolutions: 81

Simple tautologies: 39
Equational tautologies: 1
Forward subsumptions: 315

Binary resolution: 468
Factoring: 5
Forward superposition: 196
Backward superposition: 36

Unique components: 3

Memory used: 255KB
Time elapsed: 0.053 s
------------------------------
% Success in time 0.534 s
