Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 11999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_down_theoremsEP-E3.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1346
Passive clauses: 8063
Generated clauses: 15719
Final active clauses: 1062
Final passive clauses: 5015
Input formulas: 70
Initial clauses: 131

Selected by SInE selection: 70
SInE iterations: 1

Duplicate literals: 88
Fw subsumption resolutions: 750

Simple tautologies: 65
Forward subsumptions: 6620

Binary resolution: 14071
Factoring: 37

Backtracking splits: 320
Backtracking splits refuted: 69
Backtracking splits refuted at zero level: 13

Memory used: 2942KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 11995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_down_theoremsEP-E3.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 2118
Passive clauses: 15558
Generated clauses: 35205
Final active clauses: 2118
Final passive clauses: 12865
Input formulas: 70
Initial clauses: 130

Selected by SInE selection: 70
SInE iterations: 1

Duplicate literals: 711
Fw subsumption resolutions: 2529

Simple tautologies: 463
Equational tautologies: 8
Forward subsumptions: 16617

Binary resolution: 24490
Factoring: 131
Forward superposition: 7263
Backward superposition: 70

Unique components: 4

Memory used: 5756KB
Time elapsed: 1.402 s
------------------------------
% remaining time: 11981 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on theorems\conversions\codi_down_theoremsEP-E3.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 5434
Passive clauses: 21722
Generated clauses: 93036
Final active clauses: 2464
Final passive clauses: 7231
Input formulas: 70
Initial clauses: 128

Duplicate literals: 243
Trivial inequalities: 154
Bw subsumption resolutions: 2663
Fw demodulations: 15320
Condensations: 33

Simple tautologies: 360
Equational tautologies: 365
Forward subsumptions: 49567

Binary resolution: 42584
Factoring: 5
Forward superposition: 913
Backward superposition: 16705
Self superposition: 4

Backtracking splits: 2857
Backtracking splits refuted: 1496
Backtracking splits refuted at zero level: 64

Memory used: 17014KB
Time elapsed: 3.307 s
------------------------------
% remaining time: 11948 next slice time: 95
dis+2_4_bs=off:ep=on:n=on:nicw=on:nwc=1.5:ptb=off:sac=on:sio=off:spl=backtracking:ssec=off_95 on theorems\conversions\codi_down_theoremsEP-E3.all
Refutation found. Thanks to Tanya!
24149. $false (22:0) [subsumption resolution 24148,1833]
1833. p(sK1,sK1) (13:3) [subsumption resolution 1832,258]
258. ~zex(sK1) (0:2) [cnf transformation 236]
236. ~zex(sK0) & ~zex(sK1) & ~cont(sK1,sK0) & ! [X2] : (~p(X2,sK1) | ~lt(intersection(X2,sK0),X2))[skolemisation 140]
140. ? [X0,X1] : (~zex(X0) & ~zex(X1) & ~cont(X1,X0) & ! [X2] : (~p(X2,X1) | ~lt(intersection(X2,X0),X2)))[flattening 139]
139. ? [X0,X1] : ((~zex(X0) & ~zex(X1) & ~cont(X1,X0)) & ! [X2] : (~p(X2,X1) | ~lt(intersection(X2,X0),X2)))[ennf transformation 71]
71. ~! [X0,X1] : (~(~zex(X0) & ~zex(X1) & ~cont(X1,X0)) | ? [X2] : (p(X2,X1) & lt(intersection(X2,X0),X2)))[flattening 1]
1. ~! [X0] : ! [X1] : (~(~zex(X0) & ~zex(X1) & ~cont(X1,X0)) | ? [X2] : (p(X2,X1) & lt(intersection(X2,X0),X2)))[negated conjecture]
1832. zex(sK1) | p(sK1,sK1) (13:5) [forward demodulation 1816,1752]
1752. difference(sK1,sK18) = sK1 (12:5) [resolution 1649,258]
1649. zex(X4) | difference(X4,sK18) = X4 (11:7) [resolution 1645,335]
335. ~lt(X1,X0) | difference(X0,X1) = X0 (0:8) [cnf transformation 173]
173. ! [X0,X1] : (~lt(X1,X0) | difference(X0,X1) = X0)[ennf transformation 103]
103. ! [X0,X1] : (~lt(X1,X0) | difference(X0,X1) = X0)[flattening 8]
8. ! [X0] : ! [X1] : (~lt(X1,X0) | difference(X0,X1) = X0)[input]
1645. lt(sK18,X3) | zex(X3) (10:5) [backtracking split refutation 1203,1637,1644]
1644. $false {6} (11:0) [resolution 1637,382]
382. zex(sK18) (0:2) [cnf transformation 255]
255. zex(sK18)[skolemisation 70]
70. ? [X0] : zex(X0)[input]
1637. ~zex(X4) {6} (10:2) [backtracking splitting component]
1203. lt(sK18,X3) | zex(X3) | ~zex(X4) (10:7) [resolution 1144,344]
344. lt(X0,X1) | zex(X1) | ~zex(X0) (0:7) [cnf transformation 186]
186. ! [X0,X1] : (~zex(X0) | zex(X1) | lt(X0,X1))[flattening 185]
185. ! [X0,X1] : ((~zex(X0) | zex(X1)) | lt(X0,X1))[ennf transformation 112]
112. ! [X0,X1] : (~(zex(X0) & ~zex(X1)) | lt(X0,X1))[flattening 47]
47. ! [X0] : ! [X1] : (~(zex(X0) & ~zex(X1)) | lt(X0,X1))[input]
1144. ~lt(X13,X12) | lt(sK18,X12) (9:6) [resolution 1090,312]
312. ~leq(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 163]
163. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 162]
162. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[ennf transformation 89]
89. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[flattening 66]
66. ! [X0] : ! [X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[input]
1090. leq(X17,X18) | lt(sK18,X17) (8:6) [resolution 1047,606]
606. leq(X1,X0) | lt(X0,X1) (0:6) [subsumption resolution 313,323]
323. leq(X1,X0) | leq(X0,X1) (0:6) [cnf transformation 93]
93. ! [X0,X1] : (leq(X0,X1) | leq(X1,X0))[flattening 59]
59. ! [X0] : ! [X1] : (leq(X0,X1) | leq(X1,X0))[input]
313. lt(X0,X1) | leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 163]
1047. ~leq(X2,sK18) | leq(X2,X3) (7:6) [resolution 1045,370]
370. ~leq(X1,X2) | leq(X0,X2) | ~leq(X0,X1) (0:9) [cnf transformation 212]
212. ! [X0,X1,X2] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2))[flattening 211]
211. ! [X0,X1,X2] : ((~leq(X0,X1) | ~leq(X1,X2)) | leq(X0,X2))[ennf transformation 128]
128. ! [X0,X1,X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[flattening 60]
60. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[input]
1045. leq(sK18,X38) (6:3) [backtracking split refutation 1033,1037,1044]
1044. $false {3} (7:0) [resolution 1037,382]
1037. ~zex(X39) {3} (6:2) [backtracking splitting component]
1033. leq(sK18,X38) | ~zex(X39) (6:5) [resolution 973,737]
737. ~lt(X12,sK18) (5:3) [superposition 681,725]
725. difference(X0,X0) = sK18 (4:5) [subsumption resolution 714,464]
464. ~zex(X2) | difference(X2,X3) = sK18 (3:7) [resolution 450,318]
318. zex(difference(X0,X1)) | ~zex(X0) (0:6) [cnf transformation 166]
166. ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))[ennf transformation 91]
91. ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & (~(zex(X0) | cont(X0,X1)) | zex(difference(X0,X1))))[flattening 12]
12. ! [X0] : ! [X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & (~(zex(X0) | cont(X0,X1)) | zex(difference(X0,X1))))[input]
450. ~zex(X0) | sK18 = X0 (2:5) [resolution 433,382]
433. ~zex(X1) | ~zex(X0) | X0 = X1 (1:7) [resolution 387,396]
396. ~sQ1_eqProxy(X0,X1) | X0 = X1 (0:6) [equality proxy axiom]
387. sQ1_eqProxy(X0,X1) | ~zex(X1) | ~zex(X0) (0:7) [equality proxy replacement 341]
341. X0 = X1 | ~zex(X1) | ~zex(X0) (0:7) [cnf transformation 180]
180. ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)[flattening 179]
179. ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)[ennf transformation 109]
109. ! [X0,X1] : (~(zex(X0) & zex(X1)) | X0 = X1)[flattening 61]
61. ! [X0] : ! [X1] : (~(zex(X0) & zex(X1)) | X0 = X1)[input]
714. difference(X0,X0) = sK18 | zex(X0) (4:7) [resolution 463,263]
263. cont(X0,X0) | zex(X0) (0:5) [cnf transformation 141]
141. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[ennf transformation 48]
48. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[input]
463. ~cont(X0,X1) | difference(X0,X1) = sK18 (3:8) [resolution 450,319]
319. zex(difference(X0,X1)) | ~cont(X0,X1) (0:7) [cnf transformation 166]
681. ~lt(X4,difference(X4,X5)) (1:5) [subsumption resolution 677,406]
406. ~lt(X7,X8) | ~zex(X8) (1:5) [resolution 312,328]
328. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 168]
168. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 98]
98. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[flattening 62]
62. ! [X0] : ! [X1] : (~zex(X0) | leq(X0,X1))[input]
677. zex(difference(X4,X5)) | ~lt(X4,difference(X4,X5)) (1:9) [resolution 324,337]
337. ~eqdim(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 175]
175. ! [X0,X1] : (~lt(X0,X1) | ~eqdim(X0,X1))[ennf transformation 105]
105. ! [X0,X1] : (~lt(X0,X1) | ~eqdim(X0,X1))[flattening 42]
42. ! [X0] : ! [X1] : (~lt(X0,X1) | ~eqdim(X0,X1))[input]
324. eqdim(X0,difference(X0,X1)) | zex(difference(X0,X1)) (0:9) [cnf transformation 94]
94. ! [X0,X1] : (zex(difference(X0,X1)) | eqdim(X0,difference(X0,X1)))[flattening 7]
7. ! [X0] : ! [X1] : (zex(difference(X0,X1)) | eqdim(X0,difference(X0,X1)))[input]
973. leq(X26,X27) | lt(X28,X26) | ~zex(X28) (2:8) [resolution 898,606]
898. ~leq(X23,X25) | leq(X23,X24) | ~zex(X25) (1:8) [resolution 370,328]
1816. p(sK1,sK1) | zex(difference(sK1,sK18)) (13:7) [superposition 326,1752]
326. p(difference(X0,X1),X0) | zex(difference(X0,X1)) (0:9) [cnf transformation 96]
96. ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))[flattening 13]
13. ! [X0] : ! [X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))[input]
24148. ~p(sK1,sK1) (22:3) [subsumption resolution 24144,17373]
17373. cont(intersection(sK1,sK0),sK1) (19:5) [resolution 17271,362]
362. ~p(X2,intersection(X0,X1)) | cont(X2,X0) (0:8) [cnf transformation 204]
204. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 203]
203. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[ennf transformation 122]
122. ! [X0,X1,X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 30]
30. ! [X0] : ! [X1] : ! [X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[input]
17271. p(intersection(sK1,sK0),intersection(sK1,sK0)) (18:7) [subsumption resolution 17270,13425]
13425. ~zex(intersection(sK1,sK0)) (16:4) [subsumption resolution 13409,1833]
13409. ~zex(intersection(sK1,sK0)) | ~p(sK1,sK1) (16:7) [resolution 13332,260]
260. ~lt(intersection(X2,sK0),X2) | ~p(X2,sK1) (0:8) [cnf transformation 236]
13332. lt(X28,sK1) | ~zex(X28) (15:5) [resolution 4581,1838]
1838. cont(sK1,sK1) (14:3) [resolution 1833,314]
314. ~p(X0,X1) | cont(X0,X1) (0:6) [cnf transformation 165]
165. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 164]
164. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 90]
90. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 55]
55. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
4581. ~cont(X12,X11) | lt(X10,X11) | ~zex(X10) (2:8) [subsumption resolution 4568,330]
330. ~cont(X0,X1) | ~zex(X0) (0:5) [cnf transformation 169]
169. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 99]
99. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[flattening 51]
51. ! [X0] : ! [X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[input]
4568. lt(X10,X11) | ~cont(X12,X11) | zex(X12) | ~zex(X10) (2:10) [resolution 4540,344]
4540. ~lt(X5,X6) | lt(X5,X7) | ~cont(X6,X7) (1:9) [subsumption resolution 807,377]
377. ~lt(X1,X2) | lt(X0,X2) | ~lt(X0,X1) (0:9) [cnf transformation 226]
226. ! [X0,X1,X2] : (~lt(X0,X1) | ~lt(X1,X2) | lt(X0,X2))[flattening 225]
225. ! [X0,X1,X2] : ((~lt(X0,X1) | ~lt(X1,X2)) | lt(X0,X2))[ennf transformation 135]
135. ! [X0,X1,X2] : (~(lt(X0,X1) & lt(X1,X2)) | lt(X0,X2))[flattening 38]
38. ! [X0] : ! [X1] : ! [X2] : (~(lt(X0,X1) & lt(X1,X2)) | lt(X0,X2))[input]
807. ~lt(X5,X6) | lt(X5,X7) | lt(X6,X7) | ~cont(X6,X7) (1:12) [resolution 365,339]
339. eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 177]
177. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[ennf transformation 107]
107. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 56]
56. ! [X0] : ! [X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[input]
365. ~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1) (0:9) [cnf transformation 205]
205. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[ennf transformation 123]
123. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[flattening 43]
43. ! [X0] : ! [X1] : ! [X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[input]
17270. zex(intersection(sK1,sK0)) | p(intersection(sK1,sK0),intersection(sK1,sK0)) (18:11) [forward demodulation 17146,13430]
13430. intersection(sK1,sK0) = difference(intersection(sK1,sK0),sK18) (17:9) [resolution 13425,3099]
3099. zex(X0) | difference(X0,sK18) = X0 (2:7) [resolution 382,517]
517. ~zex(X1) | zex(X0) | difference(X0,X1) = X0 (1:9) [resolution 335,344]
17146. p(intersection(sK1,sK0),intersection(sK1,sK0)) | zex(difference(intersection(sK1,sK0),sK18)) (18:13) [superposition 326,13430]
24144. ~cont(intersection(sK1,sK0),sK1) | ~p(sK1,sK1) (22:8) [resolution 23789,2554]
2554. p(intersection(X48,sK0),X48) | ~cont(intersection(X48,sK0),X48) | ~p(X48,sK1) (2:13) [resolution 772,260]
772. lt(X0,X1) | ~cont(X0,X1) | p(X0,X1) (1:9) [duplicate literal removal 764]
764. lt(X0,X1) | ~cont(X0,X1) | p(X0,X1) | ~cont(X0,X1) (1:12) [resolution 339,316]
316. ~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 165]
23789. ~p(intersection(sK1,sK0),sK1) (21:5) [backward demodulation 23568,17503]
17503. ~p(intersection(sK1,sK0),difference(sK1,sK0)) (21:7) [resolution 17371,368]
368. ~po(X2,intersection(X0,X1)) | ~p(X2,difference(X0,X1)) (0:10) [cnf transformation 208]
208. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | ~po(X2,intersection(X0,X1)))[ennf transformation 126]
126. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | ~po(X2,intersection(X0,X1)))[flattening 19]
19. ! [X0] : ! [X1] : ! [X2] : (~p(X2,difference(X0,X1)) | ~po(X2,intersection(X0,X1)))[input]
17371. po(intersection(sK1,sK0),intersection(sK1,sK0)) (20:7) [resolution 17340,14114]
14114. ~lt(X70,X69) | po(X69,X69) (10:6) [subsumption resolution 14093,382]
14093. po(X69,X69) | ~zex(sK18) | ~lt(X70,X69) (10:8) [resolution 5390,1134]
1134. ~eqdim(sK18,X12) | ~lt(X13,X12) (9:6) [resolution 1085,312]
1085. leq(X7,X8) | ~eqdim(sK18,X7) (8:6) [resolution 1047,309]
309. leq(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 161]
161. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 160]
160. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[ennf transformation 88]
88. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[flattening 69]
69. ! [X0] : ! [X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[input]
5390. po(X1,X1) | eqdim(X0,X1) | ~zex(X0) (2:8) [duplicate literal removal 5386]
5386. ~zex(X0) | po(X1,X1) | eqdim(X0,X1) | eqdim(X0,X1) | ~zex(X0) (2:13) [resolution 1728,1690]
1690. p(sK12(X1,X0),X1) | eqdim(X0,X1) | ~zex(X0) (0:10) [subsumption resolution 1689,412]
412. eqdim(X1,X2) | ~p(X2,X1) (1:6) [resolution 334,315]
315. eqdim(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 165]
334. ~eqdim(X0,X1) | eqdim(X1,X0) (0:6) [cnf transformation 172]
172. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 102]
102. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[flattening 40]
40. ! [X0] : ! [X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[input]
1689. p(sK12(X1,X0),X1) | eqdim(X0,X1) | p(X1,X0) | ~zex(X0) (0:13) [subsumption resolution 345,627]
627. eqdim(X15,X16) | ~zex(X16) | ~zex(X15) (2:7) [resolution 588,328]
588. ~leq(X15,X16) | eqdim(X15,X16) | ~zex(X16) (1:8) [resolution 310,328]
310. ~leq(X1,X0) | eqdim(X0,X1) | ~leq(X0,X1) (0:9) [cnf transformation 161]
345. p(sK12(X1,X0),X1) | eqdim(X0,X1) | p(X1,X0) | zex(X1) | ~zex(X0) (0:15) [cnf transformation 249]
249. ! [X0,X1] : (~zex(X0) | zex(X1) | p(X1,X0) | eqdim(X0,X1) | (p(sK12(X1,X0),X1) & ~po(sK12(X1,X0),X0)))[skolemisation 188]
188. ! [X0,X1] : (~zex(X0) | zex(X1) | p(X1,X0) | eqdim(X0,X1) | ? [X2] : (p(X2,X1) & ~po(X2,X0)))[flattening 187]
187. ! [X0,X1] : ((~zex(X0) | zex(X1) | p(X1,X0) | eqdim(X0,X1)) | ? [X2] : (p(X2,X1) & ~po(X2,X0)))[ennf transformation 113]
113. ! [X0,X1] : (~(zex(X0) & ~zex(X1) & ~p(X1,X0) & ~eqdim(X0,X1)) | ? [X2] : (p(X2,X1) & ~po(X2,X0)))[flattening 5]
5. ! [X0] : ! [X1] : (~(zex(X0) & ~zex(X1) & ~p(X1,X0) & ~eqdim(X0,X1)) | ? [X2] : (p(X2,X1) & ~po(X2,X0)))[input]
1728. ~p(sK12(X1,X0),X2) | ~zex(X0) | po(X2,X1) | eqdim(X0,X1) (1:13) [resolution 1690,304]
304. ~p(X3,X1) | po(X0,X1) | ~p(X3,X0) (0:9) [cnf transformation 247]
247. ! [X0,X1] : ((~po(X0,X1) | (p(sK10(X1,X0),X0) & p(sK10(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[skolemisation 157]
157. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[ennf transformation 86]
86. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[flattening 85]
85. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[rectify 32]
32. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X2] : (p(X2,X0) & p(X2,X1)) | po(X0,X1)))[input]
17340. lt(sK18,intersection(sK1,sK0)) (19:5) [resolution 17269,5660]
5660. ~cont(X15,X16) | lt(sK18,X15) (13:6) [resolution 5287,4242]
4242. c(X14,X15) | ~cont(X14,X15) (3:6) [resolution 3392,555]
555. c(X0,X1) | ~cont(X1,X0) (1:6) [subsumption resolution 548,330]
548. c(X0,X1) | ~cont(X1,X0) | zex(X1) (1:8) [resolution 301,263]
301. ~cont(X3,X1) | c(X0,X1) | ~cont(X3,X0) (0:9) [cnf transformation 246]
246. ! [X0,X1] : ((~c(X0,X1) | (cont(sK9(X1,X0),X0) & cont(sK9(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 156]
156. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 84]
84. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 83]
83. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 54]
54. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
3392. ~c(X3,X2) | c(X2,X3) (2:6) [duplicate literal removal 3390]
3390. c(X2,X3) | ~c(X3,X2) | ~c(X3,X2) (2:9) [resolution 553,300]
300. cont(sK9(X1,X0),X1) | ~c(X0,X1) (0:8) [cnf transformation 246]
553. ~cont(sK9(X16,X15),X14) | c(X14,X15) | ~c(X15,X16) (1:11) [resolution 301,299]
299. cont(sK9(X1,X0),X0) | ~c(X0,X1) (0:8) [cnf transformation 246]
5287. ~c(X35,X36) | lt(sK18,X35) (12:6) [resolution 4583,299]
4583. ~cont(X35,X34) | lt(sK18,X34) (11:6) [subsumption resolution 4578,330]
4578. lt(sK18,X34) | ~cont(X35,X34) | zex(X35) (11:8) [resolution 4540,1645]
17269. cont(intersection(sK1,sK0),sK0) (18:5) [subsumption resolution 17268,13425]
17268. zex(intersection(sK1,sK0)) | cont(intersection(sK1,sK0),sK0) (18:9) [forward demodulation 17141,13430]
17141. cont(intersection(sK1,sK0),sK0) | zex(difference(intersection(sK1,sK0),sK18)) (18:11) [superposition 708,13430]
708. cont(difference(intersection(X9,X10),X11),X10) | zex(difference(intersection(X9,X10),X11)) (1:13) [resolution 326,363]
363. ~p(X2,intersection(X0,X1)) | cont(X2,X1) (0:8) [cnf transformation 204]
23568. difference(sK1,sK0) = sK1 (5:5) [resolution 23337,335]
23337. lt(sK0,sK1) (4:3) [resolution 23320,606]
23320. ~leq(sK1,sK0) (3:3) [backtracking split refutation 23275,23283,23319]
23319. $false {46} (5:0) [subsumption resolution 23318,259]
259. ~cont(sK1,sK0) (0:3) [cnf transformation 236]
23318. cont(sK1,sK0) {46} (5:3) [subsumption resolution 23305,258]
23305. zex(sK1) | cont(sK1,sK0) {46} (5:5) [resolution 23284,317]
317. ~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1) (0:9) [cnf transformation 166]
23284. zex(difference(sK1,sK0)) {46} (4:4) [resolution 23283,4526]
4526. po(difference(X4,X5),X4) | zex(difference(X4,X5)) (2:9) [subsumption resolution 4516,263]
4516. po(difference(X4,X5),X4) | zex(difference(X4,X5)) | ~cont(difference(X4,X5),difference(X4,X5)) (2:16) [resolution 704,638]
638. p(X0,X0) | ~cont(X0,X0) (1:6) [resolution 316,261]
261. eqdim(X0,X0) (0:3) [cnf transformation 39]
39. ! [X0] : eqdim(X0,X0)[input]
704. ~p(difference(X0,X1),X2) | po(X2,X0) | zex(difference(X0,X1)) (1:12) [resolution 326,304]
23283. ~po(difference(sK1,sK0),X3) {46} (3:5) [backtracking splitting component]
23275. ~leq(sK1,sK0) | ~po(difference(sK1,sK0),X3) (3:8) [duplicate literal removal 23269]
23269. ~po(difference(sK1,sK0),X3) | ~leq(sK1,sK0) | ~po(difference(sK1,sK0),X3) (3:13) [resolution 8227,528]
528. p(sK10(X3,difference(X4,X5)),X4) | ~po(difference(X4,X5),X3) (1:12) [resolution 367,302]
302. p(sK10(X1,X0),X0) | ~po(X0,X1) (0:8) [cnf transformation 247]
367. ~p(X2,difference(X0,X1)) | p(X2,X0) (0:8) [cnf transformation 207]
207. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))[ennf transformation 125]
125. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))[flattening 18]
18. ! [X0] : ! [X1] : ! [X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))[input]
8227. ~p(sK10(X1,difference(X0,sK0)),sK1) | ~po(difference(X0,sK0),X1) | ~leq(X0,sK0) (2:15) [resolution 1413,260]
1413. lt(intersection(sK10(X6,difference(X7,X8)),X8),sK10(X6,difference(X7,X8))) | ~leq(X7,X8) | ~po(difference(X7,X8),X6) (1:21) [resolution 373,302]
373. ~p(X2,difference(X0,X1)) | lt(intersection(X2,X1),X2) | ~leq(X0,X1) (0:13) [cnf transformation 218]
218. ! [X0,X1,X2] : (~leq(X0,X1) | ~p(X2,difference(X0,X1)) | lt(intersection(X2,X1),X2))[flattening 217]
217. ! [X0,X1,X2] : ((~leq(X0,X1) | ~p(X2,difference(X0,X1))) | lt(intersection(X2,X1),X2))[ennf transformation 131]
131. ! [X0,X1,X2] : (~(leq(X0,X1) & p(X2,difference(X0,X1))) | lt(intersection(X2,X1),X2))[flattening 25]
25. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & p(X2,difference(X0,X1))) | lt(intersection(X2,X1),X2))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1351
Passive clauses: 8908
Generated clauses: 24974
Final active clauses: 1076
Final passive clauses: 5894
Input formulas: 70
Initial clauses: 130

Duplicate literals: 402
Fw subsumption resolutions: 1372
Fw demodulations: 715
Bw demodulations: 143

Simple tautologies: 510
Equational tautologies: 12
Forward subsumptions: 13570

Binary resolution: 19359
Factoring: 14
Forward superposition: 1055
Backward superposition: 670

Backtracking splits: 47
Backtracking splits refuted: 23
Backtracking splits refuted at zero level: 12

Memory used: 3709KB
Time elapsed: 0.691 s
------------------------------
% Success in time 5.908 s
