% Starting sequential portfolio solving with schedule "casc"
% spawned child 13499 with time: 4 (total remaining time 2999)
% WARNING: value z3 for option sas not known
% dis+1010_4_afp=10000:afq=1.2:anc=none:irw=on:lma=on:nm=64:nwc=10:sas=z3:sac=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal21.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6908
% Time elapsed: 0.500 s
% spawned child 13500 with time: 10 (total remaining time 2994)
% ott+11_8_amm=off:anc=none:bsr=on:cond=on:irw=on:nm=2:nwc=1.1:ss=axioms:st=2.0:sac=on_8 on codi_down_sum_theorems_goal21.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8955
% Time elapsed: 1.100 s
% ------------------------------
% ------------------------------
% spawned child 13501 with time: 4 (total remaining time 2982)
% dis+1011_10_aac=none:add=large:afp=10000:afq=1.1:amm=off:anc=none:fsr=off:gs=on:gsem=off:irw=on:nm=6:newcnf=on:nwc=2.5:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal21.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 13176
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 13502 with time: 6 (total remaining time 2976)
% dis+11_28_av=off:fsr=off:irw=on:lcm=predicate:nm=2:newcnf=on:nwc=5:sp=occurrence:urr=on:updr=off_4 on codi_down_sum_theorems_goal21.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 4221
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% spawned child 13503 with time: 6 (total remaining time 2968)
% lrs+1011_5_add=large:afp=1000:afq=1.2:anc=none:fsr=off:irw=on:lma=on:nm=64:newcnf=on:nwc=1:stl=30:sd=3:ss=axioms:st=2.0:sos=on:sac=on:sp=reverse_arity:urr=on:updr=off_4 on codi_down_sum_theorems_goal21.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8955
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% spawned child 13504 with time: 4 (total remaining time 2960)
% lrs+1011_3:1_add=off:afr=on:afp=10000:afq=1.1:amm=off:bce=on:cond=on:ep=R:fsr=off:nm=16:nwc=1:stl=30:sos=all:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal21.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8827
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 13505 with time: 23 (total remaining time 2954)
% lrs+1011_3:2_aac=none:afp=10000:afq=1.2:amm=off:anc=all:cond=on:fde=none:gs=on:gsem=on:irw=on:lma=on:nm=32:newcnf=on:nwc=3:nicw=on:stl=30:sac=on:sp=reverse_arity:urr=ec_only:updr=off:uhcvi=on_21 on codi_down_sum_theorems_goal21.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 17142
% Time elapsed: 2.400 s
% ------------------------------
% ------------------------------
% spawned child 13506 with time: 4 (total remaining time 2932)
% lrs+10_6_aac=none:acc=model:add=large:afr=on:afp=10000:afq=1.2:amm=off:anc=none:bd=off:ccuc=small_ones:irw=on:lcm=reverse:nm=0:nwc=1:nicw=on:stl=30:sos=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal21.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6780
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 13507 with time: 7 (total remaining time 2927)
% dis+11_5_av=off:bd=off:bs=unit_only:bsr=on:cond=on:lcm=reverse:nm=0:nwc=1.2_5 on codi_down_sum_theorems_goal21.all
% Refutation found. Thanks to Tanya!
% SZS status Theorem for codi_down_sum_theorems_goal21.all
% SZS output start Proof for codi_down_sum_theorems_goal21.all
fof(f1,conjecture,(
  ! [X0] : ! [X1] : (! [X2] : (po(sum(X0,X1),X2) | ~(po(X1,X2) & po(X0,X2))) | ~(~zex(X1) & ~zex(X0) & eqdim(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal21.all.tptp',conjecture170)).
fof(f2,negated_conjecture,(
  ~! [X0] : ! [X1] : (! [X2] : (po(sum(X0,X1),X2) | ~(po(X1,X2) & po(X0,X2))) | ~(~zex(X1) & ~zex(X0) & eqdim(X0,X1)))),
  inference(negated_conjecture,[],[f1])).
fof(f3,axiom,(
  ! [X0] : ! [X1] : (! [X2] : (po(X2,sum(X0,X1)) | ~(po(X2,X1) | po(X2,X0))) | ~eqdim(X0,X1))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal21.all.tptp',conjecture150)).
fof(f71,axiom,(
  ! [X0] : ! [X1] : ((po(X0,X1) | ~? [X2] : (p(X2,X1) & p(X2,X0))) & (? [X2] : (p(X2,X1) & p(X2,X0)) | ~po(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal21.all.tptp',axiom88290)).
fof(f120,plain,(
  ~! [X0,X1] : (! [X2] : (po(sum(X0,X1),X2) | ~(po(X1,X2) & po(X0,X2))) | ~(~zex(X1) & ~zex(X0) & eqdim(X0,X1)))),
  inference(flattening,[],[f2])).
fof(f150,plain,(
  ! [X0] : ! [X1] : ((po(X0,X1) | ~? [X2] : (p(X2,X1) & p(X2,X0))) & (? [X3] : (p(X3,X1) & p(X3,X0)) | ~po(X0,X1)))),
  inference(rectify,[],[f71])).
fof(f151,plain,(
  ! [X0,X1] : ((po(X0,X1) | ~? [X2] : (p(X2,X1) & p(X2,X0))) & (? [X3] : (p(X3,X1) & p(X3,X0)) | ~po(X0,X1)))),
  inference(flattening,[],[f150])).
fof(f169,plain,(
  ! [X0,X1] : (! [X2] : (po(X2,sum(X0,X1)) | ~(po(X2,X1) | po(X2,X0))) | ~eqdim(X0,X1))),
  inference(flattening,[],[f3])).
fof(f240,plain,(
  ? [X0,X1] : (? [X2] : (~po(sum(X0,X1),X2) & (po(X1,X2) & po(X0,X2))) & (~zex(X1) & ~zex(X0) & eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f120])).
fof(f241,plain,(
  ? [X0,X1] : (? [X2] : (~po(sum(X0,X1),X2) & po(X1,X2) & po(X0,X2)) & ~zex(X1) & ~zex(X0) & eqdim(X0,X1))),
  inference(flattening,[],[f240])).
fof(f264,plain,(
  ! [X0,X1] : ((po(X0,X1) | ! [X2] : (~p(X2,X1) | ~p(X2,X0))) & (? [X3] : (p(X3,X1) & p(X3,X0)) | ~po(X0,X1)))),
  inference(ennf_transformation,[],[f151])).
fof(f273,plain,(
  ! [X0,X1] : (! [X2] : (po(X2,sum(X0,X1)) | (~po(X2,X1) & ~po(X2,X0))) | ~eqdim(X0,X1))),
  inference(ennf_transformation,[],[f169])).
fof(f356,plain,(
  ? [X0,X1] : (? [X2] : (~po(sum(X0,X1),X2) & po(X1,X2) & po(X0,X2)) & ~zex(X1) & ~zex(X0) & eqdim(X0,X1)) => (? [X2] : (~po(sum(sK0,sK1),X2) & po(sK1,X2) & po(sK0,X2)) & ~zex(sK1) & ~zex(sK0) & eqdim(sK0,sK1))),
  introduced(choice_axiom,[])).
fof(f357,plain,(
  ( ! [X0,X1] : (? [X2] : (~po(sum(X0,X1),X2) & po(X1,X2) & po(X0,X2)) => (~po(sum(X0,X1),sK2) & po(X1,sK2) & po(X0,sK2))) )),
  introduced(choice_axiom,[])).
fof(f358,plain,(
  (~po(sum(sK0,sK1),sK2) & po(sK1,sK2) & po(sK0,sK2)) & ~zex(sK1) & ~zex(sK0) & eqdim(sK0,sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f241,f357,f356])).
fof(f383,plain,(
  ! [X1,X0] : (? [X3] : (p(X3,X1) & p(X3,X0)) => (p(sK13(X0,X1),X1) & p(sK13(X0,X1),X0)))),
  introduced(choice_axiom,[])).
fof(f384,plain,(
  ! [X0,X1] : ((po(X0,X1) | ! [X2] : (~p(X2,X1) | ~p(X2,X0))) & ((p(sK13(X0,X1),X1) & p(sK13(X0,X1),X0)) | ~po(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f264,f383])).
fof(f409,plain,(
  eqdim(sK0,sK1)),
  inference(cnf_transformation,[],[f358])).
fof(f413,plain,(
  po(sK1,sK2)),
  inference(cnf_transformation,[],[f358])).
fof(f414,plain,(
  ~po(sum(sK0,sK1),sK2)),
  inference(cnf_transformation,[],[f358])).
fof(f478,plain,(
  ( ! [X0,X1] : (p(sK13(X0,X1),X0) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f384])).
fof(f479,plain,(
  ( ! [X0,X1] : (p(sK13(X0,X1),X1) | ~po(X0,X1)) )),
  inference(cnf_transformation,[],[f384])).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (~p(X2,X0) | ~p(X2,X1) | po(X0,X1)) )),
  inference(cnf_transformation,[],[f384])).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (po(X2,sum(X0,X1)) | ~po(X2,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f273])).
fof(f871,plain,(
  ( ! [X4,X5,X3] : (~p(sK13(X3,X4),X5) | po(X4,X5) | ~po(X3,X4)) )),
  inference(resolution,[],[f480,f479])).
fof(f7362,plain,(
  ( ! [X0,X1] : (po(X0,X1) | ~po(X1,X0) | ~po(X1,X0)) )),
  inference(resolution,[],[f871,f478])).
fof(f7385,plain,(
  ( ! [X0,X1] : (po(X0,X1) | ~po(X1,X0)) )),
  inference(duplicate_literal_removal,[],[f7362])).
fof(f7828,plain,(
  ~po(sK2,sum(sK0,sK1))),
  inference(resolution,[],[f7385,f414])).
fof(f7844,plain,(
  ~po(sK2,sK1) | ~eqdim(sK0,sK1)),
  inference(resolution,[],[f7828,f512])).
fof(f7847,plain,(
  ~po(sK2,sK1)),
  inference(subsumption_resolution,[],[f7844,f409])).
fof(f7850,plain,(
  ~po(sK1,sK2)),
  inference(resolution,[],[f7847,f7385])).
fof(f7852,plain,(
  $false),
  inference(subsumption_resolution,[],[f7850,f413])).
% SZS output end Proof for codi_down_sum_theorems_goal21.all
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 3454
% Time elapsed: 0.794 s
% ------------------------------
% ------------------------------
% Success in time 8.076 s
