Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1308
Passive clauses: 8171
Generated clauses: 15452
Final active clauses: 1216
Final passive clauses: 6079
Input formulas: 61
Initial clauses: 98

Selected by SInE selection: 61
SInE iterations: 1

Duplicate literals: 116
Fw subsumption resolutions: 504

Simple tautologies: 48
Forward subsumptions: 6340

Binary resolution: 14159
Factoring: 48

Backtracking splits: 298
Backtracking splits refuted: 63
Backtracking splits refuted at zero level: 19

Memory used: 2814KB
Time elapsed: 0.402 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1489
Passive clauses: 11814
Generated clauses: 32964
Final active clauses: 1489
Final passive clauses: 9938
Input formulas: 61
Initial clauses: 97

Selected by SInE selection: 61
SInE iterations: 1

Duplicate literals: 730
Fw subsumption resolutions: 1518

Simple tautologies: 536
Equational tautologies: 17
Forward subsumptions: 18749

Binary resolution: 25196
Factoring: 120
Forward superposition: 5208
Backward superposition: 118

Unique components: 8

Memory used: 4477KB
Time elapsed: 1.403 s
------------------------------
% remaining time: 5981 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 5360
Passive clauses: 16024
Generated clauses: 109564
Final active clauses: 1949
Final passive clauses: 1973
Input formulas: 61
Initial clauses: 95

Duplicate literals: 251
Trivial inequalities: 371
Bw subsumption resolutions: 8883
Fw demodulations: 17633
Condensations: 189

Simple tautologies: 924
Equational tautologies: 463
Forward subsumptions: 69023

Binary resolution: 43215
Factoring: 3
Forward superposition: 11282
Backward superposition: 19118
Equality resolution: 26

Backtracking splits: 728
Backtracking splits refuted: 265
Backtracking splits refuted at zero level: 16

Memory used: 15735KB
Time elapsed: 3.304 s
------------------------------
% remaining time: 5948 next slice time: 95
dis+2_4_bs=off:ep=on:n=on:nicw=on:nwc=1.5:ptb=off:sac=on:sio=off:spl=backtracking:ssec=off_95 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 4534
Passive clauses: 87811
Generated clauses: 286101
Final active clauses: 4192
Final passive clauses: 79840
Input formulas: 61
Initial clauses: 97

Duplicate literals: 2332
Fw subsumption resolutions: 16121
Fw demodulations: 13397
Bw demodulations: 185

Simple tautologies: 1204
Equational tautologies: 58
Forward subsumptions: 167935

Binary resolution: 230944
Factoring: 67
Forward superposition: 17778
Backward superposition: 5004

Backtracking splits: 52
Backtracking splits refuted: 18
Backtracking splits refuted at zero level: 11

Memory used: 39658KB
Time elapsed: 9.626 s
------------------------------
% remaining time: 5851 next slice time: 4
lrs-1003_12_bms=on:bs=unit_only:bsr=on:cond=fast:drc=off:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=10:sfv=off:sgo=on:sio=off:sp=occurrence:spo=on:ssec=off:stl=81:updr=off_4 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1074
Passive clauses: 3096
Generated clauses: 9503
Final active clauses: 1010
Final passive clauses: 1804
Input formulas: 61
Initial clauses: 97
Reactivated clauses: 7

Duplicate literals: 18
Fw subsumption resolutions: 606
Bw subsumption resolutions: 10
Fw demodulations: 525
Bw demodulations: 136
Condensations: 79

Simple tautologies: 231
Equational tautologies: 1
Forward subsumptions: 4913
Backward subsumptions: 186
Subsumed by BDD marking: 154

Binary resolution: 5471
Factoring: 71
Forward superposition: 1083
Backward superposition: 1544

Split clauses: 8
Split components: 16
Unique components: 2672
Introduced splitting names: 5
BDD propositional clauses: 4

Memory used: 2430KB
Time elapsed: 0.502 s
------------------------------
% remaining time: 5846 next slice time: 22
lrs+1010_50_bs=off:drc=off:flr=on:lcm=predicate:n=on:nwc=1.2:sfv=off:sgo=on:stl=81_22 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1700
Passive clauses: 4069
Generated clauses: 41584
Final active clauses: 1637
Final passive clauses: 2306
Input formulas: 61
Initial clauses: 95

Duplicate literals: 1274
Fw subsumption resolutions: 2502
Fw demodulations: 54
Bw demodulations: 124
Fw literal rewrites: 4505

Simple tautologies: 431
Equational tautologies: 17
Forward subsumptions: 28733
Bw demodulations to eq. taut.: 2

Binary resolution: 29710
Factoring: 24
Forward superposition: 1316
Backward superposition: 1987
Self superposition: 2
Equality resolution: 2

Unique components: 8

Memory used: 4221KB
Time elapsed: 2.302 s
------------------------------
% remaining time: 5822 next slice time: 41
dis-1010_5_bd=off:bs=off:cond=fast:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=1.3:ptb=off:sac=on:sgo=on:sio=off:sos=on:spl=backtracking:ssec=off:updr=off_41 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 3278
Passive clauses: 15191
Generated clauses: 75340
Final active clauses: 3202
Final passive clauses: 10227
Input formulas: 61
Initial clauses: 97

Duplicate literals: 1499
Fw subsumption resolutions: 3383
Condensations: 92

Simple tautologies: 620
Forward subsumptions: 55305

Binary resolution: 68462
Factoring: 595
Forward superposition: 1200

Backtracking splits: 56
Backtracking splits refuted: 34
Backtracking splits refuted at zero level: 5

Memory used: 14072KB
Time elapsed: 4.203 s
------------------------------
% remaining time: 5780 next slice time: 237
lrs+1_2:1_bs=off:ep=on:gsp=input_only:lcm=predicate:n=on:nicw=on:nwc=1.1:ptb=off:sac=on:sgo=on:sio=off:sos=on:sp=occurrence:spl=backtracking:spo=on:ssec=off:stl=81:updr=off_237 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 8294
Passive clauses: 97147
Generated clauses: 270013
Final active clauses: 1773
Final passive clauses: 20952
Input formulas: 61
Initial clauses: 97
Discarded non-redundant clauses: 27947

Duplicate literals: 2898
Fw subsumption resolutions: 10616
Fw demodulations: 45481
Bw demodulations: 45998

Simple tautologies: 1781
Equational tautologies: 128
Forward subsumptions: 110819
Bw demodulations to eq. taut.: 58

Binary resolution: 121117
Factoring: 66
Forward superposition: 5412
Backward superposition: 3818

Backtracking splits: 268
Backtracking splits refuted: 37
Backtracking splits refuted at zero level: 32

Memory used: 47461KB
Time elapsed: 23.829 s
------------------------------
% remaining time: 5541 next slice time: 311
lrs+11_40_bd=off:bs=unit_only:bsr=unit_only:drc=off:flr=on:lcm=predicate:n=on:ptb=off:sac=on:sfv=off:spl=backtracking:ssec=off:stl=81:updr=off_311 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 8199
Passive clauses: 42583
Generated clauses: 675148
Final active clauses: 7589
Final passive clauses: 16812
Input formulas: 61
Initial clauses: 95
Discarded non-redundant clauses: 104033

Duplicate literals: 1278
Trivial inequalities: 30
Fw subsumption resolutions: 22497
Bw subsumption resolutions: 109
Fw demodulations: 154499
Fw literal rewrites: 50980

Simple tautologies: 3185
Equational tautologies: 2268
Forward subsumptions: 394917
Backward subsumptions: 2551

Binary resolution: 348641
Factoring: 24
Forward superposition: 56370
Backward superposition: 40699
Self superposition: 6
Equality factoring: 2
Equality resolution: 1

Memory used: 63453KB
Time elapsed: 31.262 s
------------------------------
% remaining time: 5228 next slice time: 280
lrs-1010_10_bd=off:bs=unit_only:cond=on:flr=on:gsp=input_only:n=on:nicw=on:ptb=off:sac=on:sfv=off:sio=off:sp=occurrence:spl=backtracking:spo=on:ssec=off:stl=81_280 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 6243
Passive clauses: 45977
Generated clauses: 340222
Final active clauses: 4108
Final passive clauses: 12512
Input formulas: 61
Initial clauses: 95
Discarded non-redundant clauses: 22009

Duplicate literals: 3943
Fw subsumption resolutions: 17817
Fw demodulations: 820
Fw literal rewrites: 20078
Condensations: 10604

Simple tautologies: 4646
Equational tautologies: 83
Forward subsumptions: 235370
Backward subsumptions: 14202

Binary resolution: 205046
Factoring: 284
Forward superposition: 36863
Backward superposition: 34105
Self superposition: 22
Equality resolution: 2

Backtracking splits: 3
Backtracking splits refuted: 3
Backtracking splits refuted at zero level: 3

Memory used: 38250KB
Time elapsed: 28.121 s
------------------------------
% remaining time: 4947 next slice time: 514
lrs+10_8_bs=off:fde=none:lcm=predicate:n=on:nwc=1.7:ptb=off:sac=on:sgo=on:sio=off:sp=occurrence:spl=backtracking:ssec=off:stl=81_514 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 27178
Passive clauses: 211079
Generated clauses: 1002876
Final active clauses: 5933
Final passive clauses: 16614
Input formulas: 61
Initial clauses: 95
Discarded non-redundant clauses: 49445

Duplicate literals: 5073
Fw subsumption resolutions: 109446
Fw demodulations: 77154
Bw demodulations: 156631

Simple tautologies: 4538
Equational tautologies: 1182
Forward subsumptions: 555019
Bw demodulations to eq. taut.: 872

Binary resolution: 404508
Factoring: 112
Forward superposition: 68719
Backward superposition: 29677
Equality resolution: 7

Backtracking splits: 77
Backtracking splits refuted: 28
Backtracking splits refuted at zero level: 12

Memory used: 95819KB
Time elapsed: 51.523 s
------------------------------
% remaining time: 4431 next slice time: 11
lrs-1010_7_bms=on:bs=off:bsr=on:cond=on:drc=off:ecs=on:ep=on:fde=none:flr=on:fsr=off:lcm=predicate:n=on:nicw=on:nwc=1.3:sac=on:sgo=on:sio=off:ssec=off:stl=81_11 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1525
Passive clauses: 5584
Generated clauses: 26681
Final active clauses: 500
Final passive clauses: 2959
Input formulas: 61
Initial clauses: 97
Reactivated clauses: 22

Duplicate literals: 316
Bw subsumption resolutions: 4008
Fw literal rewrites: 3323
Condensations: 81

Simple tautologies: 295
Equational tautologies: 7
Forward subsumptions: 16034
Subsumed empty clauses: 1011
Empty clause subsumptions: 1011
Subsumed by BDD marking: 40

Binary resolution: 16247
Factoring: 51
Forward superposition: 5406
Backward superposition: 119

Split clauses: 34
Split components: 71
Unique components: 141
Introduced splitting names: 17
BDD propositional clauses: 1024

Memory used: 5373KB
Time elapsed: 1.202 s
------------------------------
% remaining time: 4419 next slice time: 264
dis+3_5:4_bd=off:bs=off:drc=off:ep=RST:lcm=predicate:n=on:nicw=on:nwc=3:ptb=off:sac=on:sfv=off:sgo=on:sio=off:sp=occurrence:spl=backtracking:ssec=off:sswn=on:updr=off_264 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 6385
Passive clauses: 252141
Generated clauses: 416308
Final active clauses: 6298
Final passive clauses: 240607
Input formulas: 61
Initial clauses: 98

Duplicate literals: 1482
Fw subsumption resolutions: 27478

Simple tautologies: 618
Forward subsumptions: 137954

Binary resolution: 387021
Factoring: 163

Backtracking splits: 62
Backtracking splits refuted: 33
Backtracking splits refuted at zero level: 15

Memory used: 90190KB
Time elapsed: 26.540 s
------------------------------
% remaining time: 4153 next slice time: 269
lrs+1002_10_bd=off:bs=off:bsr=unit_only:ecs=on:gsp=input_only:lcm=reverse:n=on:nicw=on:nwc=1.2:sd=1:sfv=off:sio=off:sos=on:spl=off:ss=included:ssec=off:st=1.2:stl=40.5_269 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 8779
Passive clauses: 40556
Generated clauses: 242221
Final active clauses: 8779
Final passive clauses: 8442
Input formulas: 61
Initial clauses: 95
Discarded non-redundant clauses: 320363

Selected by SInE selection: 61
SInE iterations: 1

Duplicate literals: 29487
Fw subsumption resolutions: 7617

Simple tautologies: 544
Equational tautologies: 89
Forward subsumptions: 108892

Binary resolution: 98038
Factoring: 1
Forward superposition: 55685
Backward superposition: 51602
Self superposition: 43

Memory used: 80339KB
Time elapsed: 27.018 s
------------------------------
% remaining time: 3883 next slice time: 3883
User error: Backtracking splitting cannot be used unless all BDD related options are disabled
% remaining time: 3883 next slice time: 3883
dis+4_8_n=on_3883 on theorems\conversions\codi_down_theoremsT6-T10_goal06.all
Refutation found. Thanks to Tanya!
5532375. $false (10:0) [subsumption resolution 5532374,186043]
186043. ~cont(sK2,sK1) (12:3) [subsumption resolution 186013,213]
213. ~p(sK2,sK1) (0:3) [cnf transformation 200]
200. p(sK1,sK0) & ~zex(sK0) & ~zex(sK1) & ~zex(sK2) & ~zex(difference(sK0,sK1)) & ~cont(sK0,sK1) & p(sK2,difference(sK0,difference(sK0,sK1))) & ~p(sK2,sK1)[skolemisation 120]
120. ? [X0,X1,X2] : (p(X1,X0) & ~zex(X0) & ~zex(X1) & ~zex(X2) & ~zex(difference(X0,X1)) & ~cont(X0,X1) & p(X2,difference(X0,difference(X0,X1))) & ~p(X2,X1))[flattening 119]
119. ? [X0,X1,X2] : ((p(X1,X0) & ~zex(X0) & ~zex(X1) & ~zex(X2) & ~zex(difference(X0,X1)) & ~cont(X0,X1) & p(X2,difference(X0,difference(X0,X1)))) & ~p(X2,X1))[ennf transformation 62]
62. ~! [X0,X1,X2] : (~(p(X1,X0) & ~zex(X0) & ~zex(X1) & ~zex(X2) & ~zex(difference(X0,X1)) & ~cont(X0,X1) & p(X2,difference(X0,difference(X0,X1)))) | p(X2,X1))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : (~(p(X1,X0) & ~zex(X0) & ~zex(X1) & ~zex(X2) & ~zex(difference(X0,X1)) & ~cont(X0,X1) & p(X2,difference(X0,difference(X0,X1)))) | p(X2,X1))[negated conjecture]
186013. p(sK2,sK1) | ~cont(sK2,sK1) (12:6) [resolution 185439,243]
243. ~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 136]
136. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 135]
135. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 73]
73. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 50]
50. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
185439. eqdim(sK2,sK1) (11:3) [subsumption resolution 185430,185389]
185389. leq(sK2,sK1) (9:3) [subsumption resolution 185310,35142]
35142. p(sK4(sK0,sK1),sK0) (5:5) [resolution 35112,17662]
17662. ~p(X0,sK1) | p(X0,sK0) (4:6) [subsumption resolution 17606,1624]
1624. cont(X1,sK0) | ~p(X1,sK1) (7:6) [resolution 1616,241]
241. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 136]
1616. ~cont(X11,sK1) | cont(X11,sK0) (6:6) [resolution 293,966]
966. cont(sK1,sK0) (5:3) [subsumption resolution 965,208]
208. ~zex(sK1) (0:2) [cnf transformation 200]
965. zex(sK1) | cont(sK1,sK0) (5:5) [subsumption resolution 937,298]
298. zex(sK6) (0:2) [cnf transformation 204]
204. zex(sK6)[skolemisation 61]
61. ? [X0] : zex(X0)[input]
937. ~zex(sK6) | zex(sK1) | cont(sK1,sK0) (5:7) [superposition 244,929]
929. difference(sK1,sK0) = sK6 (4:5) [resolution 633,206]
206. p(sK1,sK0) (0:3) [cnf transformation 200]
633. ~p(X1,X2) | difference(X1,X2) = sK6 (3:8) [resolution 346,241]
346. ~cont(X2,X3) | difference(X2,X3) = sK6 (2:8) [resolution 339,246]
246. zex(difference(X0,X1)) | ~cont(X0,X1) (0:7) [cnf transformation 137]
137. ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))[ennf transformation 74]
74. ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & (~(zex(X0) | cont(X0,X1)) | zex(difference(X0,X1))))[flattening 15]
15. ! [X0] : ! [X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & (~(zex(X0) | cont(X0,X1)) | zex(difference(X0,X1))))[input]
339. ~zex(X0) | sK6 = X0 (1:5) [resolution 262,298]
262. ~zex(X1) | X0 = X1 | ~zex(X0) (0:7) [cnf transformation 149]
149. ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)[flattening 148]
148. ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)[ennf transformation 89]
89. ! [X0,X1] : (~(zex(X0) & zex(X1)) | X0 = X1)[flattening 41]
41. ! [X0] : ! [X1] : (~(zex(X0) & zex(X1)) | X0 = X1)[input]
244. ~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1) (0:9) [cnf transformation 137]
293. ~cont(X1,X2) | cont(X0,X2) | ~cont(X0,X1) (0:9) [cnf transformation 190]
190. ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))[flattening 189]
189. ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))[ennf transformation 114]
114. ! [X0,X1,X2] : (~(cont(X0,X1) & cont(X1,X2)) | cont(X0,X2))[flattening 45]
45. ! [X0] : ! [X1] : ! [X2] : (~(cont(X0,X1) & cont(X1,X2)) | cont(X0,X2))[input]
17606. ~p(X0,sK1) | p(X0,sK0) | ~cont(X0,sK0) (4:9) [resolution 17568,243]
17568. eqdim(X56,sK0) | ~p(X56,sK1) (3:6) [resolution 4485,206]
4485. ~p(X7,X6) | eqdim(X5,X6) | ~p(X5,X7) (2:9) [resolution 1515,242]
242. eqdim(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 136]
1515. ~eqdim(X5,X7) | eqdim(X5,X6) | ~p(X7,X6) (1:9) [resolution 288,242]
288. ~eqdim(X1,X2) | eqdim(X0,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 180]
180. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~eqdim(X1,X2) | eqdim(X0,X2))[flattening 179]
179. ! [X0,X1,X2] : ((~eqdim(X0,X1) | ~eqdim(X1,X2)) | eqdim(X0,X2))[ennf transformation 109]
109. ! [X0,X1,X2] : (~(eqdim(X0,X1) & eqdim(X1,X2)) | eqdim(X0,X2))[flattening 36]
36. ! [X0] : ! [X1] : ! [X2] : (~(eqdim(X0,X1) & eqdim(X1,X2)) | eqdim(X0,X2))[input]
35112. p(sK4(sK0,sK1),sK1) (4:5) [forward demodulation 35111,439]
439. difference(sK1,sK6) = sK1 (3:5) [resolution 425,208]
425. zex(X0) | difference(X0,sK6) = X0 (2:7) [resolution 256,349]
349. lt(sK6,X0) | zex(X0) (1:5) [resolution 265,298]
265. ~zex(X0) | zex(X1) | lt(X0,X1) (0:7) [cnf transformation 155]
155. ! [X0,X1] : (~zex(X0) | zex(X1) | lt(X0,X1))[flattening 154]
154. ! [X0,X1] : ((~zex(X0) | zex(X1)) | lt(X0,X1))[ennf transformation 92]
92. ! [X0,X1] : (~(zex(X0) & ~zex(X1)) | lt(X0,X1))[flattening 42]
42. ! [X0] : ! [X1] : (~(zex(X0) & ~zex(X1)) | lt(X0,X1))[input]
256. ~lt(X1,X0) | difference(X0,X1) = X0 (0:8) [cnf transformation 142]
142. ! [X0,X1] : (~lt(X1,X0) | difference(X0,X1) = X0)[ennf transformation 83]
83. ! [X0,X1] : (~lt(X1,X0) | difference(X0,X1) = X0)[flattening 11]
11. ! [X0] : ! [X1] : (~lt(X1,X0) | difference(X0,X1) = X0)[input]
35111. p(sK4(sK0,difference(sK1,sK6)),sK1) (4:7) [subsumption resolution 35094,3741]
3741. p(sK1,sK1) (7:3) [subsumption resolution 3740,1896]
1896. ~eqdim(sK1,sK6) (5:3) [resolution 1886,312]
312. ~lt(X1,X0) | ~eqdim(X0,X1) (1:6) [resolution 238,236]
236. ~leq(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 132]
132. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 131]
131. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[ennf transformation 71]
71. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[flattening 57]
57. ! [X0] : ! [X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[input]
238. leq(X0,X1) | ~eqdim(X0,X1) (0:6) [cnf transformation 134]
134. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 133]
133. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[ennf transformation 72]
72. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[flattening 60]
60. ! [X0] : ! [X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[input]
1886. lt(sK6,sK1) (4:3) [subsumption resolution 1878,207]
207. ~zex(sK0) (0:2) [cnf transformation 200]
1878. zex(sK0) | lt(sK6,sK1) (4:5) [resolution 1445,206]
1445. ~p(X1,X2) | zex(X2) | lt(sK6,X1) (3:8) [resolution 1396,325]
325. eqdim(X1,X2) | ~p(X2,X1) (1:6) [resolution 259,242]
259. ~eqdim(X0,X1) | eqdim(X1,X0) (0:6) [cnf transformation 145]
145. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 86]
86. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[flattening 35]
35. ! [X0] : ! [X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[input]
1396. ~eqdim(X3,X2) | lt(sK6,X2) | zex(X3) (2:8) [resolution 281,349]
281. ~lt(X2,X0) | lt(X2,X1) | ~eqdim(X0,X1) (0:9) [cnf transformation 171]
171. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[ennf transformation 102]
102. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[flattening 38]
38. ! [X0] : ! [X1] : ! [X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[input]
3740. p(sK1,sK1) | eqdim(sK1,sK6) (7:6) [resolution 3714,684]
684. ~leq(X1,sK6) | eqdim(X1,sK6) (2:6) [resolution 240,303]
303. leq(sK6,X0) (1:3) [resolution 251,298]
251. ~zex(X0) | leq(X0,X1) (0:5) [cnf transformation 138]
138. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 79]
79. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[flattening 56]
56. ! [X0] : ! [X1] : (~zex(X0) | leq(X0,X1))[input]
240. ~leq(X1,X0) | eqdim(X0,X1) | ~leq(X0,X1) (0:9) [cnf transformation 134]
3714. leq(sK1,X8) | p(sK1,sK1) (6:6) [superposition 1202,439]
1202. p(X1,X1) | leq(difference(X1,sK6),X2) (5:8) [resolution 1066,251]
1066. p(X1,X1) | zex(difference(X1,sK6)) (4:7) [subsumption resolution 1054,254]
254. zex(difference(X0,X1)) | X0 != X1 (0:7) [cnf transformation 140]
140. ! [X0,X1] : (X0 != X1 | zex(difference(X0,X1)))[ennf transformation 81]
81. ! [X0,X1] : (~X0 = X1 | zex(difference(X0,X1)))[flattening 7]
7. ! [X0] : ! [X1] : (~X0 = X1 | zex(difference(X0,X1)))[input]
1054. p(X1,X1) | zex(difference(X1,sK6)) | sK6 = X1 (4:10) [superposition 250,434]
434. difference(X19,sK6) = X19 | sK6 = X19 (3:8) [resolution 425,339]
250. p(difference(X0,X1),X0) | zex(difference(X0,X1)) (0:9) [cnf transformation 78]
78. ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))[flattening 16]
16. ! [X0] : ! [X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))[input]
35094. ~p(sK1,sK1) | p(sK4(sK0,difference(sK1,sK6)),sK1) (4:10) [superposition 7494,439]
7494. ~p(sK1,difference(X21,X22)) | p(sK4(sK0,difference(X21,X22)),X21) (2:12) [resolution 478,615]
615. po(X3,sK0) | ~p(sK1,X3) (1:6) [resolution 231,206]
231. ~p(X3,X1) | po(X0,X1) | ~p(X3,X0) (0:9) [cnf transformation 202]
202. ! [X0,X1] : ((~po(X0,X1) | (p(sK4(X1,X0),X0) & p(sK4(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[skolemisation 128]
128. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[ennf transformation 69]
69. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[flattening 68]
68. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[rectify 29]
29. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X2] : (p(X2,X0) & p(X2,X1)) | po(X0,X1)))[input]
478. ~po(difference(X1,X2),X0) | p(sK4(X0,difference(X1,X2)),X1) (1:12) [resolution 282,229]
229. p(sK4(X1,X0),X0) | ~po(X0,X1) (0:8) [cnf transformation 202]
282. ~p(X2,difference(X0,X1)) | p(X2,X0) (0:8) [cnf transformation 172]
172. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))[ennf transformation 103]
103. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))[flattening 21]
21. ! [X0] : ! [X1] : ! [X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))[input]
185310. ~p(sK4(sK0,sK1),sK0) | leq(sK2,sK1) (9:8) [resolution 185143,37395]
37395. ~eqdim(sK4(sK0,sK1),X6) | leq(X6,sK1) (8:8) [resolution 37340,4620]
4620. ~eqdim(X12,X11) | leq(X10,X11) | ~eqdim(X12,X10) (2:9) [resolution 1556,239]
239. leq(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 134]
1556. ~leq(X7,X9) | leq(X7,X8) | ~eqdim(X9,X8) (1:9) [resolution 289,238]
289. ~leq(X1,X2) | leq(X0,X2) | ~leq(X0,X1) (0:9) [cnf transformation 182]
182. ! [X0,X1,X2] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2))[flattening 181]
181. ! [X0,X1,X2] : ((~leq(X0,X1) | ~leq(X1,X2)) | leq(X0,X2))[ennf transformation 110]
110. ! [X0,X1,X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[flattening 54]
54. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[input]
37340. eqdim(sK4(sK0,sK1),sK1) (7:5) [subsumption resolution 37329,36405]
36405. ~eqdim(difference(sK4(sK0,sK1),sK6),sK6) (13:7) [resolution 36178,312]
36178. lt(sK6,difference(sK4(sK0,sK1),sK6)) (12:7) [resolution 35143,5743]
5743. ~po(X18,X19) | lt(sK6,difference(X18,sK6)) (6:8) [resolution 4376,229]
4376. ~p(X6,X5) | lt(sK6,difference(X5,sK6)) (5:8) [resolution 4355,2570]
2570. p(X3,difference(X2,sK6)) | ~p(X3,X2) (4:8) [subsumption resolution 2569,420]
420. ~po(X0,sK6) (3:3) [resolution 230,315]
315. ~p(X0,sK6) (2:3) [resolution 241,304]
304. ~cont(X0,sK6) (1:3) [resolution 252,298]
252. ~zex(X0) | ~cont(X1,X0) (0:5) [cnf transformation 139]
139. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 80]
80. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[flattening 46]
46. ! [X0] : ! [X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[input]
230. p(sK4(X1,X0),X1) | ~po(X0,X1) (0:8) [cnf transformation 202]
2569. po(X3,sK6) | p(X3,difference(X2,sK6)) | ~p(X3,X2) (4:11) [superposition 295,639]
639. intersection(X1,sK6) = sK6 (3:5) [resolution 348,397]
397. ~c(X0,sK6) (2:3) [resolution 227,304]
227. cont(sK3(X1,X0),X1) | ~c(X0,X1) (0:8) [cnf transformation 201]
201. ! [X0,X1] : ((~c(X0,X1) | (cont(sK3(X1,X0),X0) & cont(sK3(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 127]
127. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 67]
67. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 66]
66. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 49]
49. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
348. c(X6,X7) | intersection(X6,X7) = sK6 (2:8) [resolution 339,219]
219. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 123]
123. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[ennf transformation 63]
63. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[flattening 24]
24. ! [X0] : ! [X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[input]
295. po(X2,intersection(X0,X1)) | p(X2,difference(X0,X1)) | ~p(X2,X0) (0:13) [cnf transformation 194]
194. ! [X0,X1,X2] : (~p(X2,X0) | po(X2,intersection(X0,X1)) | p(X2,difference(X0,X1)))[flattening 193]
193. ! [X0,X1,X2] : ((~p(X2,X0) | po(X2,intersection(X0,X1))) | p(X2,difference(X0,X1)))[ennf transformation 116]
116. ! [X0,X1,X2] : (~(p(X2,X0) & ~po(X2,intersection(X0,X1))) | p(X2,difference(X0,X1)))[flattening 23]
23. ! [X0] : ! [X1] : ! [X2] : (~(p(X2,X0) & ~po(X2,intersection(X0,X1))) | p(X2,difference(X0,X1)))[input]
4355. ~p(X4,X3) | lt(sK6,X3) (3:6) [resolution 4352,241]
4352. ~cont(X1,X0) | lt(sK6,X0) (2:6) [subsumption resolution 4327,253]
253. ~zex(X0) | ~cont(X0,X1) (0:5) [cnf transformation 139]
4327. lt(sK6,X0) | ~cont(X1,X0) | zex(X1) (2:8) [resolution 1482,349]
1482. ~lt(X6,X8) | lt(X6,X7) | ~cont(X8,X7) (1:9) [subsumption resolution 1481,281]
1481. lt(X6,X7) | ~lt(X6,X8) | eqdim(X8,X7) | ~cont(X8,X7) (1:12) [resolution 287,260]
260. lt(X0,X1) | eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 146]
146. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[ennf transformation 87]
87. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 30]
30. ! [X0] : ! [X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[input]
287. ~lt(X1,X2) | lt(X0,X2) | ~lt(X0,X1) (0:9) [cnf transformation 178]
178. ! [X0,X1,X2] : (~lt(X0,X1) | ~lt(X1,X2) | lt(X0,X2))[flattening 177]
177. ! [X0,X1,X2] : ((~lt(X0,X1) | ~lt(X1,X2)) | lt(X0,X2))[ennf transformation 108]
108. ! [X0,X1,X2] : (~(lt(X0,X1) & lt(X1,X2)) | lt(X0,X2))[flattening 33]
33. ! [X0] : ! [X1] : ! [X2] : (~(lt(X0,X1) & lt(X1,X2)) | lt(X0,X2))[input]
35143. po(sK4(sK0,sK1),sK4(sK0,sK1)) (11:7) [resolution 35112,11225]
11225. ~p(X7,sK1) | po(X7,X7) (10:6) [resolution 11083,6458]
6458. ~po(X0,X1) | po(X0,X0) (2:6) [duplicate literal removal 6444]
6444. po(X0,X0) | ~po(X0,X1) | ~po(X0,X1) (2:9) [resolution 618,229]
618. ~p(sK4(X8,X7),X6) | po(X6,X7) | ~po(X7,X8) (1:11) [resolution 231,229]
11083. po(X7,sK1) | ~p(X7,sK1) (9:6) [subsumption resolution 11082,315]
11082. p(X7,sK6) | po(X7,sK1) | ~p(X7,sK1) (9:9) [forward demodulation 10991,929]
10991. po(X7,sK1) | p(X7,difference(sK1,sK0)) | ~p(X7,sK1) (9:11) [superposition 295,10938]
10938. intersection(sK1,sK0) = sK1 (8:5) [resolution 10933,1304]
1304. ~cont(X5,intersection(X5,X6)) | intersection(X5,X6) = X5 (1:10) [subsumption resolution 1299,252]
1299. intersection(X5,X6) = X5 | ~cont(X5,intersection(X5,X6)) | zex(intersection(X5,X6)) (1:14) [resolution 267,249]
249. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 77]
77. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 25]
25. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
267. ~cont(X1,X0) | X0 = X1 | ~cont(X0,X1) (0:9) [cnf transformation 159]
159. ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1)[flattening 158]
158. ! [X0,X1] : ((~cont(X0,X1) | ~cont(X1,X0)) | X0 = X1)[ennf transformation 94]
94. ! [X0,X1] : (~(cont(X0,X1) & cont(X1,X0)) | X0 = X1)[flattening 44]
44. ! [X0] : ! [X1] : (~(cont(X0,X1) & cont(X1,X0)) | X0 = X1)[input]
10933. cont(sK1,intersection(sK1,sK0)) (7:5) [subsumption resolution 10914,208]
10914. zex(sK1) | cont(sK1,intersection(sK1,sK0)) (7:7) [resolution 10901,244]
10901. zex(difference(sK1,intersection(sK1,sK0))) (6:6) [resolution 10842,7071]
7071. po(X4,X4) | zex(X4) (5:5) [duplicate literal removal 7040]
7040. zex(X4) | po(X4,X4) | zex(X4) (5:7) [resolution 6481,743]
743. p(X0,X0) | zex(X0) (2:5) [resolution 740,216]
216. cont(X0,X0) | zex(X0) (0:5) [cnf transformation 121]
121. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[ennf transformation 43]
43. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[input]
740. ~cont(X0,X0) | p(X0,X0) (1:6) [resolution 243,214]
214. eqdim(X0,X0) (0:3) [cnf transformation 34]
34. ! [X0] : eqdim(X0,X0)[input]
6481. ~p(X1,X0) | zex(X1) | po(X0,X0) (4:8) [resolution 6458,746]
746. po(X1,X0) | zex(X0) | ~p(X0,X1) (3:8) [resolution 743,231]
10842. ~po(difference(sK1,intersection(sK1,sK0)),X14) (5:7) [subsumption resolution 10820,315]
10820. p(sK4(X14,difference(sK1,intersection(sK1,sK0))),sK6) | ~po(difference(sK1,intersection(sK1,sK0)),X14) (5:16) [superposition 2148,929]
2148. p(sK4(X10,difference(X11,intersection(X11,X12))),difference(X11,X12)) | ~po(difference(X11,intersection(X11,X12)),X10) (1:18) [resolution 285,229]
285. ~p(X2,difference(X0,intersection(X0,X1))) | p(X2,difference(X0,X1)) (0:12) [cnf transformation 175]
175. ! [X0,X1,X2] : (~p(X2,difference(X0,intersection(X0,X1))) | p(X2,difference(X0,X1)))[ennf transformation 106]
106. ! [X0,X1,X2] : (~p(X2,difference(X0,intersection(X0,X1))) | p(X2,difference(X0,X1)))[flattening 4]
4. ! [X0] : ! [X1] : ! [X2] : (~p(X2,difference(X0,intersection(X0,X1))) | p(X2,difference(X0,X1)))[input]
37329. eqdim(sK4(sK0,sK1),sK1) | eqdim(difference(sK4(sK0,sK1),sK6),sK6) (7:12) [resolution 35155,3708]
3708. p(X9,X9) | eqdim(difference(X9,sK6),sK6) (6:8) [resolution 1202,684]
35155. ~p(sK4(sK0,sK1),X3) | eqdim(X3,sK1) (5:8) [resolution 35112,4484]
4484. ~p(X4,X3) | eqdim(X2,X3) | ~p(X4,X2) (2:9) [resolution 1515,325]
185143. eqdim(X170,sK2) | ~p(X170,sK0) (4:6) [resolution 16177,477]
477. p(sK2,sK0) (1:3) [resolution 282,212]
212. p(sK2,difference(sK0,difference(sK0,sK1))) (0:7) [cnf transformation 200]
16177. ~p(X12,X13) | eqdim(X11,X12) | ~p(X11,X13) (3:9) [resolution 1514,242]
1514. ~eqdim(X2,X4) | eqdim(X2,X3) | ~p(X3,X4) (2:9) [resolution 288,325]
185430. eqdim(sK2,sK1) | ~leq(sK2,sK1) (11:6) [resolution 185388,240]
185388. leq(sK1,sK2) (10:3) [subsumption resolution 185309,35142]
185309. ~p(sK4(sK0,sK1),sK0) | leq(sK1,sK2) (10:8) [resolution 185143,67406]
67406. ~eqdim(sK4(sK0,sK1),X9) | leq(sK1,X9) (9:8) [resolution 37403,312]
37403. lt(X14,sK4(sK0,sK1)) | leq(sK1,X14) (8:8) [resolution 37340,4985]
4985. ~eqdim(X12,X10) | lt(X11,X12) | leq(X10,X11) (2:9) [resolution 1561,239]
1561. ~leq(X22,X24) | leq(X22,X23) | lt(X23,X24) (1:9) [resolution 289,636]
636. leq(X1,X0) | lt(X0,X1) (0:6) [subsumption resolution 237,247]
247. leq(X1,X0) | leq(X0,X1) (0:6) [cnf transformation 75]
75. ! [X0,X1] : (leq(X0,X1) | leq(X1,X0))[flattening 53]
53. ! [X0] : ! [X1] : (leq(X0,X1) | leq(X1,X0))[input]
237. lt(X0,X1) | leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 132]
5532374. cont(sK2,sK1) (10:3) [subsumption resolution 5532268,209]
209. ~zex(sK2) (0:2) [cnf transformation 200]
5532268. zex(sK2) | cont(sK2,sK1) (10:5) [resolution 5532258,244]
5532258. zex(difference(sK2,sK1)) (9:4) [subsumption resolution 5532257,477]
5532257. zex(difference(sK2,sK1)) | ~p(sK2,sK0) (9:7) [subsumption resolution 5532249,125183]
125183. leq(X1,sK1) | ~p(X1,sK0) (11:6) [resolution 37896,37395]
37896. eqdim(sK4(sK0,sK1),X3) | ~p(X3,sK0) (10:8) [resolution 37883,1514]
37883. eqdim(sK4(sK0,sK1),sK0) (9:5) [resolution 37368,37340]
37368. ~eqdim(X0,sK1) | eqdim(X0,sK0) (8:6) [resolution 37345,288]
37345. eqdim(sK1,sK0) (7:3) [resolution 37320,259]
37320. eqdim(sK0,sK1) (6:3) [resolution 35155,35142]
5532249. zex(difference(sK2,sK1)) | ~leq(sK2,sK1) | ~p(sK2,sK0) (9:10) [resolution 5532243,7891]
7891. cont(difference(X30,X31),X32) | zex(difference(X30,X31)) | ~leq(X30,X31) | ~p(X30,X32) (2:15) [resolution 1964,1614]
1614. ~cont(X5,X7) | cont(X5,X6) | ~p(X7,X6) (1:9) [resolution 293,241]
1964. cont(difference(X0,X1),X0) | ~leq(X0,X1) | zex(difference(X0,X1)) (1:12) [resolution 291,216]
291. ~cont(X2,difference(X0,X1)) | cont(X2,X0) | ~leq(X0,X1) (0:11) [cnf transformation 186]
186. ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,difference(X0,X1)) | cont(X2,X0))[flattening 185]
185. ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,difference(X0,X1))) | cont(X2,X0))[ennf transformation 112]
112. ! [X0,X1,X2] : (~(leq(X0,X1) & cont(X2,difference(X0,X1))) | cont(X2,X0))[flattening 13]
13. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & cont(X2,difference(X0,X1))) | cont(X2,X0))[input]
5532243. ~cont(difference(sK2,sK1),sK0) (8:5) [superposition 5532211,174449]
174449. difference(X53,sK6) = X53 (5:5) [subsumption resolution 174448,218]
218. ~lt(X0,X0) (0:3) [cnf transformation 122]
122. ! [X0] : ~lt(X0,X0)[ennf transformation 31]
31. ! [X0] : ~lt(X0,X0)[input]
174448. lt(sK6,sK6) | difference(X53,sK6) = X53 (5:8) [subsumption resolution 174220,65884]
65884. ~eqdim(sK6,X5) | difference(X5,X6) = X5 (4:8) [resolution 65773,258]
258. ~lt(X0,X1) | ~eqdim(X0,X1) (0:6) [cnf transformation 144]
144. ! [X0,X1] : (~lt(X0,X1) | ~eqdim(X0,X1))[ennf transformation 85]
85. ! [X0,X1] : (~lt(X0,X1) | ~eqdim(X0,X1))[flattening 37]
37. ! [X0] : ! [X1] : (~lt(X0,X1) | ~eqdim(X0,X1))[input]
65773. lt(sK6,X45) | difference(X45,X46) = X45 (3:8) [subsumption resolution 65677,1573]
1573. leq(X14,X15) | lt(sK6,X14) (3:6) [resolution 1554,636]
1554. ~leq(X2,sK6) | leq(X2,X3) (2:6) [resolution 289,303]
65677. difference(X45,X46) = X45 | ~leq(X45,X46) | lt(sK6,X45) (3:11) [resolution 13721,4352]
13721. cont(sK5(difference(X0,X1),X0),X0) | difference(X0,X1) = X0 | ~leq(X0,X1) (2:15) [factoring 2318]
2318. cont(sK5(difference(X18,X19),X20),X20) | cont(sK5(difference(X18,X19),X20),X18) | difference(X18,X19) = X20 | ~leq(X18,X19) (1:22) [resolution 271,291]
271. cont(sK5(X1,X0),X1) | cont(sK5(X1,X0),X0) | X0 = X1 (0:13) [cnf transformation 203]
203. ! [X0,X1] : (((cont(sK5(X1,X0),X0) & ~cont(sK5(X1,X0),X1)) | (cont(sK5(X1,X0),X1) & ~cont(sK5(X1,X0),X0))) | X0 = X1)[skolemisation 166]
166. ! [X0,X1] : (? [X2] : ((cont(X2,X0) & ~cont(X2,X1)) | (cont(X2,X1) & ~cont(X2,X0))) | X0 = X1)[ennf transformation 98]
98. ! [X0,X1] : (~! [X2] : ((~cont(X2,X0) | cont(X2,X1)) & (~cont(X2,X1) | cont(X2,X0))) | X0 = X1)[flattening 48]
48. ! [X0] : ! [X1] : (~! [X2] : ((~cont(X2,X0) | cont(X2,X1)) & (~cont(X2,X1) | cont(X2,X0))) | X0 = X1)[input]
174220. lt(sK6,sK6) | eqdim(sK6,X53) | difference(X53,sK6) = X53 (5:11) [resolution 4022,709]
709. eqdim(X0,sK6) | difference(X0,sK6) = X0 (4:8) [resolution 696,256]
696. lt(sK6,X0) | eqdim(X0,sK6) (3:6) [resolution 684,636]
4022. ~eqdim(X5,X6) | lt(sK6,X6) | eqdim(sK6,X5) (3:9) [resolution 4002,281]
4002. lt(sK6,X1) | eqdim(sK6,X1) (2:6) [resolution 685,303]
685. ~leq(X2,X3) | eqdim(X2,X3) | lt(X2,X3) (1:9) [resolution 240,636]
5532211. ~cont(difference(difference(sK2,sK1),X0),sK0) (7:7) [subsumption resolution 5532210,65642]
65642. leq(sK0,sK1) (9:3) [subsumption resolution 65629,35110]
35110. p(sK4(sK0,sK0),sK0) (4:5) [forward demodulation 35109,438]
438. difference(sK0,sK6) = sK0 (3:5) [resolution 425,207]
35109. p(sK4(sK0,difference(sK0,sK6)),sK0) (4:7) [subsumption resolution 35093,206]
35093. ~p(sK1,sK0) | p(sK4(sK0,difference(sK0,sK6)),sK0) (4:10) [superposition 7494,438]
65629. leq(sK0,sK1) | ~p(sK4(sK0,sK0),sK0) (9:8) [resolution 36914,37372]
37372. eqdim(sK1,X3) | ~p(X3,sK0) (8:6) [resolution 37345,1514]
36914. ~eqdim(X10,sK4(sK0,sK0)) | leq(sK0,X10) (8:8) [resolution 36899,4706]
4706. ~eqdim(X12,X10) | ~eqdim(X11,X12) | leq(X10,X11) (2:9) [resolution 1557,239]
1557. ~leq(X10,X12) | leq(X10,X11) | ~eqdim(X11,X12) (1:9) [resolution 289,239]
36899. eqdim(sK4(sK0,sK0),sK0) (7:5) [subsumption resolution 36890,36491]
36491. ~eqdim(difference(sK4(sK0,sK0),sK6),sK6) (13:7) [resolution 36197,312]
36197. lt(sK6,difference(sK4(sK0,sK0),sK6)) (12:7) [resolution 36141,4376]
36141. p(sK4(sK0,sK0),sK4(sK0,sK0)) (11:7) [resolution 35121,8701]
8701. ~c(X19,X20) | p(X19,X19) (10:6) [resolution 8686,226]
226. cont(sK3(X1,X0),X0) | ~c(X0,X1) (0:8) [cnf transformation 201]
8686. ~cont(X1,X0) | p(X0,X0) (9:6) [subsumption resolution 8647,8427]
8427. ~p(X1,X0) | p(X0,X0) (7:6) [resolution 3826,2570]
3826. ~p(X4,difference(X3,sK6)) | p(X3,X3) (6:8) [resolution 1203,241]
1203. ~cont(X4,difference(X3,sK6)) | p(X3,X3) (5:8) [resolution 1066,252]
8647. p(X0,X0) | p(X1,X0) | ~cont(X1,X0) (9:9) [resolution 8463,2766]
2766. ~zex(difference(X7,X6)) | p(X6,X7) | ~cont(X6,X7) (1:10) [subsumption resolution 2765,253]
2765. zex(X6) | p(X6,X7) | ~zex(difference(X7,X6)) | ~cont(X6,X7) (1:12) [subsumption resolution 2745,252]
2745. zex(X6) | zex(X7) | p(X6,X7) | ~zex(difference(X7,X6)) | ~cont(X6,X7) (1:14) [resolution 268,246]
268. ~zex(difference(X1,X0)) | zex(X1) | zex(X0) | p(X1,X0) | ~zex(difference(X0,X1)) (0:15) [cnf transformation 161]
161. ! [X0,X1] : (~zex(difference(X0,X1)) | ~zex(difference(X1,X0)) | zex(X0) | zex(X1) | p(X1,X0))[flattening 160]
160. ! [X0,X1] : ((~zex(difference(X0,X1)) | ~zex(difference(X1,X0)) | zex(X0) | zex(X1)) | p(X1,X0))[ennf transformation 95]
95. ! [X0,X1] : (~(zex(difference(X0,X1)) & zex(difference(X1,X0)) & ~zex(X0) & ~zex(X1)) | p(X1,X0))[flattening 9]
9. ! [X0] : ! [X1] : (~(zex(difference(X0,X1)) & zex(difference(X1,X0)) & ~zex(X0) & ~zex(X1)) | p(X1,X0))[input]
8463. p(X16,X16) | zex(difference(X16,X17)) (8:7) [resolution 8427,250]
35121. c(sK4(sK0,sK0),sK4(sK0,sK0)) (8:7) [resolution 35110,1344]
1344. ~p(X3,X4) | c(X3,X3) (7:6) [resolution 1325,241]
1325. ~cont(X2,X3) | c(X2,X2) (6:6) [resolution 1295,558]
558. c(X0,X1) | ~cont(X1,X0) (1:6) [subsumption resolution 554,253]
554. c(X0,X1) | ~cont(X1,X0) | zex(X1) (1:8) [resolution 228,216]
228. ~cont(X3,X1) | c(X0,X1) | ~cont(X3,X0) (0:9) [cnf transformation 201]
1295. ~c(X10,X9) | c(X9,X9) (5:6) [resolution 1284,227]
1284. ~cont(X3,X2) | c(X2,X2) (4:6) [resolution 1268,558]
1268. ~c(X0,X1) | c(X0,X0) (3:6) [resolution 1254,220]
220. ~zex(intersection(X0,X1)) | ~c(X0,X1) (0:7) [cnf transformation 123]
1254. c(X0,X0) | zex(intersection(X0,X1)) (2:7) [resolution 986,249]
986. ~cont(intersection(X8,X9),X10) | c(X10,X8) (1:8) [subsumption resolution 978,253]
978. zex(intersection(X8,X9)) | c(X10,X8) | ~cont(intersection(X8,X9),X10) (1:12) [resolution 249,228]
36890. eqdim(sK4(sK0,sK0),sK0) | eqdim(difference(sK4(sK0,sK0),sK6),sK6) (7:12) [resolution 35127,3708]
35127. ~p(sK4(sK0,sK0),X3) | eqdim(X3,sK0) (5:8) [resolution 35110,4484]
5532210. ~cont(difference(difference(sK2,sK1),X0),sK0) | ~leq(sK0,sK1) (7:10) [subsumption resolution 5532179,185389]
5532179. ~leq(sK2,sK1) | ~cont(difference(difference(sK2,sK1),X0),sK0) | ~leq(sK0,sK1) (7:13) [resolution 905740,14310]
14310. cont(difference(difference(X0,X1),X2),difference(X3,X1)) | ~leq(X0,X1) | ~cont(difference(difference(X0,X1),X2),X3) | ~leq(X3,X1) (2:22) [subsumption resolution 14270,253]
14270. ~leq(X0,X1) | zex(difference(difference(X0,X1),X2)) | cont(difference(difference(X0,X1),X2),difference(X3,X1)) | ~cont(difference(difference(X0,X1),X2),X3) | ~leq(X3,X1) (2:28) [resolution 2494,290]
290. ~lt(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)) | ~cont(X2,X0) | ~leq(X0,X1) (0:16) [cnf transformation 184]
184. ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X0) | ~lt(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)))[flattening 183]
183. ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X0) | ~lt(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))[ennf transformation 111]
111. ! [X0,X1,X2] : (~(leq(X0,X1) & cont(X2,X0) & lt(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))[flattening 12]
12. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & cont(X2,X0) & lt(intersection(X2,X1),X2)) | cont(X2,difference(X0,X1)))[input]
2494. lt(intersection(difference(difference(X4,X5),X6),X5),difference(difference(X4,X5),X6)) | ~leq(X4,X5) | zex(difference(difference(X4,X5),X6)) (1:22) [resolution 292,250]
292. ~p(X2,difference(X0,X1)) | lt(intersection(X2,X1),X2) | ~leq(X0,X1) (0:13) [cnf transformation 188]
188. ! [X0,X1,X2] : (~leq(X0,X1) | ~p(X2,difference(X0,X1)) | lt(intersection(X2,X1),X2))[flattening 187]
187. ! [X0,X1,X2] : ((~leq(X0,X1) | ~p(X2,difference(X0,X1))) | lt(intersection(X2,X1),X2))[ennf transformation 113]
113. ! [X0,X1,X2] : (~(leq(X0,X1) & p(X2,difference(X0,X1))) | lt(intersection(X2,X1),X2))[flattening 14]
14. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & p(X2,difference(X0,X1))) | lt(intersection(X2,X1),X2))[input]
905740. ~cont(difference(difference(sK2,X1101),X1102),difference(sK0,sK1)) (6:9) [subsumption resolution 904715,368385]
368385. leq(sK2,difference(sK0,sK1)) (11:5) [superposition 368284,174449]
368284. leq(difference(sK2,X61),difference(sK0,sK1)) (10:7) [subsumption resolution 368224,3727]
3727. p(sK0,sK0) (7:3) [subsumption resolution 3726,2137]
2137. ~eqdim(sK0,sK6) (5:3) [resolution 2128,312]
2128. lt(sK6,sK0) (4:3) [subsumption resolution 2118,208]
2118. zex(sK1) | lt(sK6,sK0) (4:5) [resolution 1446,206]
1446. ~p(X4,X3) | zex(X4) | lt(sK6,X3) (3:8) [resolution 1396,242]
3726. p(sK0,sK0) | eqdim(sK0,sK6) (7:6) [resolution 3713,684]
3713. leq(sK0,X7) | p(sK0,sK0) (6:6) [superposition 1202,438]
368224. leq(difference(sK2,X61),difference(sK0,sK1)) | ~p(sK0,sK0) (10:10) [resolution 334733,167388]
167388. eqdim(X9,difference(sK0,sK1)) | ~p(sK0,X9) (6:8) [resolution 165308,259]
165308. eqdim(difference(sK0,sK1),X24) | ~p(sK0,X24) (5:8) [subsumption resolution 165282,43484]
43484. difference(sK0,difference(sK0,sK1)) != sK0 (16:7) [resolution 43435,23362]
23362. p(sK0,X24) | sK0 != X24 (12:6) [duplicate literal removal 23323]
23323. p(sK0,X24) | sK0 != X24 | sK0 != X24 (12:9) [resolution 6052,16397]
16397. cont(sK0,X36) | sK0 != X36 (11:6) [resolution 1620,12382]
12382. cont(sK0,sK0) (10:3) [backward demodulation 12369,8896]
8896. cont(intersection(sK0,sK0),sK0) (10:5) [resolution 8501,278]
278. ~p(X2,intersection(X0,X1)) | cont(X2,X1) (0:8) [cnf transformation 169]
169. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 168]
168. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[ennf transformation 100]
100. ! [X0,X1,X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 27]
27. ! [X0] : ! [X1] : ! [X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[input]
8501. p(intersection(sK0,sK0),intersection(sK0,sK0)) (9:7) [resolution 8479,6549]
6549. po(intersection(sK0,sK0),intersection(sK0,sK0)) (5:7) [subsumption resolution 6537,477]
6537. po(intersection(sK0,sK0),intersection(sK0,sK0)) | ~p(sK2,sK0) (5:10) [resolution 6482,2092]
2092. p(X10,intersection(sK0,sK0)) | ~p(X10,sK0) (4:8) [forward demodulation 2085,438]
2085. ~p(X10,sK0) | p(X10,intersection(sK0,difference(sK0,sK6))) (4:10) [superposition 283,438]
283. ~p(X2,difference(X0,X1)) | p(X2,intersection(X0,difference(X0,X1))) (0:12) [cnf transformation 173]
173. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,intersection(X0,difference(X0,X1))))[ennf transformation 104]
104. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,intersection(X0,difference(X0,X1))))[flattening 5]
5. ! [X0] : ! [X1] : ! [X2] : (~p(X2,difference(X0,X1)) | p(X2,intersection(X0,difference(X0,X1))))[input]
6482. ~p(sK2,X2) | po(X2,X2) (3:6) [resolution 6458,616]
616. po(X4,difference(sK0,difference(sK0,sK1))) | ~p(sK2,X4) (1:10) [resolution 231,212]
8479. ~po(X25,X26) | p(X25,X25) (8:6) [resolution 8427,229]
12369. intersection(sK0,sK0) = sK0 (9:5) [resolution 11938,207]
11938. zex(X4) | intersection(X4,X4) = X4 (8:7) [resolution 10877,1304]
10877. cont(X2,intersection(X2,X2)) | zex(X2) (7:7) [resolution 10859,244]
10859. zex(difference(X2,intersection(X2,X2))) (6:6) [resolution 10837,7071]
10837. ~po(difference(X0,intersection(X0,X0)),X1) (4:7) [subsumption resolution 10812,315]
10812. p(sK4(X1,difference(X0,intersection(X0,X0))),sK6) | ~po(difference(X0,intersection(X0,X0)),X1) (4:16) [superposition 2148,581]
581. difference(X0,X0) = sK6 (3:5) [equality resolution 345]
345. X0 != X1 | difference(X0,X1) = sK6 (2:8) [resolution 339,254]
1620. ~cont(X2,X4) | cont(X2,X3) | X3 != X4 (2:9) [subsumption resolution 1613,252]
1613. cont(X2,X3) | ~cont(X2,X4) | zex(X4) | X3 != X4 (2:11) [resolution 293,800]
800. cont(X0,X1) | zex(X0) | X0 != X1 (1:8) [resolution 244,254]
6052. ~cont(X4,X5) | p(X4,X5) | X4 != X5 (2:9) [resolution 2766,254]
43435. ~p(sK0,difference(sK0,difference(sK0,sK1))) (15:7) [resolution 43335,284]
284. ~po(X2,intersection(X0,X1)) | ~p(X2,difference(X0,X1)) (0:10) [cnf transformation 174]
174. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | ~po(X2,intersection(X0,X1)))[ennf transformation 105]
105. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | ~po(X2,intersection(X0,X1)))[flattening 22]
22. ! [X0] : ! [X1] : ! [X2] : (~p(X2,difference(X0,X1)) | ~po(X2,intersection(X0,X1)))[input]
43335. po(sK0,intersection(sK0,difference(sK0,sK1))) (14:7) [resolution 42509,6457]
6457. ~po(X3,X2) | po(X2,X3) (2:6) [duplicate literal removal 6445]
6445. po(X2,X3) | ~po(X3,X2) | ~po(X3,X2) (2:9) [resolution 618,230]
42509. po(intersection(sK0,difference(sK0,sK1)),sK0) (13:7) [resolution 9843,9696]
9696. po(difference(sK0,sK1),sK0) (12:5) [subsumption resolution 9691,210]
210. ~zex(difference(sK0,sK1)) (0:4) [cnf transformation 200]
9691. po(difference(sK0,sK1),sK0) | zex(difference(sK0,sK1)) (12:9) [resolution 8735,1043]
1043. ~p(difference(X2,X3),X4) | po(X4,X2) | zex(difference(X2,X3)) (1:12) [resolution 250,231]
8735. p(difference(sK0,sK1),difference(sK0,sK1)) (11:7) [resolution 8701,7456]
7456. c(difference(sK0,sK1),difference(sK0,sK1)) (7:7) [equality resolution 3904]
3904. difference(sK0,sK1) != X40 | c(X40,X40) (6:8) [resolution 1289,210]
1289. zex(X2) | c(X1,X1) | X1 != X2 (5:8) [resolution 1284,800]
9843. ~po(difference(X55,X56),X57) | po(intersection(X55,difference(X55,X56)),X57) (2:12) [duplicate literal removal 9822]
9822. ~po(difference(X55,X56),X57) | po(intersection(X55,difference(X55,X56)),X57) | ~po(difference(X55,X56),X57) (2:17) [resolution 2079,619]
619. ~p(sK4(X10,X11),X9) | po(X9,X10) | ~po(X11,X10) (1:11) [resolution 231,230]
2079. p(sK4(X10,difference(X11,X12)),intersection(X11,difference(X11,X12))) | ~po(difference(X11,X12),X10) (1:16) [resolution 283,229]
165282. difference(sK0,difference(sK0,sK1)) = sK0 | eqdim(difference(sK0,sK1),X24) | ~p(sK0,X24) (5:15) [resolution 164763,6630]
6630. ~cont(X0,X1) | difference(X1,X0) = X1 | eqdim(X0,X2) | ~p(X1,X2) (2:14) [resolution 1224,1515]
1224. eqdim(X1,X2) | ~cont(X1,X2) | difference(X2,X1) = X2 (1:11) [resolution 260,256]
164763. cont(difference(sK0,sK1),sK0) (4:5) [subsumption resolution 164762,210]
164762. zex(difference(sK0,sK1)) | cont(difference(sK0,sK1),sK0) (4:9) [forward demodulation 164103,436]
436. difference(sK0,sK1) = difference(difference(sK0,sK1),sK6) (3:9) [resolution 425,210]
164103. cont(difference(sK0,sK1),sK0) | zex(difference(difference(sK0,sK1),sK6)) (4:11) [superposition 10555,436]
10555. cont(difference(difference(X15,X16),X17),X15) | zex(difference(difference(X15,X16),X17)) (2:13) [resolution 2076,277]
277. ~p(X2,intersection(X0,X1)) | cont(X2,X0) (0:8) [cnf transformation 169]
2076. p(difference(difference(X4,X5),X6),intersection(X4,difference(X4,X5))) | zex(difference(difference(X4,X5),X6)) (1:17) [resolution 283,250]
334733. ~eqdim(sK0,X16) | leq(difference(sK2,X15),X16) (9:8) [resolution 334712,1556]
334712. leq(difference(sK2,X56),sK0) (8:5) [subsumption resolution 334680,310336]
310336. ~lt(difference(sK0,sK1),sK0) (11:5) [resolution 310313,37820]
37820. lt(X1,sK1) | ~lt(X1,sK0) (8:6) [resolution 37360,236]
37360. leq(sK0,X13) | lt(X13,sK1) (7:6) [resolution 37320,4984]
4984. ~eqdim(X7,X9) | lt(X8,X9) | leq(X7,X8) (2:9) [resolution 1561,238]
310313. ~lt(difference(sK0,sK1),sK1) (10:5) [resolution 310269,236]
310269. leq(sK1,difference(sK0,sK1)) (9:5) [superposition 309357,174449]
309357. leq(difference(sK1,X39),difference(sK0,sK1)) (8:7) [subsumption resolution 309278,3727]
309278. leq(difference(sK1,X39),difference(sK0,sK1)) | ~p(sK0,sK0) (8:10) [resolution 304430,165308]
304430. ~eqdim(X389,sK0) | leq(difference(sK1,X388),X389) (7:8) [resolution 20007,37320]
20007. ~eqdim(X18,X19) | leq(difference(X19,X20),X17) | ~eqdim(X17,X18) (3:11) [subsumption resolution 19971,251]
19971. ~eqdim(X17,X18) | leq(difference(X19,X20),X17) | ~eqdim(X18,X19) | zex(difference(X19,X20)) (3:15) [resolution 4706,1516]
1516. eqdim(X8,difference(X9,X10)) | ~eqdim(X8,X9) | zex(difference(X9,X10)) (1:12) [resolution 288,248]
248. eqdim(X0,difference(X0,X1)) | zex(difference(X0,X1)) (0:9) [cnf transformation 76]
76. ! [X0,X1] : (zex(difference(X0,X1)) | eqdim(X0,difference(X0,X1)))[flattening 10]
10. ! [X0] : ! [X1] : (zex(difference(X0,X1)) | eqdim(X0,difference(X0,X1)))[input]
334680. lt(difference(sK0,sK1),sK0) | leq(difference(sK2,X56),sK0) (8:10) [resolution 329946,248950]
248950. ~leq(X3,difference(sK0,sK1)) | leq(X3,sK0) (7:8) [resolution 248789,289]
248789. leq(difference(sK0,sK1),sK0) (6:5) [resolution 207129,218]
207129. lt(difference(sK0,sK1),X156) | leq(X156,sK0) (5:8) [resolution 22922,164763]
22922. ~cont(X31,X30) | lt(X31,X29) | leq(X29,X30) (3:9) [resolution 4989,1233]
1233. ~lt(X8,X7) | ~cont(X7,X8) (1:6) [subsumption resolution 1227,312]
1227. eqdim(X7,X8) | ~cont(X7,X8) | ~lt(X8,X7) (1:9) [resolution 260,257]
257. ~lt(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 143]
143. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[ennf transformation 84]
84. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[flattening 32]
32. ! [X0] : ! [X1] : (~lt(X0,X1) | ~lt(X1,X0))[input]
4989. leq(X22,X23) | lt(X23,X24) | lt(X24,X22) (2:9) [resolution 1561,636]
329946. leq(difference(sK2,X249),X248) | lt(X248,sK0) (4:8) [resolution 38705,477]
38705. ~p(X69,X70) | lt(X72,X70) | leq(difference(X69,X71),X72) (3:11) [subsumption resolution 38628,251]
38628. ~p(X69,X70) | zex(difference(X69,X71)) | lt(X72,X70) | leq(difference(X69,X71),X72) (3:15) [resolution 4491,4984]
4491. eqdim(difference(X21,X22),X23) | ~p(X21,X23) | zex(difference(X21,X22)) (2:12) [resolution 1515,875]
875. eqdim(difference(X2,X3),X2) | zex(difference(X2,X3)) (1:9) [resolution 248,259]
904715. ~cont(difference(difference(sK2,X1101),X1102),difference(sK0,sK1)) | ~leq(sK2,difference(sK0,sK1)) (6:14) [superposition 142985,903591]
903591. difference(sK2,difference(sK0,sK1)) = sK2 (5:7) [resolution 903575,606415]
606415. ~p(X4,difference(X4,X5)) | difference(X4,X5) = X4 (3:10) [subsumption resolution 50457,306637]
306637. leq(X0,X1) | difference(X0,X1) = X0 (4:8) [resolution 22914,218]
22914. leq(X2,X3) | lt(X4,X2) | difference(X4,X3) = X4 (3:11) [resolution 4989,256]
50457. difference(X4,X5) = X4 | ~leq(X4,X5) | ~p(X4,difference(X4,X5)) (3:13) [resolution 7927,241]
7927. ~cont(X10,difference(X10,X11)) | difference(X10,X11) = X10 | ~leq(X10,X11) (2:13) [subsumption resolution 7882,252]
7882. ~leq(X10,X11) | zex(difference(X10,X11)) | difference(X10,X11) = X10 | ~cont(X10,difference(X10,X11)) (2:17) [resolution 1964,267]
903575. p(sK2,difference(sK2,difference(sK0,sK1))) (4:7) [subsumption resolution 903574,209]
903574. zex(sK2) | p(sK2,difference(sK2,difference(sK0,sK1))) (4:9) [forward demodulation 903461,676928]
676928. difference(sK2,intersection(sK2,difference(sK0,sK1))) = sK2 (3:9) [subsumption resolution 51626,167501]
167501. eqdim(sK0,difference(sK0,sK1)) (11:5) [resolution 167492,40392]
40392. ~eqdim(sK1,X5) | eqdim(sK0,X5) (10:6) [subsumption resolution 40343,38149]
38149. ~lt(sK0,X0) | ~eqdim(sK1,X0) (9:6) [resolution 37375,236]
37375. leq(X6,sK0) | ~eqdim(sK1,X6) (8:6) [resolution 37345,4620]
40343. ~eqdim(sK1,X5) | lt(sK0,X5) | eqdim(sK0,X5) (10:9) [resolution 39108,4009]
4009. lt(X15,X14) | lt(X14,X15) | eqdim(X14,X15) (2:9) [resolution 685,636]
39108. ~lt(X10,sK0) | ~eqdim(sK1,X10) (9:6) [resolution 37820,312]
167492. eqdim(sK1,difference(sK0,sK1)) (9:5) [subsumption resolution 167426,3727]
167426. ~p(sK0,sK0) | eqdim(sK1,difference(sK0,sK1)) (9:8) [resolution 165308,37523]
37523. ~eqdim(X1,sK0) | eqdim(sK1,X1) (8:6) [resolution 37346,259]
37346. eqdim(X0,sK1) | ~eqdim(X0,sK0) (7:6) [resolution 37320,288]
51626. difference(sK2,intersection(sK2,difference(sK0,sK1))) = sK2 | ~eqdim(sK0,difference(sK0,sK1)) (3:14) [resolution 7158,238]
7158. ~leq(sK0,difference(sK0,sK1)) | difference(sK2,intersection(sK2,difference(sK0,sK1))) = sK2 (2:14) [resolution 2496,256]
2496. lt(intersection(sK2,difference(sK0,sK1)),sK2) | ~leq(sK0,difference(sK0,sK1)) (1:12) [resolution 292,212]
903461. p(sK2,difference(sK2,difference(sK0,sK1))) | zex(difference(sK2,intersection(sK2,difference(sK0,sK1)))) (4:15) [superposition 2145,676928]
2145. p(difference(X2,intersection(X2,X3)),difference(X2,X3)) | zex(difference(X2,intersection(X2,X3))) (3:15) [resolution 285,743]
142985. ~cont(difference(difference(difference(X10,X11),X12),X13),X11) | ~leq(X10,X11) (3:12) [subsumption resolution 142984,253]
142984. zex(difference(difference(difference(X10,X11),X12),X13)) | ~leq(X10,X11) | ~cont(difference(difference(difference(X10,X11),X12),X13),X11) (3:20) [subsumption resolution 142910,216]
142910. zex(difference(difference(difference(X10,X11),X12),X13)) | ~leq(X10,X11) | ~cont(difference(difference(difference(X10,X11),X12),X13),difference(difference(difference(X10,X11),X12),X13)) | ~cont(difference(difference(difference(X10,X11),X12),X13),X11) (3:35) [resolution 7999,2030]
2030. ~lt(intersection(X9,X8),X7) | ~cont(X7,X9) | ~cont(X7,X8) (1:11) [resolution 294,236]
294. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 192]
192. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 191]
191. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 115]
115. ! [X0,X1,X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[flattening 26]
26. ! [X0] : ! [X1] : ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[input]
7999. lt(intersection(difference(difference(difference(X48,X49),X50),X51),X49),difference(difference(difference(X48,X49),X50),X51)) | zex(difference(difference(difference(X48,X49),X50),X51)) | ~leq(X48,X49) (2:28) [resolution 1048,292]
1048. p(difference(difference(X15,X16),X17),X15) | zex(difference(difference(X15,X16),X17)) (1:13) [resolution 250,282]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 26745
Passive clauses: 2732950
Generated clauses: 5532137
Final active clauses: 24152
Final passive clauses: 2671768
Input formulas: 61
Initial clauses: 95

Duplicate literals: 49070
Trivial inequalities: 8929
Fw subsumption resolutions: 316872
Fw demodulations: 381729
Bw demodulations: 775

Simple tautologies: 5879
Equational tautologies: 4764
Forward subsumptions: 2067843
Backward subsumptions: 1818

Binary resolution: 4462890
Factoring: 199
Forward superposition: 215754
Backward superposition: 96535
Self superposition: 5
Equality resolution: 850

Unique components: 89

Memory used: 750819KB
Time elapsed: 384.528 s
------------------------------
% Success in time 596.33 s
