Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal10.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 535
Passive clauses: 2049
Generated clauses: 3947
Final active clauses: 531
Final passive clauses: 1400
Input formulas: 43
Initial clauses: 90

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 29
Fw subsumption resolutions: 97

Simple tautologies: 42
Forward subsumptions: 1756

Binary resolution: 3619
Factoring: 25

Backtracking splits: 57
Backtracking splits refuted: 13
Backtracking splits refuted at zero level: 4

Memory used: 895KB
Time elapsed: 0.403 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal10.all
Refutation found. Thanks to Tanya!
899. $false (5:0) [subsumption resolution 891,435]
435. ~zex(sK1) (3:2) [subsumption resolution 429,329]
329. zex(sK0) | ~zex(sK1) (2:4) [resolution 239,310]
310. ~lt(sK1,sK0) (1:3) [resolution 235,214]
214. lt(sK0,sK1) (0:3) [cnf transformation 168]
168. inc(sK0,sK1) & lt(sK0,sK1) & ~lt(intersection(sK0,sK1),sK1)[skolemisation 120]
120. ? [X0,X1] : (inc(X0,X1) & lt(X0,X1) & ~lt(intersection(X0,X1),X1))[flattening 119]
119. ? [X0,X1] : ((inc(X0,X1) & lt(X0,X1)) & ~lt(intersection(X0,X1),X1))[ennf transformation 44]
44. ~! [X0,X1] : ((inc(X0,X1) & lt(X0,X1)) => lt(intersection(X0,X1),X1))[flattening 1]
1. ~! [X0] : ! [X1] : ((inc(X0,X1) & lt(X0,X1)) => lt(intersection(X0,X1),X1))[negated conjecture]
235. ~lt(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 133]
133. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[ennf transformation 64]
64. ! [X0,X1] : (lt(X0,X1) => ~lt(X1,X0))[flattening 63]
63. ! [X0] : ! [X1] : (lt(X0,X1) => ~lt(X1,X0))[rectify 3]
3. ! [X3] : ! [X4] : (lt(X3,X4) => ~lt(X4,X3))[input]
239. lt(X0,X1) | zex(X1) | ~zex(X0) (0:7) [cnf transformation 141]
141. ! [X0,X1] : (~zex(X0) | zex(X1) | lt(X0,X1))[flattening 140]
140. ! [X0,X1] : ((~zex(X0) | zex(X1)) | lt(X0,X1))[ennf transformation 72]
72. ! [X0,X1] : ((zex(X0) & ~zex(X1)) => lt(X0,X1))[flattening 71]
71. ! [X0] : ! [X1] : ((zex(X0) & ~zex(X1)) => lt(X0,X1))[rectify 13]
13. ! [X24] : ! [X25] : ((zex(X24) & ~zex(X25)) => lt(X24,X25))[input]
429. ~zex(sK1) | ~zex(sK0) (3:4) [resolution 349,227]
227. ~lt(X0,X0) (0:3) [cnf transformation 124]
124. ! [X0] : ~lt(X0,X0)[ennf transformation 50]
50. ! [X0] : ~lt(X0,X0)[rectify 2]
2. ! [X2] : ~lt(X2,X2)[input]
349. lt(X0,sK1) | ~zex(X0) | ~zex(sK0) (2:7) [superposition 214,332]
332. X0 = X1 | ~zex(X1) | ~zex(X0) (1:7) [resolution 296,303]
303. ~sQ1_eqProxy(X0,X1) | X0 = X1 (0:6) [equality proxy axiom]
296. sQ1_eqProxy(X0,X1) | ~zex(X1) | ~zex(X0) (0:7) [equality proxy replacement 236]
236. X0 = X1 | ~zex(X1) | ~zex(X0) (0:7) [cnf transformation 135]
135. ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)[flattening 134]
134. ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)[ennf transformation 66]
66. ! [X0,X1] : ((zex(X0) & zex(X1)) => X0 = X1)[flattening 65]
65. ! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => X0 = X1)[rectify 27]
27. ! [X53] : ! [X54] : ((zex(X53) & zex(X54)) => X53 = X54)[input]
891. zex(sK1) (5:2) [resolution 888,328]
328. ~zex(intersection(sK0,sK1)) | zex(sK1) (1:6) [resolution 239,215]
215. ~lt(intersection(sK0,sK1),sK1) (0:5) [cnf transformation 168]
888. zex(intersection(sK0,sK1)) (4:4) [subsumption resolution 878,214]
878. ~lt(sK0,sK1) | zex(intersection(sK0,sK1)) (4:7) [resolution 776,241]
241. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 144]
144. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 76]
76. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 75]
75. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 30]
30. ! [X59] : ! [X60] : (~zex(intersection(X59,X60)) => cont(intersection(X59,X60),X59))[input]
776. ~cont(intersection(sK0,sK1),X4) | ~lt(X4,sK1) (3:8) [subsumption resolution 766,689]
689. ~lt(intersection(sK0,sK1),X16) | ~lt(X16,sK1) (1:8) [resolution 287,215]
287. lt(X0,X2) | ~lt(X1,X2) | ~lt(X0,X1) (0:9) [cnf transformation 164]
164. ! [X0,X1,X2] : (~lt(X0,X1) | ~lt(X1,X2) | lt(X0,X2))[flattening 163]
163. ! [X0,X1,X2] : ((~lt(X0,X1) | ~lt(X1,X2)) | lt(X0,X2))[ennf transformation 114]
114. ! [X0,X1,X2] : ((lt(X0,X1) & lt(X1,X2)) => lt(X0,X2))[flattening 113]
113. ! [X0] : ! [X1] : ! [X2] : ((lt(X0,X1) & lt(X1,X2)) => lt(X0,X2))[rectify 4]
4. ! [X5] : ! [X6] : ! [X7] : ((lt(X5,X6) & lt(X6,X7)) => lt(X5,X7))[input]
766. ~lt(X4,sK1) | lt(intersection(sK0,sK1),X4) | ~cont(intersection(sK0,sK1),X4) (3:13) [resolution 709,231]
231. eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 128]
128. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 127]
127. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 56]
56. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 55]
55. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 34]
34. ! [X70] : ! [X71] : (cont(X70,X71) => (lt(X70,X71) | eqdim(X70,X71)))[input]
709. ~eqdim(intersection(sK0,sK1),X6) | ~lt(X6,sK1) (2:8) [resolution 607,233]
233. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 131]
131. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 60]
60. ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))[flattening 59]
59. ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))[rectify 6]
6. ! [X9] : ! [X10] : (eqdim(X9,X10) => eqdim(X10,X9))[input]
607. ~eqdim(X16,intersection(sK0,sK1)) | ~lt(X16,sK1) (1:8) [resolution 281,215]
281. lt(X1,X2) | ~lt(X0,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 152]
152. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X0,X2) | lt(X1,X2))[flattening 151]
151. ! [X0,X1,X2] : (~eqdim(X0,X1) | (~lt(X0,X2) | lt(X1,X2)))[ennf transformation 102]
102. ! [X0,X1,X2] : (eqdim(X0,X1) => (lt(X0,X2) => lt(X1,X2)))[flattening 101]
101. ! [X0] : ! [X1] : ! [X2] : (eqdim(X0,X1) => (lt(X0,X2) => lt(X1,X2)))[rectify 10]
10. ! [X19] : ! [X20] : ! [X21] : (eqdim(X19,X20) => (lt(X19,X21) => lt(X20,X21)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 137
Passive clauses: 362
Generated clauses: 681
Final active clauses: 137
Final passive clauses: 199
Input formulas: 43
Initial clauses: 85

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 29
Fw subsumption resolutions: 48

Simple tautologies: 33
Equational tautologies: 1
Forward subsumptions: 229

Binary resolution: 297
Factoring: 1
Forward superposition: 162
Backward superposition: 60

Unique components: 3

Memory used: 255KB
Time elapsed: 0.055 s
------------------------------
% Success in time 0.636 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 00:58:15 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal10.all.tptp"
Input read from 
============================ end of footer ===========================
