Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_down_theoremsT6-T10_goal01a.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1526
Passive clauses: 8573
Generated clauses: 17592
Final active clauses: 1491
Final passive clauses: 6416
Input formulas: 67
Initial clauses: 111

Selected by SInE selection: 67
SInE iterations: 1

Duplicate literals: 145
Fw subsumption resolutions: 593

Simple tautologies: 58
Forward subsumptions: 8158

Binary resolution: 16408
Factoring: 21

Backtracking splits: 259
Backtracking splits refuted: 31
Backtracking splits refuted at zero level: 7

Memory used: 3198KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_down_theoremsT6-T10_goal01a.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1938
Passive clauses: 17630
Generated clauses: 40026
Final active clauses: 1938
Final passive clauses: 15228
Input formulas: 67
Initial clauses: 110

Selected by SInE selection: 67
SInE iterations: 1

Duplicate literals: 922
Fw subsumption resolutions: 2254

Simple tautologies: 548
Equational tautologies: 16
Forward subsumptions: 19187

Binary resolution: 28934
Factoring: 108
Forward superposition: 7760
Backward superposition: 97

Unique components: 5

Memory used: 6396KB
Time elapsed: 1.404 s
------------------------------
% remaining time: 5981 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on theorems\conversions\codi_down_theoremsT6-T10_goal01a.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 8955
Passive clauses: 31190
Generated clauses: 120989
Final active clauses: 1923
Final passive clauses: 3855
Input formulas: 67
Initial clauses: 108

Duplicate literals: 421
Trivial inequalities: 239
Bw subsumption resolutions: 19699
Fw demodulations: 11199
Condensations: 336

Simple tautologies: 1196
Equational tautologies: 319
Forward subsumptions: 68501

Binary resolution: 47035
Factoring: 4
Forward superposition: 8005
Backward superposition: 8305

Backtracking splits: 1303
Backtracking splits refuted: 714
Backtracking splits refuted at zero level: 114

Memory used: 15351KB
Time elapsed: 3.305 s
------------------------------
% remaining time: 5948 next slice time: 95
dis+2_4_bs=off:ep=on:n=on:nicw=on:nwc=1.5:ptb=off:sac=on:sio=off:spl=backtracking:ssec=off_95 on theorems\conversions\codi_down_theoremsT6-T10_goal01a.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 5104
Passive clauses: 94711
Generated clauses: 247102
Final active clauses: 4155
Final passive clauses: 82719
Input formulas: 67
Initial clauses: 110

Duplicate literals: 5607
Fw subsumption resolutions: 16924
Fw demodulations: 18576
Bw demodulations: 343

Simple tautologies: 1496
Equational tautologies: 20
Forward subsumptions: 113340

Binary resolution: 178589
Factoring: 34
Forward superposition: 18413
Backward superposition: 9014

Backtracking splits: 41
Backtracking splits refuted: 11
Backtracking splits refuted at zero level: 3

Memory used: 44391KB
Time elapsed: 9.614 s
------------------------------
% remaining time: 5851 next slice time: 4
lrs-1003_12_bms=on:bs=unit_only:bsr=on:cond=fast:drc=off:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=10:sfv=off:sgo=on:sio=off:sp=occurrence:spo=on:ssec=off:stl=81:updr=off_4 on theorems\conversions\codi_down_theoremsT6-T10_goal01a.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 969
Passive clauses: 3307
Generated clauses: 9745
Final active clauses: 939
Final passive clauses: 2137
Input formulas: 67
Initial clauses: 110

Duplicate literals: 102
Fw subsumption resolutions: 169
Bw subsumption resolutions: 3
Fw demodulations: 46
Bw demodulations: 280
Condensations: 306

Simple tautologies: 56
Equational tautologies: 27
Forward subsumptions: 5725
Backward subsumptions: 217
Bw demodulations to eq. taut.: 5
Subsumed by BDD marking: 34

Binary resolution: 8005
Factoring: 151
Forward superposition: 303
Backward superposition: 552

Split clauses: 3
Split components: 6
Unique components: 1778
Introduced splitting names: 1
BDD propositional clauses: 1

Memory used: 2814KB
Time elapsed: 0.503 s
------------------------------
% remaining time: 5846 next slice time: 22
lrs+1010_50_bs=off:drc=off:flr=on:lcm=predicate:n=on:nwc=1.2:sfv=off:sgo=on:stl=81_22 on theorems\conversions\codi_down_theoremsT6-T10_goal01a.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 3256
Passive clauses: 8654
Generated clauses: 79498
Final active clauses: 3084
Final passive clauses: 5190
Input formulas: 67
Initial clauses: 108

Duplicate literals: 3934
Fw subsumption resolutions: 2341
Fw demodulations: 189
Bw demodulations: 371
Fw literal rewrites: 3764

Simple tautologies: 1578
Equational tautologies: 48
Forward subsumptions: 58210
Bw demodulations to eq. taut.: 9

Binary resolution: 47852
Factoring: 19
Forward superposition: 7393
Backward superposition: 13683
Equality resolution: 1

Unique components: 5

Memory used: 7164KB
Time elapsed: 2.306 s
------------------------------
% remaining time: 5823 next slice time: 41
dis-1010_5_bd=off:bs=off:cond=fast:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=1.3:ptb=off:sac=on:sgo=on:sio=off:sos=on:spl=backtracking:ssec=off:updr=off_41 on theorems\conversions\codi_down_theoremsT6-T10_goal01a.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 8581
Passive clauses: 27497
Generated clauses: 106408
Final active clauses: 8298
Final passive clauses: 17019
Input formulas: 67
Initial clauses: 110

Duplicate literals: 2407
Fw subsumption resolutions: 5717
Condensations: 157

Simple tautologies: 222
Forward subsumptions: 71847

Binary resolution: 94133
Factoring: 924
Forward superposition: 3239

Backtracking splits: 93
Backtracking splits refuted: 48
Backtracking splits refuted at zero level: 9

Memory used: 21364KB
Time elapsed: 4.207 s
------------------------------
% remaining time: 5780 next slice time: 237
lrs+1_2:1_bs=off:ep=on:gsp=input_only:lcm=predicate:n=on:nicw=on:nwc=1.1:ptb=off:sac=on:sgo=on:sio=off:sos=on:sp=occurrence:spl=backtracking:spo=on:ssec=off:stl=81:updr=off_237 on theorems\conversions\codi_down_theoremsT6-T10_goal01a.all
Refutation found. Thanks to Tanya!
234301. $false (9:0) [subsumption resolution 234296,240]
240. p(sK2,difference(sK0,sK1)) (0:5) [cnf transformation 222]
222. p(sK2,difference(sK0,sK1)) & ~zex(sK2) & ~zex(sK0) & ~zex(intersection(sK0,sK1)) & ~p(sK2,difference(sK0,intersection(sK0,sK1)))[skolemisation 138]
138. ? [X0,X1,X2] : (p(X2,difference(X0,X1)) & ~zex(X2) & ~zex(X0) & ~zex(intersection(X0,X1)) & ~p(X2,difference(X0,intersection(X0,X1))))[flattening 137]
137. ? [X0,X1,X2] : ((p(X2,difference(X0,X1)) & ~zex(X2) & ~zex(X0) & ~zex(intersection(X0,X1))) & ~p(X2,difference(X0,intersection(X0,X1))))[ennf transformation 68]
68. ~! [X0,X1,X2] : (~(p(X2,difference(X0,X1)) & ~zex(X2) & ~zex(X0) & ~zex(intersection(X0,X1))) | p(X2,difference(X0,intersection(X0,X1))))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : (~(p(X2,difference(X0,X1)) & ~zex(X2) & ~zex(X0) & ~zex(intersection(X0,X1))) | p(X2,difference(X0,intersection(X0,X1))))[negated conjecture]
234296. ~p(sK2,difference(sK0,sK1)) (9:5) [resolution 233908,336]
336. ~po(X2,intersection(X0,X1)) | ~p(X2,difference(X0,X1)) (0:10) [cnf transformation 201]
201. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | ~po(X2,intersection(X0,X1)))[ennf transformation 127]
127. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | ~po(X2,intersection(X0,X1)))[flattening 66]
66. ! [X0] : ! [X1] : ! [X2] : (~p(X2,difference(X0,X1)) | ~po(X2,intersection(X0,X1)))[input]
233908. po(sK2,intersection(sK0,sK1)) (8:5) [subsumption resolution 233467,244]
244. ~p(sK2,difference(sK0,intersection(sK0,sK1))) (0:7) [cnf transformation 222]
233467. po(sK2,intersection(sK0,sK1)) | p(sK2,difference(sK0,intersection(sK0,sK1))) (8:12) [superposition 364,69809]
69809. intersection(sK0,sK1) = intersection(sK0,intersection(sK0,sK1)) (7:9) [resolution 11175,354]
354. ~sQ1_eqProxy(X0,X1) | X0 = X1 (0:6) [equality proxy axiom]
11175. sQ1_eqProxy(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) (6:9) [resolution 3303,575]
575. ~cont(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) | sQ1_eqProxy(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) (3:18) [resolution 413,352]
352. ~cont(X1,X0) | sQ1_eqProxy(X0,X1) | ~cont(X0,X1) (0:9) [equality proxy replacement 325]
325. X0 = X1 | ~cont(X1,X0) | ~cont(X0,X1) (0:9) [cnf transformation 191]
191. ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1)[flattening 190]
190. ! [X0,X1] : ((~cont(X0,X1) | ~cont(X1,X0)) | X0 = X1)[ennf transformation 119]
119. ! [X0,X1] : (~(cont(X0,X1) & cont(X1,X0)) | X0 = X1)[flattening 20]
20. ! [X0] : ! [X1] : (~(cont(X0,X1) & cont(X1,X0)) | X0 = X1)[input]
413. cont(intersection(sK0,sK1),intersection(sK0,intersection(sK0,sK1))) (2:9) [resolution 355,296]
296. ~p(X0,X1) | cont(X0,X1) (0:6) [cnf transformation 168]
168. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 167]
167. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 95]
95. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 3]
3. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
355. p(intersection(sK0,sK1),intersection(sK0,intersection(sK0,sK1))) (1:9) [resolution 243,308]
308. zex(intersection(X0,X1)) | p(intersection(X0,X1),intersection(X0,intersection(X0,X1))) (0:13) [cnf transformation 103]
103. ! [X0,X1] : (zex(intersection(X0,X1)) | p(intersection(X0,X1),intersection(X0,intersection(X0,X1))))[flattening 51]
51. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | p(intersection(X0,X1),intersection(X0,intersection(X0,X1))))[input]
243. ~zex(intersection(sK0,sK1)) (0:4) [cnf transformation 222]
3303. cont(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) (5:9) [resolution 719,296]
719. p(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) (4:9) [subsumption resolution 715,243]
715. p(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) | zex(intersection(sK0,sK1)) (4:13) [resolution 570,327]
327. zex(intersection(X0,intersection(X0,X1))) | p(intersection(X0,intersection(X0,X1)),intersection(X0,X1)) | zex(intersection(X0,X1)) (0:19) [cnf transformation 195]
195. ! [X0,X1] : (zex(intersection(X0,X1)) | zex(intersection(X0,intersection(X0,X1))) | p(intersection(X0,intersection(X0,X1)),intersection(X0,X1)))[flattening 194]
194. ! [X0,X1] : ((zex(intersection(X0,X1)) | zex(intersection(X0,intersection(X0,X1)))) | p(intersection(X0,intersection(X0,X1)),intersection(X0,X1)))[ennf transformation 121]
121. ! [X0,X1] : (~(~zex(intersection(X0,X1)) & ~zex(intersection(X0,intersection(X0,X1)))) | p(intersection(X0,intersection(X0,X1)),intersection(X0,X1)))[flattening 54]
54. ! [X0] : ! [X1] : (~(~zex(intersection(X0,X1)) & ~zex(intersection(X0,intersection(X0,X1)))) | p(intersection(X0,intersection(X0,X1)),intersection(X0,X1)))[input]
570. ~zex(intersection(sK0,intersection(sK0,sK1))) (3:6) [resolution 413,311]
311. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 172]
172. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 106]
106. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[flattening 22]
22. ! [X0] : ! [X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[input]
364. po(sK2,intersection(sK0,X0)) | p(sK2,difference(sK0,X0)) (2:10) [resolution 358,342]
342. ~p(X2,X0) | po(X2,intersection(X0,X1)) | p(X2,difference(X0,X1)) (0:13) [cnf transformation 213]
213. ! [X0,X1,X2] : (~p(X2,X0) | po(X2,intersection(X0,X1)) | p(X2,difference(X0,X1)))[flattening 212]
212. ! [X0,X1,X2] : ((~p(X2,X0) | po(X2,intersection(X0,X1))) | p(X2,difference(X0,X1)))[ennf transformation 133]
133. ! [X0,X1,X2] : (~(p(X2,X0) & ~po(X2,intersection(X0,X1))) | p(X2,difference(X0,X1)))[flattening 67]
67. ! [X0] : ! [X1] : ! [X2] : (~(p(X2,X0) & ~po(X2,intersection(X0,X1))) | p(X2,difference(X0,X1)))[input]
358. p(sK2,sK0) (1:3) [resolution 240,335]
335. ~p(X2,difference(X0,X1)) | p(X2,X0) (0:8) [cnf transformation 200]
200. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))[ennf transformation 126]
126. ! [X0,X1,X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))[flattening 65]
65. ! [X0] : ! [X1] : ! [X2] : (~p(X2,difference(X0,X1)) | p(X2,X0))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 10456
Passive clauses: 74254
Generated clauses: 250531
Final active clauses: 6360
Final passive clauses: 36355
Input formulas: 67
Initial clauses: 110

Duplicate literals: 2939
Fw subsumption resolutions: 6284
Fw demodulations: 45079
Bw demodulations: 26135

Simple tautologies: 1337
Forward subsumptions: 120316
Bw demodulations to eq. taut.: 12

Binary resolution: 143803
Factoring: 19
Forward superposition: 6245
Backward superposition: 2853

Backtracking splits: 547
Backtracking splits refuted: 40
Backtracking splits refuted at zero level: 24

Memory used: 49764KB
Time elapsed: 11.771 s
------------------------------
% Success in time 33.724 s
