Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal22.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 970
Passive clauses: 5226
Generated clauses: 13326
Final active clauses: 954
Final passive clauses: 3901
Input formulas: 50
Initial clauses: 86

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 83
Fw subsumption resolutions: 573

Simple tautologies: 66
Forward subsumptions: 7498

Binary resolution: 12359
Factoring: 28

Backtracking splits: 170
Backtracking splits refuted: 12

Memory used: 2174KB
Time elapsed: 0.404 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal22.all
Refutation found. Thanks to Tanya!
561. $false (4:0) [subsumption resolution 560,420]
420. c(sK1,sK0) (3:3) [resolution 419,187]
187. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 102]
102. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[ennf transformation 54]
54. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[flattening 11]
11. ! [X0] : ! [X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[input]
419. ~zex(intersection(sK1,sK0)) (2:4) [resolution 413,304]
304. c(intersection(sK1,sK0),sK0) (1:5) [resolution 187,173]
173. ~zex(intersection(intersection(sK1,sK0),sK0)) (0:6) [cnf transformation 162]
162. zex(intersection(sK0,sK1)) & ~zex(intersection(intersection(sK1,sK0),sK0))[skolemisation 95]
95. ? [X0,X1] : (zex(intersection(X0,X1)) & ~zex(intersection(intersection(X1,X0),X0)))[ennf transformation 51]
51. ~! [X0,X1] : (~zex(intersection(X0,X1)) | zex(intersection(intersection(X1,X0),X0)))[flattening 1]
1. ~! [X0] : ! [X1] : (~zex(intersection(X0,X1)) | zex(intersection(intersection(X1,X0),X0)))[negated conjecture]
413. ~c(X2,X3) | ~zex(X2) (1:5) [resolution 198,213]
213. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 116]
116. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 67]
67. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[flattening 31]
31. ! [X0] : ! [X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[input]
198. cont(sK5(X1,X0),X0) | ~c(X0,X1) (0:8) [cnf transformation 166]
166. ! [X0,X1] : ((~c(X0,X1) | (cont(sK5(X1,X0),X0) & cont(sK5(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 108]
108. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 60]
60. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 59]
59. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 34]
34. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
560. ~c(sK1,sK0) (4:3) [duplicate literal removal 552]
552. ~c(sK1,sK0) | ~c(sK1,sK0) (4:6) [resolution 478,199]
199. cont(sK5(X1,X0),X1) | ~c(X0,X1) (0:8) [cnf transformation 166]
478. ~cont(sK5(X1,sK1),sK0) | ~c(sK1,X1) (3:8) [resolution 471,198]
471. ~cont(X0,sK1) | ~cont(X0,sK0) (2:6) [resolution 200,305]
305. ~c(sK0,sK1) (1:3) [resolution 188,172]
172. zex(intersection(sK0,sK1)) (0:4) [cnf transformation 162]
188. ~zex(intersection(X0,X1)) | ~c(X0,X1) (0:7) [cnf transformation 102]
200. c(X0,X1) | ~cont(X3,X1) | ~cont(X3,X0) (0:9) [cnf transformation 166]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 98
Passive clauses: 196
Generated clauses: 377
Final active clauses: 98
Final passive clauses: 75
Input formulas: 50
Initial clauses: 85

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 21
Fw subsumption resolutions: 22

Simple tautologies: 30
Equational tautologies: 1
Forward subsumptions: 123

Binary resolution: 127
Factoring: 1
Forward superposition: 76
Backward superposition: 46

Unique components: 2

Memory used: 127KB
Time elapsed: 0.014 s
------------------------------
% Success in time 0.485 s
