% ott+11_2:1_add=large:afp=40000:afq=2.0:amm=sco:anc=none:br=off:cond=on:irw=on:nwc=1:sd=2:ss=axioms:st=2.0:sos=all:urr=on:updr=off_3 on codi_down_sum_theorems_goal05.all
% Refutation found. Thanks to Tanya!
% SZS status Theorem for codi_down_sum_theorems_goal05.all
% SZS output start Proof for codi_down_sum_theorems_goal05.all
fof(f1283,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f757,f1063,f1282])).
fof(f1282,plain,(
  ~spl7_4),
  inference(avatar_contradiction_clause,[],[f1281])).
fof(f1281,plain,(
  $false | ~spl7_4),
  inference(subsumption_resolution,[],[f1280,f272])).
fof(f272,plain,(
  ~zex(sK0)),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ~zex(sK0) & zex(sum(sK0,sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f188,f259])).
fof(f259,plain,(
  ? [X0,X1] : (~zex(X0) & zex(sum(X0,X1))) => (~zex(sK0) & zex(sum(sK0,sK1)))),
  introduced(choice_axiom,[])).
fof(f188,plain,(
  ? [X0,X1] : (~zex(X0) & zex(sum(X0,X1)))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ~! [X0,X1] : (zex(X0) | ~zex(sum(X0,X1)))),
  inference(flattening,[],[f2])).
fof(f2,negated_conjecture,(
  ~! [X0] : ! [X1] : (zex(X0) | ~zex(sum(X0,X1)))),
  inference(negated_conjecture,[],[f1])).
fof(f1,conjecture,(
  ! [X0] : ! [X1] : (zex(X0) | ~zex(sum(X0,X1)))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal05.all.tptp',conjecture50)).
fof(f1280,plain,(
  zex(sK0) | ~spl7_4),
  inference(forward_demodulation,[],[f1075,f1146])).
fof(f1146,plain,(
  sK0 = sK1 | ~spl7_4),
  inference(backward_demodulation,[],[f1138,f1102])).
fof(f1102,plain,(
  sK1 = sum(sK0,sK1) | ~spl7_4),
  inference(backward_demodulation,[],[f1098,f418])).
fof(f418,plain,(
  sum(sK0,sK1) = sK2),
  inference(unit_resulting_resolution,[],[f273,f271,f360])).
fof(f360,plain,(
  ( ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1] : (X0 = X1 | ~zex(X1) | ~zex(X0))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0,X1] : (X0 = X1 | (~zex(X1) | ~zex(X0)))),
  inference(ennf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : (X0 = X1 | ~(zex(X1) & zex(X0)))),
  inference(flattening,[],[f89])).
fof(f89,axiom,(
  ! [X0] : ! [X1] : (X0 = X1 | ~(zex(X1) & zex(X0)))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal05.all.tptp',axiom89550)).
fof(f271,plain,(
  zex(sum(sK0,sK1))),
  inference(cnf_transformation,[],[f260])).
fof(f273,plain,(
  zex(sK2)),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  zex(sK2)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f107,f261])).
fof(f261,plain,(
  ? [X0] : zex(X0) => zex(sK2)),
  introduced(choice_axiom,[])).
fof(f107,axiom,(
  ? [X0] : zex(X0)),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal05.all.tptp',axiom87750)).
fof(f1098,plain,(
  sK1 = sK2 | ~spl7_4),
  inference(unit_resulting_resolution,[],[f441,f756,f333])).
fof(f333,plain,(
  ( ! [X0,X1] : (X0 = X1 | cont(sK4(X0,X1),X1) | cont(sK4(X0,X1),X0)) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1] : (X0 = X1 | ((~cont(sK4(X0,X1),X0) & cont(sK4(X0,X1),X1)) | (~cont(sK4(X0,X1),X1) & cont(sK4(X0,X1),X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f237,f265])).
fof(f265,plain,(
  ! [X1,X0] : (? [X2] : ((~cont(X2,X0) & cont(X2,X1)) | (~cont(X2,X1) & cont(X2,X0))) => ((~cont(sK4(X0,X1),X0) & cont(sK4(X0,X1),X1)) | (~cont(sK4(X0,X1),X1) & cont(sK4(X0,X1),X0))))),
  introduced(choice_axiom,[])).
fof(f237,plain,(
  ! [X0,X1] : (X0 = X1 | ? [X2] : ((~cont(X2,X0) & cont(X2,X1)) | (~cont(X2,X1) & cont(X2,X0))))),
  inference(ennf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1] : (X0 = X1 | ~! [X2] : ((cont(X2,X0) | ~cont(X2,X1)) & (cont(X2,X1) | ~cont(X2,X0))))),
  inference(flattening,[],[f90])).
fof(f90,axiom,(
  ! [X0] : ! [X1] : (X0 = X1 | ~! [X2] : ((cont(X2,X0) | ~cont(X2,X1)) & (cont(X2,X1) | ~cont(X2,X0))))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal05.all.tptp',axiom89560)).
fof(f756,plain,(
  ( ! [X0] : (~cont(X0,sK1)) ) | ~spl7_4),
  inference(avatar_component_clause,[],[f755])).
fof(f755,plain,(
  spl7_4 <=> ! [X0] : ~cont(X0,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl7_4])])).
fof(f441,plain,(
  ( ! [X0] : (~cont(X0,sK2)) )),
  inference(forward_demodulation,[],[f399,f418])).
fof(f399,plain,(
  ( ! [X0] : (~cont(X0,sum(sK0,sK1))) )),
  inference(unit_resulting_resolution,[],[f271,f340])).
fof(f340,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  inference(flattening,[],[f88])).
fof(f88,axiom,(
  ! [X0] : ! [X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal05.all.tptp',axiom89540)).
fof(f1138,plain,(
  sK0 = sum(sK0,sK1) | ~spl7_4),
  inference(backward_demodulation,[],[f1098,f990])).
fof(f990,plain,(
  sK0 = sum(sK0,sK2)),
  inference(superposition,[],[f472,f299])).
fof(f299,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = sum(X1,X0)) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : sum(X0,X1) = sum(X1,X0)),
  inference(flattening,[],[f3])).
fof(f3,axiom,(
  ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal05.all.tptp',axiom10)).
fof(f472,plain,(
  sK0 = sum(sK2,sK0)),
  inference(unit_resulting_resolution,[],[f387,f298])).
fof(f298,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = X1 | ~lt(X0,X1)) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : (sum(X0,X1) = X1 | ~lt(X0,X1))),
  inference(flattening,[],[f4])).
fof(f4,axiom,(
  ! [X0] : ! [X1] : (sum(X0,X1) = X1 | ~lt(X0,X1))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal05.all.tptp',axiom20)).
fof(f387,plain,(
  lt(sK2,sK0)),
  inference(unit_resulting_resolution,[],[f273,f272,f372])).
fof(f372,plain,(
  ( ! [X0,X1] : (lt(X0,X1) | zex(X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1] : (lt(X0,X1) | zex(X1) | ~zex(X0))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X0,X1] : (lt(X0,X1) | (zex(X1) | ~zex(X0)))),
  inference(ennf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : (lt(X0,X1) | ~(~zex(X1) & zex(X0)))),
  inference(flattening,[],[f84])).
fof(f84,axiom,(
  ! [X0] : ! [X1] : (lt(X0,X1) | ~(~zex(X1) & zex(X0)))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal05.all.tptp',axiom89500)).
fof(f1075,plain,(
  zex(sK1) | ~spl7_4),
  inference(unit_resulting_resolution,[],[f756,f342])).
fof(f342,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,axiom,(
  ! [X0] : ((~zex(X0) | ~cont(X0,X0)) & (cont(X0,X0) | zex(X0)))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal05.all.tptp',axiom89510)).
fof(f1063,plain,(
  spl7_3),
  inference(avatar_contradiction_clause,[],[f1062])).
fof(f1062,plain,(
  $false | spl7_3),
  inference(subsumption_resolution,[],[f1061,f441])).
fof(f1061,plain,(
  cont(sK0,sK2) | spl7_3),
  inference(forward_demodulation,[],[f1060,f418])).
fof(f1060,plain,(
  cont(sK0,sum(sK0,sK1)) | spl7_3),
  inference(forward_demodulation,[],[f1048,f299])).
fof(f1048,plain,(
  cont(sK0,sum(sK1,sK0)) | spl7_3),
  inference(unit_resulting_resolution,[],[f385,f889,f297])).
fof(f297,plain,(
  ( ! [X2,X0,X1] : (cont(X2,sum(X0,X1)) | ~cont(X2,X1) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | ~cont(X2,X1) | ~leq(X0,X1))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | (~cont(X2,X1) | ~leq(X0,X1)))),
  inference(ennf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | ~(cont(X2,X1) & leq(X0,X1)))),
  inference(flattening,[],[f5])).
fof(f5,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (cont(X2,sum(X0,X1)) | ~(cont(X2,X1) & leq(X0,X1)))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal05.all.tptp',axiom30)).
fof(f889,plain,(
  leq(sK1,sK0) | spl7_3),
  inference(unit_resulting_resolution,[],[f753,f345])).
fof(f345,plain,(
  ( ! [X0,X1] : (leq(X1,X0) | leq(X0,X1)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : (leq(X1,X0) | leq(X0,X1))),
  inference(flattening,[],[f96])).
fof(f96,axiom,(
  ! [X0] : ! [X1] : (leq(X1,X0) | leq(X0,X1))),
  file('Dropbox/Shirly_PhD_Work/codi_down_sum_theorems_rerun/codi_down_sum_theorems_goal05.all.tptp',axiom89110)).
fof(f753,plain,(
  ~leq(sK0,sK1) | spl7_3),
  inference(avatar_component_clause,[],[f751])).
fof(f751,plain,(
  spl7_3 <=> leq(sK0,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl7_3])])).
fof(f385,plain,(
  cont(sK0,sK0)),
  inference(unit_resulting_resolution,[],[f272,f342])).
fof(f757,plain,(
  ~spl7_3 | spl7_4),
  inference(avatar_split_clause,[],[f749,f755,f751])).
fof(f749,plain,(
  ( ! [X0] : (~cont(X0,sK1) | ~leq(sK0,sK1)) )),
  inference(subsumption_resolution,[],[f747,f441])).
fof(f747,plain,(
  ( ! [X0] : (cont(X0,sK2) | ~cont(X0,sK1) | ~leq(sK0,sK1)) )),
  inference(superposition,[],[f297,f418])).
% SZS output end Proof for codi_down_sum_theorems_goal05.all
% ------------------------------
% Version: Vampire 4.4.0 (commit f9be976 on 2020-01-20 19:56:10 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5756
% Time elapsed: 0.078 s
% ------------------------------
% ------------------------------
% Success in time 0.098 s
