Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal09.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 747
Passive clauses: 4272
Generated clauses: 10581
Final active clauses: 686
Final passive clauses: 2656
Input formulas: 50
Initial clauses: 99

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 145
Fw subsumption resolutions: 633

Simple tautologies: 72
Forward subsumptions: 5411

Binary resolution: 9417
Factoring: 50

Backtracking splits: 171
Backtracking splits refuted: 66
Backtracking splits refuted at zero level: 25

Memory used: 1918KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal09.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1686
Passive clauses: 8973
Generated clauses: 24547
Final active clauses: 1686
Final passive clauses: 6871
Input formulas: 50
Initial clauses: 98

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 309
Fw subsumption resolutions: 1127

Simple tautologies: 293
Equational tautologies: 5
Forward subsumptions: 14264

Binary resolution: 17425
Factoring: 159
Forward superposition: 5374
Backward superposition: 76

Unique components: 3

Memory used: 3070KB
Time elapsed: 1.403 s
------------------------------
% remaining time: 5981 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on theorems\conversions\codi_int_theorems_goal09.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 6258
Passive clauses: 14154
Generated clauses: 54853
Final active clauses: 5215
Final passive clauses: 5940
Input formulas: 50
Initial clauses: 96

Duplicate literals: 397
Trivial inequalities: 3
Bw subsumption resolutions: 471
Fw demodulations: 3440
Condensations: 39

Simple tautologies: 152
Equational tautologies: 435
Forward subsumptions: 34751

Binary resolution: 46526
Forward superposition: 502
Backward superposition: 1097

Backtracking splits: 1034
Backtracking splits refuted: 323
Backtracking splits refuted at zero level: 23

Memory used: 11129KB
Time elapsed: 3.310 s
------------------------------
% remaining time: 5947 next slice time: 95
dis+2_4_bs=off:ep=on:n=on:nicw=on:nwc=1.5:ptb=off:sac=on:sio=off:spl=backtracking:ssec=off_95 on theorems\conversions\codi_int_theorems_goal09.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 2670
Passive clauses: 44848
Generated clauses: 115364
Final active clauses: 2517
Final passive clauses: 40346
Input formulas: 50
Initial clauses: 98

Duplicate literals: 5274
Fw subsumption resolutions: 6141
Fw demodulations: 3839
Bw demodulations: 14

Simple tautologies: 1023
Equational tautologies: 18
Forward subsumptions: 56211

Binary resolution: 94599
Factoring: 75
Forward superposition: 3134
Backward superposition: 2568

Backtracking splits: 79
Backtracking splits refuted: 31
Backtracking splits refuted at zero level: 24

Memory used: 20468KB
Time elapsed: 9.612 s
------------------------------
% remaining time: 5850 next slice time: 4
lrs-1003_12_bms=on:bs=unit_only:bsr=on:cond=fast:drc=off:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=10:sfv=off:sgo=on:sio=off:sp=occurrence:spo=on:ssec=off:stl=81:updr=off_4 on theorems\conversions\codi_int_theorems_goal09.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 678
Passive clauses: 1416
Generated clauses: 4246
Final active clauses: 656
Final passive clauses: 596
Input formulas: 50
Initial clauses: 98
Reactivated clauses: 7

Duplicate literals: 64
Fw subsumption resolutions: 92
Bw subsumption resolutions: 79
Fw demodulations: 47
Bw demodulations: 449
Condensations: 18

Simple tautologies: 36
Equational tautologies: 73
Forward subsumptions: 2274
Backward subsumptions: 333
Bw demodulations to eq. taut.: 12
Subsumed empty clauses: 4
Subsumed by BDD marking: 159

Binary resolution: 3452
Factoring: 105
Forward superposition: 30
Backward superposition: 214

Split clauses: 22
Split components: 44
Unique components: 719
Introduced splitting names: 6
BDD propositional clauses: 8

Memory used: 1407KB
Time elapsed: 0.502 s
------------------------------
% remaining time: 5845 next slice time: 22
lrs+1010_50_bs=off:drc=off:flr=on:lcm=predicate:n=on:nwc=1.2:sfv=off:sgo=on:stl=81_22 on theorems\conversions\codi_int_theorems_goal09.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1400
Passive clauses: 5232
Generated clauses: 23971
Final active clauses: 1400
Final passive clauses: 3832
Input formulas: 50
Initial clauses: 96

Duplicate literals: 470
Fw subsumption resolutions: 509
Fw literal rewrites: 411

Simple tautologies: 216
Equational tautologies: 3
Forward subsumptions: 17061

Binary resolution: 20614
Factoring: 40
Forward superposition: 412
Backward superposition: 1431
Self superposition: 1

Unique components: 3

Memory used: 5373KB
Time elapsed: 2.309 s
------------------------------
% remaining time: 5821 next slice time: 41
dis-1010_5_bd=off:bs=off:cond=fast:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=1.3:ptb=off:sac=on:sgo=on:sio=off:sos=on:spl=backtracking:ssec=off:updr=off_41 on theorems\conversions\codi_int_theorems_goal09.all
Refutation found. Thanks to Tanya!
40394. $false (6:0) [subsumption resolution 40391,182]
182. inc(sK0,sK1) (0:3) [cnf transformation 166]
166. inc(sK0,sK1) & lt(sK0,sK1) & ~eqdim(intersection(sK0,sK1),sK0)[skolemisation 99]
99. ? [X0,X1] : (inc(X0,X1) & lt(X0,X1) & ~eqdim(intersection(X0,X1),X0))[flattening 98]
98. ? [X0,X1] : ((inc(X0,X1) & lt(X0,X1)) & ~eqdim(intersection(X0,X1),X0))[ennf transformation 51]
51. ~! [X0,X1] : (~(inc(X0,X1) & lt(X0,X1)) | eqdim(intersection(X0,X1),X0))[flattening 1]
1. ~! [X0] : ! [X1] : (~(inc(X0,X1) & lt(X0,X1)) | eqdim(intersection(X0,X1),X0))[negated conjecture]
40391. ~inc(sK0,sK1) (6:3) [resolution 40387,244]
244. inc(X1,X0) | ~inc(X0,X1) (0:6) [cnf transformation 126]
126. ! [X0,X1] : (~inc(X0,X1) | inc(X1,X0))[ennf transformation 74]
74. ! [X0,X1] : (~inc(X0,X1) | inc(X1,X0))[flattening 3]
3. ! [X0] : ! [X1] : (~inc(X0,X1) | inc(X1,X0))[input]
40387. ~inc(sK1,sK0) (5:3) [backtracking split refutation 2648,2961,40386]
40386. $false {19} (7:0) [subsumption resolution 40385,2961]
40385. ~sP0(sK0,sK1) (7:3) [duplicate literal removal 40383]
40383. ~sP0(sK0,sK1) | ~sP0(sK0,sK1) (7:6) [resolution 40382,224]
224. p(sK9(X1,X0),X0) | ~sP0(X0,X1) (0:8) [cnf transformation 175]
175. ! [X0,X1] : (~sP0(X0,X1) | (lt(sK9(X1,X0),X1) & cont(sK9(X1,X0),X1) & p(sK9(X1,X0),X0)))[skolemisation 174]
174. ! [X0,X1] : (~sP0(X0,X1) | ? [X2] : (lt(X2,X1) & cont(X2,X1) & p(X2,X0)))[rectify 173]
173. ! [X1,X0] : (~sP0(X1,X0) | ? [X2] : (lt(X2,X0) & cont(X2,X0) & p(X2,X1)))[nnf transformation 164]
164. ! [X1,X0] : (~sP0(X1,X0) | ? [X2] : (lt(X2,X0) & cont(X2,X0) & p(X2,X1)))[predicate definition introduction]
40382. ~p(sK9(sK1,X2),sK0) | ~sP0(X2,sK1) (6:8) [subsumption resolution 40381,231]
231. eqdim(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 119]
119. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 118]
118. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 67]
67. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 35]
35. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
40381. ~eqdim(sK9(sK1,X2),sK0) | ~sP0(X2,sK1) | ~p(sK9(sK1,X2),sK0) (6:13) [resolution 4041,230]
230. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 119]
4041. ~cont(sK9(sK1,X5),sK0) | ~eqdim(sK9(sK1,X5),sK0) | ~sP0(X5,sK1) (5:13) [resolution 3273,223]
223. cont(sK9(X1,X0),X1) | ~sP0(X0,X1) (0:8) [cnf transformation 175]
3273. ~cont(X6,sK1) | ~cont(X6,sK0) | ~eqdim(X6,sK0) (4:9) [resolution 2763,234]
234. leq(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 121]
121. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 120]
120. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[ennf transformation 68]
68. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[flattening 49]
49. ! [X0] : ! [X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[input]
2763. ~leq(sK0,X0) | ~cont(X0,sK1) | ~cont(X0,sK0) (3:9) [resolution 2734,272]
272. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 160]
160. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 159]
159. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 96]
96. ! [X0,X1,X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[flattening 10]
10. ! [X0] : ! [X1] : ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[input]
2734. ~leq(X0,intersection(sK0,sK1)) | ~leq(sK0,X0) (2:8) [resolution 2720,270]
270. leq(X0,X2) | ~leq(X1,X2) | ~leq(X0,X1) (0:9) [cnf transformation 156]
156. ! [X0,X1,X2] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2))[flattening 155]
155. ! [X0,X1,X2] : ((~leq(X0,X1) | ~leq(X1,X2)) | leq(X0,X2))[ennf transformation 94]
94. ! [X0,X1,X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[flattening 40]
40. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & leq(X1,X2)) | leq(X0,X2))[input]
2720. ~leq(sK0,intersection(sK0,sK1)) (1:5) [backtracking split refutation 302,2691,2718]
2718. $false {16} (2:0) [subsumption resolution 2712,2657]
2657. lt(intersection(sK0,sK1),sK0) (1:5) [backtracking split refutation 299,373,2656]
2656. $false {2} (6:0) [subsumption resolution 2653,182]
2653. ~inc(sK0,sK1) {2} (6:3) [resolution 2649,244]
2649. ~inc(sK1,sK0) {2} (5:3) [subsumption resolution 2648,2611]
2611. ~sP0(sK0,sK1) {2} (7:3) [duplicate literal removal 2609]
2609. ~sP0(sK0,sK1) | ~sP0(sK0,sK1) {2} (7:6) [resolution 1909,224]
1909. ~p(sK9(sK1,X2),sK0) | ~sP0(X2,sK1) {2} (6:8) [resolution 391,230]
391. ~cont(sK9(sK1,X5),sK0) | ~sP0(X5,sK1) {2} (5:8) [resolution 378,223]
378. ~cont(X0,sK1) | ~cont(X0,sK0) {2} (4:6) [resolution 377,221]
221. c(X0,X1) | ~cont(X3,X1) | ~cont(X3,X0) (0:9) [cnf transformation 172]
172. ! [X0,X1] : ((~c(X0,X1) | (cont(sK8(X1,X0),X0) & cont(sK8(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 116]
116. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 64]
64. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 63]
63. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 34]
34. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
377. ~c(sK0,sK1) {2} (3:3) [resolution 374,199]
199. ~zex(intersection(X0,X1)) | ~c(X0,X1) (0:7) [cnf transformation 107]
107. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[ennf transformation 54]
54. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[flattening 8]
8. ! [X0] : ! [X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[input]
374. zex(intersection(sK0,sK1)) {2} (2:4) [resolution 373,240]
240. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 71]
71. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 9]
9. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
373. ~cont(intersection(sK0,sK1),sK0) {2} (1:5) [backtracking splitting component]
299. lt(intersection(sK0,sK1),sK0) | ~cont(intersection(sK0,sK1),sK0) (1:10) [resolution 184,249]
249. eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 131]
131. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[ennf transformation 79]
79. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 36]
36. ! [X0] : ! [X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[input]
184. ~eqdim(intersection(sK0,sK1),sK0) (0:5) [cnf transformation 166]
2712. ~lt(intersection(sK0,sK1),sK0) {16} (2:5) [resolution 2691,236]
236. leq(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 123]
123. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 122]
122. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[ennf transformation 69]
69. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[flattening 46]
46. ! [X0] : ! [X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[input]
2691. ~leq(intersection(sK0,sK1),sK0) {16} (1:5) [backtracking splitting component]
302. ~leq(sK0,intersection(sK0,sK1)) | ~leq(intersection(sK0,sK1),sK0) (1:10) [resolution 184,235]
235. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 121]
2961. sP0(sK0,sK1) {19} (5:3) [backtracking splitting component]
2648. ~inc(sK1,sK0) | sP0(sK0,sK1) (5:6) [duplicate literal removal 2635]
2635. sP0(sK0,sK1) | ~inc(sK1,sK0) | sP0(sK0,sK1) | ~inc(sK1,sK0) (5:12) [resolution 372,225]
225. lt(sK10(X1,X0),X1) | sP0(X1,X0) | ~inc(X0,X1) (0:11) [cnf transformation 177]
177. ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | (lt(sK10(X1,X0),X1) & cont(sK10(X1,X0),X1) & p(sK10(X1,X0),X0))) & ((! [X3] : (~lt(X3,X0) | ~cont(X3,X0) | ~p(X3,X1)) & ! [X4] : (~lt(X4,X1) | ~cont(X4,X1) | ~p(X4,X0))) | inc(X0,X1)))[skolemisation 176]
176. ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | ? [X2] : (lt(X2,X1) & cont(X2,X1) & p(X2,X0))) & ((! [X3] : (~lt(X3,X0) | ~cont(X3,X0) | ~p(X3,X1)) & ! [X4] : (~lt(X4,X1) | ~cont(X4,X1) | ~p(X4,X0))) | inc(X0,X1)))[rectify 165]
165. ! [X0,X1] : ((~inc(X0,X1) | sP0(X1,X0) | ? [X3] : (lt(X3,X1) & cont(X3,X1) & p(X3,X0))) & ((! [X4] : (~lt(X4,X0) | ~cont(X4,X0) | ~p(X4,X1)) & ! [X5] : (~lt(X5,X1) | ~cont(X5,X1) | ~p(X5,X0))) | inc(X0,X1)))[definition folding 164]
372. ~lt(sK10(X4,sK1),sK0) | sP0(X4,sK1) | ~inc(sK1,X4) (4:11) [resolution 327,227]
227. p(sK10(X1,X0),X0) | sP0(X1,X0) | ~inc(X0,X1) (0:11) [cnf transformation 177]
327. ~p(X6,sK1) | ~lt(X6,sK0) (3:6) [resolution 295,231]
295. ~eqdim(X3,sK1) | ~lt(X3,sK0) (2:6) [resolution 288,267]
267. lt(X1,X2) | ~lt(X0,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 150]
150. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X0,X2) | lt(X1,X2))[ennf transformation 91]
91. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X0,X2) | lt(X1,X2))[flattening 24]
24. ! [X0] : ! [X1] : ! [X2] : (~eqdim(X0,X1) | ~lt(X0,X2) | lt(X1,X2))[input]
288. ~lt(sK1,sK0) (1:3) [resolution 183,251]
251. ~lt(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 133]
133. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[ennf transformation 81]
81. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[flattening 17]
17. ! [X0] : ! [X1] : (~lt(X0,X1) | ~lt(X1,X0))[input]
183. lt(sK0,sK1) (0:3) [cnf transformation 166]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 2874
Passive clauses: 10292
Generated clauses: 40160
Final active clauses: 2680
Final passive clauses: 6300
Input formulas: 50
Initial clauses: 98

Duplicate literals: 556
Fw subsumption resolutions: 918
Condensations: 93

Simple tautologies: 86
Forward subsumptions: 28982

Binary resolution: 38188
Factoring: 316

Backtracking splits: 42
Backtracking splits refuted: 8
Backtracking splits refuted at zero level: 4

Memory used: 5500KB
Time elapsed: 2.675 s
------------------------------
% Success in time 20.564 s
