% Starting sequential portfolio solving with schedule "casc"
% spawned child 10160 with time: 4 (total remaining time 2999)
% WARNING: value z3 for option sas not known
% dis+1010_4_afp=10000:afq=1.2:anc=none:irw=on:lma=on:nm=64:nwc=10:sas=z3:sac=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6780
% Time elapsed: 0.500 s
% spawned child 10161 with time: 10 (total remaining time 2994)
% ott+11_8_amm=off:anc=none:bsr=on:cond=on:irw=on:nm=2:nwc=1.1:ss=axioms:st=2.0:sac=on_8 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8059
% Time elapsed: 1.100 s
% ------------------------------
% ------------------------------
% spawned child 10162 with time: 4 (total remaining time 2982)
% dis+1011_10_aac=none:add=large:afp=10000:afq=1.1:amm=off:anc=none:fsr=off:gs=on:gsem=off:irw=on:nm=6:newcnf=on:nwc=2.5:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 12409
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 10164 with time: 6 (total remaining time 2977)
% dis+11_28_av=off:fsr=off:irw=on:lcm=predicate:nm=2:newcnf=on:nwc=5:sp=occurrence:urr=on:updr=off_4 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 3709
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% spawned child 10165 with time: 6 (total remaining time 2969)
% lrs+1011_5_add=large:afp=1000:afq=1.2:anc=none:fsr=off:irw=on:lma=on:nm=64:newcnf=on:nwc=1:stl=30:sd=3:ss=axioms:st=2.0:sos=on:sac=on:sp=reverse_arity:urr=on:updr=off_4 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7419
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% spawned child 10167 with time: 4 (total remaining time 2961)
% lrs+1011_3:1_add=off:afr=on:afp=10000:afq=1.1:amm=off:bce=on:cond=on:ep=R:fsr=off:nm=16:nwc=1:stl=30:sos=all:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit

% Memory used [KB]: 8571
% Time elapsed: 0.517 s
% ------------------------------
% ------------------------------
% spawned child 10168 with time: 23 (total remaining time 2956)
% lrs+1011_3:2_aac=none:afp=10000:afq=1.2:amm=off:anc=all:cond=on:fde=none:gs=on:gsem=on:irw=on:lma=on:nm=32:newcnf=on:nwc=3:nicw=on:stl=30:sac=on:sp=reverse_arity:urr=ec_only:updr=off:uhcvi=on_21 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 17782
% Time elapsed: 2.400 s
% ------------------------------
% ------------------------------
% spawned child 10169 with time: 4 (total remaining time 2932)
% lrs+10_6_aac=none:acc=model:add=large:afr=on:afp=10000:afq=1.2:amm=off:anc=none:bd=off:ccuc=small_ones:irw=on:lcm=reverse:nm=0:nwc=1:nicw=on:stl=30:sos=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7547
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 10170 with time: 7 (total remaining time 2926)
% dis+11_5_av=off:bd=off:bs=unit_only:bsr=on:cond=on:lcm=reverse:nm=0:nwc=1.2_5 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 2942
% Time elapsed: 0.800 s
% ------------------------------
% ------------------------------
% spawned child 10171 with time: 4 (total remaining time 2919)
% dis+11_32_av=off:ep=RST:fsr=off:lwlo=on:nm=6:nwc=1.1:sd=5:ss=axioms:st=5.0:sp=reverse_arity:uhcvi=on_2 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 1918
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 10172 with time: 7 (total remaining time 2913)
% dis+1003_64_add=off:afr=on:afp=100000:afq=1.1:anc=none:cond=fast:fde=none:irw=on:nm=6:newcnf=on:nwc=1.3:uhcvi=on_5 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7164
% Time elapsed: 0.800 s
% ------------------------------
% ------------------------------
% spawned child 10173 with time: 5 (total remaining time 2905)
% dis+11_4_afr=on:afp=1000:afq=1.4:amm=off:anc=none:lcm=reverse:lma=on:lwlo=on:nm=6:newcnf=on:nwc=1:sd=2:ss=axioms:st=2.0:sp=reverse_arity_3 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6908
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% spawned child 10177 with time: 87 (total remaining time 2898)
% ott+1002_128_av=off:bd=off:bs=on:bsr=on:cond=on:fsr=off:nm=6:newcnf=on:nwc=1:sp=reverse_arity:updr=off_82 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 5756
% Time elapsed: 8.800 s
% ------------------------------
% ------------------------------
% spawned child 10178 with time: 8 (total remaining time 2816)
% lrs+1002_8_add=large:afp=40000:afq=1.0:amm=off:anc=none:cond=on:gs=on:irw=on:nm=16:newcnf=on:nwc=1:stl=30:sos=on:sp=reverse_arity:updr=off_6 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 10490
% Time elapsed: 0.900 s
% ------------------------------
% ------------------------------
% spawned child 10180 with time: 54 (total remaining time 2807)
% lrs+1011_3:1_av=off:cond=on:irw=on:lma=on:nm=16:nwc=1:stl=30:sos=all:updr=off_51 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 9083
% Time elapsed: 5.500 s
% ------------------------------
% ------------------------------
% spawned child 10181 with time: 7 (total remaining time 2754)
% lrs+1_3:2_aac=none:add=large:anc=all_dependent:bce=on:cond=fast:ep=RST:fsr=off:lma=on:nm=2:nwc=1:stl=30:sos=on:sp=occurrence:urr=on:updr=off:uhcvi=on_5 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit

% Memory used [KB]: 9083
% Time elapsed: 0.805 s
% ------------------------------
% ------------------------------
% spawned child 10182 with time: 11 (total remaining time 2746)
% ott+1_8_av=off:bd=off:bs=on:cond=on:gsp=input_only:gs=on:gsem=off:irw=on:lcm=predicate:lwlo=on:nwc=1:sos=on_10 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7675
% Time elapsed: 1.200 s
% ------------------------------
% ------------------------------
% spawned child 10184 with time: 7 (total remaining time 2734)
% dis+11_5_afp=40000:afq=1.4:anc=none:br=off:cond=on:fsr=off:irw=on:lma=on:nm=4:nwc=1:sd=3:ss=axioms:st=2.0:urr=on:updr=off_5 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8955
% Time elapsed: 0.800 s
% ------------------------------
% ------------------------------
% spawned child 10185 with time: 4 (total remaining time 2725)
% dis+1010_10_afr=on:afp=100000:afq=2.0:amm=sco:anc=none:ep=RS:fde=unused:gs=on:nwc=1:sos=on:sac=on:sp=occurrence_2 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 11385
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 10186 with time: 9 (total remaining time 2719)
% lrs+1011_8:1_afr=on:afp=1000:afq=2.0:br=off:gsp=input_only:gs=on:nm=16:nwc=1:stl=30:sos=all:sp=occurrence:urr=on_7 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 14456
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% spawned child 10187 with time: 9 (total remaining time 2709)
% lrs+1003_2_acc=on:afp=4000:afq=2.0:amm=sco:bd=off:ccuc=small_ones:fsr=off:fde=unused:gsp=input_only:nm=64:newcnf=on:nwc=2.5:nicw=on:stl=30:urr=ec_only_7 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 9210
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% spawned child 10189 with time: 4 (total remaining time 2698)
% ott+4_40_av=off:bce=on:cond=fast:fde=none:nm=0:nwc=1:sos=all:updr=off_2 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 2558
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 10190 with time: 5 (total remaining time 2693)
% dis-3_5_av=off:cond=on:gs=on:gsem=off:lcm=predicate:nm=0:nwc=1:sos=on:sp=reverse_arity_3 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8443
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% spawned child 10191 with time: 56 (total remaining time 2686)
% lrs+1011_14_av=off:fsr=off:irw=on:nwc=1:stl=30:sos=on:sp=occurrence:urr=ec_only:updr=off_53 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 10746
% Time elapsed: 5.700 s
% ------------------------------
% ------------------------------
% spawned child 10192 with time: 5 (total remaining time 2631)
% dis+11_10_av=off:lma=on:nm=64:nwc=5:sp=reverse_arity_3 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 1535
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% spawned child 10193 with time: 13 (total remaining time 2625)
% lrs+1011_4:1_av=off:fsr=off:irw=on:nwc=1:stl=30:sd=4:ss=axioms:st=1.5:sp=reverse_arity_12 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 4733
% Time elapsed: 1.400 s
% ------------------------------
% ------------------------------
% spawned child 10194 with time: 245 (total remaining time 2611)
% lrs+10_2:3_afp=1000:afq=1.1:amm=sco:anc=none:er=known:lcm=reverse:lwlo=on:nm=64:newcnf=on:nwc=2:stl=30:sd=5:ss=axioms:sos=theory:sac=on:sp=occurrence_233 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 38890
% Time elapsed: 24.600 s
% ------------------------------
% ------------------------------
% spawned child 10196 with time: 7 (total remaining time 2387)
% lrs+1011_3:2_av=off:er=known:lma=on:nm=2:newcnf=on:nwc=2:stl=30:sd=2:ss=axioms:st=3.0:urr=on:updr=off_5 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit

% Memory used [KB]: 767
% Time elapsed: 2.325 s
% ------------------------------
% ------------------------------
% spawned child 10197 with time: 38 (total remaining time 2363)
% lrs+1011_2:3_av=off:gs=on:gsem=off:nwc=1.5:stl=30:sos=theory:sp=occurrence:urr=ec_only:updr=off_36 on codi_down_sum_theorems_goal07.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 24690
% Time elapsed: 3.900 s
% ------------------------------
% ------------------------------
% spawned child 10199 with time: 62 (total remaining time 2327)
% lrs+11_4_av=off:gsp=input_only:irw=on:lma=on:nm=0:nwc=1.2:stl=30:sd=2:ss=axioms:sp=reverse_arity:urr=on:updr=off_59 on codi_down_sum_theorems_goal07.all
% Refutation found. Thanks to Tanya!
% SZS status Theorem for codi_down_sum_theorems_goal07.all
% SZS output start Proof for codi_down_sum_theorems_goal07.all
fof(f1,conjecture,(
  ! [X0] : ! [X1] : (eqdim(X0,sum(X0,X1)) | ~(~zex(X1) & ~zex(X0) & eqdim(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',conjecture70)).
fof(f2,negated_conjecture,(
  ~! [X0] : ! [X1] : (eqdim(X0,sum(X0,X1)) | ~(~zex(X1) & ~zex(X0) & eqdim(X0,X1)))),
  inference(negated_conjecture,[],[f1])).
fof(f3,axiom,(
  ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom10)).
fof(f5,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (cont(X2,sum(X0,X1)) | ~(cont(X2,X1) & leq(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom30)).
fof(f6,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom40)).
fof(f74,axiom,(
  ! [X0] : ! [X1] : (~lt(X1,X0) | ~lt(X0,X1))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom89400)).
fof(f75,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (lt(X0,X2) | ~(lt(X1,X2) & lt(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom89410)).
fof(f77,axiom,(
  ! [X0] : ! [X1] : (eqdim(X1,X0) | ~eqdim(X0,X1))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom89430)).
fof(f78,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (eqdim(X0,X2) | ~(eqdim(X1,X2) & eqdim(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom89440)).
fof(f79,axiom,(
  ! [X0] : ! [X1] : (~eqdim(X0,X1) | ~lt(X0,X1))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom89450)).
fof(f85,axiom,(
  ! [X0] : ((~zex(X0) | ~cont(X0,X0)) & (cont(X0,X0) | zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom89510)).
fof(f88,axiom,(
  ! [X0] : ! [X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom89540)).
fof(f93,axiom,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom88850)).
fof(f106,axiom,(
  ! [X0] : ! [X1] : ((eqdim(X0,X1) | ~(leq(X1,X0) & leq(X0,X1))) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom89310)).
fof(f109,axiom,(
  ! [X0] : ! [X1] : (difference(X0,X1) = X0 | ~lt(X1,X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal07.all.tptp',axiom31670)).
fof(f118,plain,(
  ~! [X0,X1] : (eqdim(X0,sum(X0,X1)) | ~(~zex(X1) & ~zex(X0) & eqdim(X0,X1)))),
  inference(flattening,[],[f2])).
fof(f119,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  inference(flattening,[],[f6])).
fof(f120,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | ~(cont(X2,X1) & leq(X0,X1)))),
  inference(flattening,[],[f5])).
fof(f122,plain,(
  ! [X0,X1] : sum(X0,X1) = sum(X1,X0)),
  inference(flattening,[],[f3])).
fof(f124,plain,(
  ! [X0,X1,X2] : (eqdim(X0,X2) | ~(eqdim(X1,X2) & eqdim(X0,X1)))),
  inference(flattening,[],[f78])).
fof(f130,plain,(
  ! [X0,X1] : (eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1))),
  inference(flattening,[],[f93])).
fof(f132,plain,(
  ! [X0,X1] : (~eqdim(X0,X1) | ~lt(X0,X1))),
  inference(flattening,[],[f79])).
fof(f133,plain,(
  ! [X0,X1] : (eqdim(X1,X0) | ~eqdim(X0,X1))),
  inference(flattening,[],[f77])).
fof(f149,plain,(
  ! [X0,X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  inference(flattening,[],[f88])).
fof(f163,plain,(
  ! [X0,X1] : ((eqdim(X0,X1) | ~(leq(X1,X0) & leq(X0,X1))) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1)))),
  inference(flattening,[],[f106])).
fof(f170,plain,(
  ! [X0,X1] : (difference(X0,X1) = X0 | ~lt(X1,X0))),
  inference(flattening,[],[f109])).
fof(f173,plain,(
  ! [X0,X1,X2] : (lt(X0,X2) | ~(lt(X1,X2) & lt(X0,X1)))),
  inference(flattening,[],[f75])).
fof(f175,plain,(
  ! [X0,X1] : (~lt(X1,X0) | ~lt(X0,X1))),
  inference(flattening,[],[f74])).
fof(f195,plain,(
  ? [X0,X1] : (~eqdim(X0,sum(X0,X1)) & (~zex(X1) & ~zex(X0) & eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f118])).
fof(f196,plain,(
  ? [X0,X1] : (~eqdim(X0,sum(X0,X1)) & ~zex(X1) & ~zex(X0) & eqdim(X0,X1))),
  inference(flattening,[],[f195])).
fof(f197,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | (cont(X2,X0) | ~cont(X2,sum(X0,X1))))),
  inference(ennf_transformation,[],[f119])).
fof(f198,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f197])).
fof(f199,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | (~cont(X2,X1) | ~leq(X0,X1)))),
  inference(ennf_transformation,[],[f120])).
fof(f200,plain,(
  ! [X0,X1,X2] : (cont(X2,sum(X0,X1)) | ~cont(X2,X1) | ~leq(X0,X1))),
  inference(flattening,[],[f199])).
fof(f203,plain,(
  ! [X0,X1,X2] : (eqdim(X0,X2) | (~eqdim(X1,X2) | ~eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f204,plain,(
  ! [X0,X1,X2] : (eqdim(X0,X2) | ~eqdim(X1,X2) | ~eqdim(X0,X1))),
  inference(flattening,[],[f203])).
fof(f245,plain,(
  ! [X0,X1] : ((eqdim(X0,X1) | (~leq(X1,X0) | ~leq(X0,X1))) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f163])).
fof(f246,plain,(
  ! [X0,X1] : ((eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1)) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1)))),
  inference(flattening,[],[f245])).
fof(f257,plain,(
  ! [X0,X1,X2] : (lt(X0,X2) | (~lt(X1,X2) | ~lt(X0,X1)))),
  inference(ennf_transformation,[],[f173])).
fof(f258,plain,(
  ! [X0,X1,X2] : (lt(X0,X2) | ~lt(X1,X2) | ~lt(X0,X1))),
  inference(flattening,[],[f257])).
fof(f267,plain,(
  ? [X0,X1] : (~eqdim(X0,sum(X0,X1)) & ~zex(X1) & ~zex(X0) & eqdim(X0,X1)) => (~eqdim(sK0,sum(sK0,sK1)) & ~zex(sK1) & ~zex(sK0) & eqdim(sK0,sK1))),
  introduced(choice_axiom,[])).
fof(f268,plain,(
  ~eqdim(sK0,sum(sK0,sK1)) & ~zex(sK1) & ~zex(sK0) & eqdim(sK0,sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f196,f267])).
fof(f285,plain,(
  eqdim(sK0,sK1)),
  inference(cnf_transformation,[],[f268])).
fof(f286,plain,(
  ~zex(sK0)),
  inference(cnf_transformation,[],[f268])).
fof(f287,plain,(
  ~zex(sK1)),
  inference(cnf_transformation,[],[f268])).
fof(f288,plain,(
  ~eqdim(sK0,sum(sK0,sK1))),
  inference(cnf_transformation,[],[f268])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (~cont(X2,sum(X0,X1)) | cont(X2,X0) | cont(difference(X2,X0),X1)) )),
  inference(cnf_transformation,[],[f198])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (~cont(X2,X1) | cont(X2,sum(X0,X1)) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f200])).
fof(f292,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = sum(X1,X0)) )),
  inference(cnf_transformation,[],[f122])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (~eqdim(X0,X1) | ~eqdim(X1,X2) | eqdim(X0,X2)) )),
  inference(cnf_transformation,[],[f204])).
fof(f304,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f130])).
fof(f306,plain,(
  ( ! [X0,X1] : (~lt(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f132])).
fof(f307,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0)) )),
  inference(cnf_transformation,[],[f133])).
fof(f327,plain,(
  ( ! [X0,X1] : (~cont(X1,X0) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f149])).
fof(f335,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f85])).
fof(f350,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f366,plain,(
  ( ! [X0,X1] : (~lt(X1,X0) | difference(X0,X1) = X0) )),
  inference(cnf_transformation,[],[f170])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (~lt(X0,X1) | ~lt(X1,X2) | lt(X0,X2)) )),
  inference(cnf_transformation,[],[f258])).
fof(f373,plain,(
  ( ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0)) )),
  inference(cnf_transformation,[],[f175])).
fof(f402,plain,(
  cont(sK0,sK0)),
  inference(unit_resulting_resolution,[],[f286,f335])).
fof(f403,plain,(
  cont(sK1,sK1)),
  inference(unit_resulting_resolution,[],[f287,f335])).
fof(f432,plain,(
  ~eqdim(sum(sK0,sK1),sK0)),
  inference(unit_resulting_resolution,[],[f288,f307])).
fof(f433,plain,(
  eqdim(sK1,sK0)),
  inference(unit_resulting_resolution,[],[f285,f307])).
fof(f438,plain,(
  ~lt(sK1,sK0)),
  inference(unit_resulting_resolution,[],[f433,f306])).
fof(f466,plain,(
  leq(sK0,sK1)),
  inference(unit_resulting_resolution,[],[f285,f350])).
fof(f468,plain,(
  leq(sK1,sK0)),
  inference(unit_resulting_resolution,[],[f433,f350])).
fof(f859,plain,(
  ~eqdim(sum(sK0,sK1),sK1)),
  inference(unit_resulting_resolution,[],[f432,f433,f295])).
fof(f861,plain,(
  ~eqdim(sK1,sum(sK0,sK1))),
  inference(unit_resulting_resolution,[],[f288,f285,f295])).
fof(f2167,plain,(
  cont(sK0,sum(sK1,sK0))),
  inference(unit_resulting_resolution,[],[f468,f402,f290])).
fof(f2171,plain,(
  cont(sK1,sum(sK0,sK1))),
  inference(unit_resulting_resolution,[],[f466,f403,f290])).
fof(f2195,plain,(
  cont(sK0,sum(sK0,sK1))),
  inference(forward_demodulation,[],[f2167,f292])).
fof(f2225,plain,(
  lt(sK1,sum(sK0,sK1))),
  inference(unit_resulting_resolution,[],[f861,f2171,f304])).
fof(f2230,plain,(
  ~zex(sum(sK0,sK1))),
  inference(unit_resulting_resolution,[],[f2171,f327])).
fof(f2241,plain,(
  cont(sum(sK0,sK1),sum(sK0,sK1))),
  inference(unit_resulting_resolution,[],[f2230,f335])).
fof(f2397,plain,(
  lt(sK0,sum(sK0,sK1))),
  inference(unit_resulting_resolution,[],[f288,f2195,f304])).
fof(f2438,plain,(
  ~lt(sum(sK0,sK1),sK0)),
  inference(unit_resulting_resolution,[],[f438,f2225,f371])).
fof(f2444,plain,(
  ~lt(sum(sK0,sK1),sK1)),
  inference(unit_resulting_resolution,[],[f2225,f373])).
fof(f2714,plain,(
  sum(sK0,sK1) = difference(sum(sK0,sK1),sK0)),
  inference(unit_resulting_resolution,[],[f2397,f366])).
fof(f2817,plain,(
  ~cont(sum(sK0,sK1),sK0)),
  inference(unit_resulting_resolution,[],[f432,f2438,f304])).
fof(f2827,plain,(
  ~cont(sum(sK0,sK1),sK1)),
  inference(unit_resulting_resolution,[],[f859,f2444,f304])).
fof(f3923,plain,(
  cont(difference(sum(sK0,sK1),sK0),sK1)),
  inference(unit_resulting_resolution,[],[f2817,f2241,f289])).
fof(f3961,plain,(
  cont(sum(sK0,sK1),sK1)),
  inference(forward_demodulation,[],[f3923,f2714])).
fof(f3962,plain,(
  $false),
  inference(subsumption_resolution,[],[f3961,f2827])).
% SZS output end Proof for codi_down_sum_theorems_goal07.all
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 1407
% Time elapsed: 0.727 s
% ------------------------------
% ------------------------------
% Success in time 67.998 s
