Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal07.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 475
Passive clauses: 1751
Generated clauses: 3938
Final active clauses: 475
Final passive clauses: 1183
Input formulas: 43
Initial clauses: 92

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 39
Fw subsumption resolutions: 118

Simple tautologies: 49
Forward subsumptions: 2036

Binary resolution: 3615
Factoring: 30

Backtracking splits: 32
Backtracking splits refuted: 6

Memory used: 767KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 5993 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal07.all
Refutation found. Thanks to Tanya!
2233. $false (9:0) [subsumption resolution 2228,216]
216. eqdim(sK2,sK0) (0:3) [cnf transformation 169]
169. cont(sK2,sK0) & cont(sK2,sK1) & eqdim(sK2,sK0) & eqdim(sK2,sK1) & ~cont(sK2,intersection(sK0,sK1))[skolemisation 121]
121. ? [X0,X1,X2] : (cont(X2,X0) & cont(X2,X1) & eqdim(X2,X0) & eqdim(X2,X1) & ~cont(X2,intersection(X0,X1)))[flattening 120]
120. ? [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,X0) & eqdim(X2,X1)) & ~cont(X2,intersection(X0,X1)))[ennf transformation 44]
44. ~! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,X0) & eqdim(X2,X1)) => cont(X2,intersection(X0,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,X0) & eqdim(X2,X1)) => cont(X2,intersection(X0,X1)))[negated conjecture]
2228. ~eqdim(sK2,sK0) (9:3) [resolution 2189,238]
238. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 134]
134. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 64]
64. ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))[flattening 63]
63. ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))[rectify 33]
33. ! [X67] : ! [X68] : (eqdim(X67,X68) => eqdim(X68,X67))[input]
2189. ~eqdim(sK0,sK2) (8:3) [subsumption resolution 2179,1524]
1524. ~zex(intersection(sK0,sK1)) (5:4) [resolution 1519,231]
231. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 126]
126. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 52]
52. ! [X0,X1] : (zex(X0) => leq(X0,X1))[flattening 51]
51. ! [X0] : ! [X1] : (zex(X0) => leq(X0,X1))[rectify 16]
16. ! [X31] : ! [X32] : (zex(X31) => leq(X31,X32))[input]
1519. ~leq(intersection(sK0,sK1),sK2) (4:5) [subsumption resolution 1518,214]
214. cont(sK2,sK0) (0:3) [cnf transformation 169]
1518. ~leq(intersection(sK0,sK1),sK2) | ~cont(sK2,sK0) (4:8) [subsumption resolution 1502,215]
215. cont(sK2,sK1) (0:3) [cnf transformation 169]
1502. ~leq(intersection(sK0,sK1),sK2) | ~cont(sK2,sK1) | ~cont(sK2,sK0) (4:11) [resolution 1173,288]
288. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 161]
161. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 160]
160. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 110]
110. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[flattening 109]
109. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[rectify 26]
26. ! [X52] : ! [X53] : ! [X54] : ((cont(X54,X52) & cont(X54,X53)) => leq(X54,intersection(X52,X53)))[input]
1173. ~leq(sK2,intersection(sK0,sK1)) | ~leq(intersection(sK0,sK1),sK2) (3:10) [resolution 1169,281]
281. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 207]
207. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 206]
206. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[nnf transformation 98]
98. ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[flattening 97]
97. ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[rectify 41]
41. ! [X86] : ! [X87] : (eqdim(X86,X87) <=> (leq(X86,X87) & leq(X87,X86)))[input]
1169. ~eqdim(sK2,intersection(sK0,sK1)) (2:5) [subsumption resolution 1168,214]
1168. ~eqdim(sK2,intersection(sK0,sK1)) | ~cont(sK2,sK0) (2:8) [subsumption resolution 1159,215]
1159. ~eqdim(sK2,intersection(sK0,sK1)) | ~cont(sK2,sK1) | ~cont(sK2,sK0) (2:11) [resolution 291,320]
320. ~p(sK2,intersection(sK0,sK1)) (1:5) [resolution 273,218]
218. ~cont(sK2,intersection(sK0,sK1)) (0:5) [cnf transformation 169]
273. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 203]
203. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 202]
202. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[nnf transformation 94]
94. ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[flattening 93]
93. ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[rectify 7]
7. ! [X15] : ! [X16] : (p(X15,X16) <=> (cont(X15,X16) & eqdim(X15,X16)))[input]
291. p(X2,intersection(X0,X1)) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:16) [cnf transformation 210]
210. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 209]
209. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[nnf transformation 116]
116. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[flattening 115]
115. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[rectify 27]
27. ! [X55] : ! [X56] : ! [X57] : ((cont(X57,X55) & cont(X57,X56) & eqdim(X57,intersection(X55,X56))) <=> p(X57,intersection(X55,X56)))[input]
2179. ~eqdim(sK0,sK2) | zex(intersection(sK0,sK1)) (8:7) [resolution 1815,244]
244. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 145]
145. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 76]
76. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 75]
75. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 25]
25. ! [X50] : ! [X51] : (~zex(intersection(X50,X51)) => cont(intersection(X50,X51),X50))[input]
1815. ~cont(intersection(sK0,sK1),X0) | ~eqdim(X0,sK2) (7:8) [subsumption resolution 1799,1184]
1184. ~eqdim(intersection(sK0,sK1),X0) | ~eqdim(X0,sK2) (4:8) [resolution 1175,290]
290. eqdim(X0,X2) | ~eqdim(X1,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 165]
165. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~eqdim(X1,X2) | eqdim(X0,X2))[flattening 164]
164. ! [X0,X1,X2] : ((~eqdim(X0,X1) | ~eqdim(X1,X2)) | eqdim(X0,X2))[ennf transformation 114]
114. ! [X0,X1,X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))[flattening 113]
113. ! [X0] : ! [X1] : ! [X2] : ((eqdim(X0,X1) & eqdim(X1,X2)) => eqdim(X0,X2))[rectify 34]
34. ! [X69] : ! [X70] : ! [X71] : ((eqdim(X69,X70) & eqdim(X70,X71)) => eqdim(X69,X71))[input]
1175. ~eqdim(intersection(sK0,sK1),sK2) (3:5) [resolution 1169,238]
1799. ~eqdim(X0,sK2) | eqdim(intersection(sK0,sK1),X0) | ~cont(intersection(sK0,sK1),X0) (7:13) [resolution 1589,234]
234. lt(X0,X1) | eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 129]
129. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 128]
128. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 56]
56. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 55]
55. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 28]
28. ! [X58] : ! [X59] : (cont(X58,X59) => (lt(X58,X59) | eqdim(X58,X59)))[input]
1589. ~lt(intersection(sK0,sK1),X2) | ~eqdim(X2,sK2) (6:8) [resolution 1525,284]
284. lt(X2,X1) | ~lt(X2,X0) | ~eqdim(X0,X1) (0:9) [cnf transformation 153]
153. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X2,X0) | lt(X2,X1))[flattening 152]
152. ! [X0,X1,X2] : (~eqdim(X0,X1) | (~lt(X2,X0) | lt(X2,X1)))[ennf transformation 102]
102. ! [X0,X1,X2] : (eqdim(X0,X1) => (lt(X2,X0) => lt(X2,X1)))[flattening 101]
101. ! [X0] : ! [X1] : ! [X2] : (eqdim(X0,X1) => (lt(X2,X0) => lt(X2,X1)))[rectify 36]
36. ! [X74] : ! [X75] : ! [X76] : (eqdim(X74,X75) => (lt(X76,X74) => lt(X76,X75)))[input]
1525. ~lt(intersection(sK0,sK1),sK2) (5:5) [resolution 1519,276]
276. leq(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 205]
205. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | lt(X0,X1)))[flattening 204]
204. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | lt(X0,X1)))[nnf transformation 150]
150. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[ennf transformation 96]
96. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[flattening 95]
95. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[rectify 6]
6. ! [X13] : ! [X14] : (lt(X13,X14) <=> (leq(X13,X14) & ~eqdim(X13,X14)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 280
Passive clauses: 923
Generated clauses: 2014
Final active clauses: 280
Final passive clauses: 614
Input formulas: 43
Initial clauses: 87

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 50
Fw subsumption resolutions: 132

Simple tautologies: 42
Equational tautologies: 3
Forward subsumptions: 891

Binary resolution: 1052
Factoring: 16
Forward superposition: 620
Backward superposition: 60

Unique components: 5

Memory used: 511KB
Time elapsed: 0.188 s
------------------------------
% Success in time 0.992 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 23:33:15 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal07.all.tptp"
Input read from 
============================ end of footer ===========================
