Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal15.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 558
Passive clauses: 2103
Generated clauses: 3914
Final active clauses: 556
Final passive clauses: 1435
Input formulas: 43
Initial clauses: 89

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 36
Fw subsumption resolutions: 120

Simple tautologies: 45
Forward subsumptions: 1507

Binary resolution: 3527
Factoring: 15

Backtracking splits: 104
Backtracking splits refuted: 10

Memory used: 895KB
Time elapsed: 0.406 s
------------------------------
% remaining time: 5992 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal15.all
Refutation found. Thanks to Tanya!
7019. $false (3:0) [subsumption resolution 7018,212]
212. ~zex(intersection(sK0,sK1)) (0:4) [cnf transformation 167]
167. ~zex(intersection(sK0,sK1)) & ~p(intersection(sK0,sK1),intersection(sK0,intersection(sK0,sK1)))[skolemisation 119]
119. ? [X0,X1] : (~zex(intersection(X0,X1)) & ~p(intersection(X0,X1),intersection(X0,intersection(X0,X1))))[ennf transformation 44]
44. ~! [X0,X1] : (~zex(intersection(X0,X1)) => p(intersection(X0,X1),intersection(X0,intersection(X0,X1))))[flattening 1]
1. ~! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => p(intersection(X0,X1),intersection(X0,intersection(X0,X1))))[negated conjecture]
7018. zex(intersection(sK0,sK1)) (3:4) [subsumption resolution 7009,239]
239. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 143]
143. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 76]
76. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 75]
75. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 30]
30. ! [X59] : ! [X60] : (~zex(intersection(X59,X60)) => cont(intersection(X59,X60),X59))[input]
7009. ~cont(intersection(sK0,sK1),sK0) | zex(intersection(sK0,sK1)) (3:9) [resolution 6971,223]
223. cont(X0,X0) | zex(X0) (0:5) [cnf transformation 176]
176. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[nnf transformation 122]
122. ! [X0] : (~zex(X0) <=> cont(X0,X0))[ennf transformation 49]
49. ! [X0] : (~zex(X0) <=> cont(X0,X0))[rectify 23]
23. ! [X45] : (~zex(X45) <=> cont(X45,X45))[input]
6971. ~cont(intersection(sK0,sK1),intersection(sK0,sK1)) | ~cont(intersection(sK0,sK1),sK0) (2:12) [subsumption resolution 1716,3535]
3535. ~lt(X7,intersection(X6,X7)) (3:5) [subsumption resolution 3522,572]
572. ~lt(X9,X10) | ~zex(X10) (2:5) [resolution 566,226]
226. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 124]
124. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 52]
52. ! [X0,X1] : (zex(X0) => leq(X0,X1))[flattening 51]
51. ! [X0] : ! [X1] : (zex(X0) => leq(X0,X1))[rectify 40]
40. ! [X83] : ! [X84] : (zex(X83) => leq(X83,X84))[input]
566. ~leq(X2,X3) | ~lt(X3,X2) (1:6) [subsumption resolution 564,274]
274. leq(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 205]
205. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | lt(X0,X1)))[flattening 204]
204. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | lt(X0,X1)))[nnf transformation 148]
148. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[ennf transformation 97]
97. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[flattening 96]
96. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[rectify 18]
18. ! [X34] : ! [X35] : (lt(X34,X35) <=> (leq(X34,X35) & ~eqdim(X34,X35)))[input]
564. ~leq(X2,X3) | ~leq(X3,X2) | ~lt(X3,X2) (1:9) [resolution 273,232]
232. ~eqdim(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 131]
131. ! [X0,X1] : (~lt(X0,X1) | ~eqdim(X0,X1))[ennf transformation 62]
62. ! [X0,X1] : (lt(X0,X1) => ~eqdim(X0,X1))[flattening 61]
61. ! [X0] : ! [X1] : (lt(X0,X1) => ~eqdim(X0,X1))[rectify 8]
8. ! [X14] : ! [X15] : (lt(X14,X15) => ~eqdim(X14,X15))[input]
273. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 203]
203. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 202]
202. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[nnf transformation 95]
95. ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[flattening 94]
94. ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[rectify 28]
28. ! [X55] : ! [X56] : (eqdim(X55,X56) <=> (leq(X55,X56) & leq(X56,X55)))[input]
3522. zex(intersection(X6,X7)) | ~lt(X7,intersection(X6,X7)) (3:9) [resolution 1593,513]
513. ~cont(X0,X1) | ~lt(X1,X0) (2:6) [subsumption resolution 511,233]
233. ~lt(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 132]
132. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[ennf transformation 64]
64. ! [X0,X1] : (lt(X0,X1) => ~lt(X1,X0))[flattening 63]
63. ! [X0] : ! [X1] : (lt(X0,X1) => ~lt(X1,X0))[rectify 3]
3. ! [X3] : ! [X4] : (lt(X3,X4) => ~lt(X4,X3))[input]
511. lt(X0,X1) | ~cont(X0,X1) | ~lt(X1,X0) (2:9) [resolution 229,307]
307. ~eqdim(X2,X1) | ~lt(X1,X2) (1:6) [resolution 232,231]
231. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 130]
130. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 60]
60. ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))[flattening 59]
59. ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))[rectify 6]
6. ! [X9] : ! [X10] : (eqdim(X9,X10) => eqdim(X10,X9))[input]
229. eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 127]
127. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 126]
126. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 56]
56. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 55]
55. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 34]
34. ! [X70] : ! [X71] : (cont(X70,X71) => (lt(X70,X71) | eqdim(X70,X71)))[input]
1593. cont(intersection(X0,X1),X1) | zex(intersection(X0,X1)) (2:9) [subsumption resolution 1571,215]
215. eqdim(X0,X0) (0:3) [cnf transformation 46]
46. ! [X0] : eqdim(X0,X0)[rectify 5]
5. ! [X8] : eqdim(X8,X8)[input]
1571. ~eqdim(intersection(X0,X1),intersection(X0,X1)) | cont(intersection(X0,X1),X1) | zex(intersection(X0,X1)) (2:16) [resolution 561,223]
561. ~cont(X5,intersection(X6,X7)) | ~eqdim(X5,intersection(X6,X7)) | cont(X5,X7) (1:13) [resolution 270,288]
288. ~p(X2,intersection(X0,X1)) | cont(X2,X1) (0:8) [cnf transformation 208]
208. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 207]
207. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[nnf transformation 115]
115. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[flattening 114]
114. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[rectify 32]
32. ! [X64] : ! [X65] : ! [X66] : ((cont(X66,X64) & cont(X66,X65) & eqdim(X66,intersection(X64,X65))) <=> p(X66,intersection(X64,X65)))[input]
270. p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 201]
201. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 200]
200. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[nnf transformation 93]
93. ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[flattening 92]
92. ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[rectify 19]
19. ! [X36] : ! [X37] : (p(X36,X37) <=> (cont(X36,X37) & eqdim(X36,X37)))[input]
1716. ~cont(intersection(sK0,sK1),intersection(sK0,sK1)) | ~cont(intersection(sK0,sK1),sK0) | lt(intersection(sK0,sK1),intersection(sK0,intersection(sK0,sK1))) (2:21) [subsumption resolution 1703,282]
282. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 157]
157. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 156]
156. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 107]
107. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[flattening 106]
106. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[rectify 31]
31. ! [X61] : ! [X62] : ! [X63] : ((cont(X63,X61) & cont(X63,X62)) => leq(X63,intersection(X61,X62)))[input]
1703. ~cont(intersection(sK0,sK1),intersection(sK0,sK1)) | ~cont(intersection(sK0,sK1),sK0) | lt(intersection(sK0,sK1),intersection(sK0,intersection(sK0,sK1))) | ~leq(intersection(sK0,sK1),intersection(sK0,intersection(sK0,sK1))) (2:30) [resolution 938,276]
276. eqdim(X0,X1) | lt(X0,X1) | ~leq(X0,X1) (0:9) [cnf transformation 205]
938. ~eqdim(intersection(sK0,sK1),intersection(sK0,intersection(sK0,sK1))) | ~cont(intersection(sK0,sK1),intersection(sK0,sK1)) | ~cont(intersection(sK0,sK1),sK0) (1:21) [resolution 286,213]
213. ~p(intersection(sK0,sK1),intersection(sK0,intersection(sK0,sK1))) (0:9) [cnf transformation 167]
286. p(X2,intersection(X0,X1)) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:16) [cnf transformation 208]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 654
Passive clauses: 2940
Generated clauses: 6802
Final active clauses: 654
Final passive clauses: 2176
Input formulas: 43
Initial clauses: 84

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 103
Fw subsumption resolutions: 376

Simple tautologies: 95
Equational tautologies: 6
Forward subsumptions: 3392

Binary resolution: 4845
Factoring: 75
Forward superposition: 1290
Backward superposition: 38

Unique components: 2

Memory used: 1151KB
Time elapsed: 1.136 s
------------------------------
% Success in time 2.575 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 00:58:59 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal15.all.tptp"
Input read from 
============================ end of footer ===========================
