Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal20.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 875
Passive clauses: 4552
Generated clauses: 10591
Final active clauses: 867
Final passive clauses: 3418
Input formulas: 50
Initial clauses: 87

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 63
Fw subsumption resolutions: 455

Simple tautologies: 56
Forward subsumptions: 5573

Binary resolution: 9808
Factoring: 26

Backtracking splits: 123
Backtracking splits refuted: 12

Memory used: 1918KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal20.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1235
Passive clauses: 7846
Generated clauses: 22781
Final active clauses: 1235
Final passive clauses: 6283
Input formulas: 50
Initial clauses: 86

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 244
Fw subsumption resolutions: 1114

Simple tautologies: 220
Equational tautologies: 11
Forward subsumptions: 13686

Binary resolution: 18497
Factoring: 114
Forward superposition: 2694
Backward superposition: 50

Unique components: 3

Memory used: 2558KB
Time elapsed: 1.402 s
------------------------------
% remaining time: 5981 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on theorems\conversions\codi_int_theorems_goal20.all
Refutation found. Thanks to Tanya!
32403. $false (1:0) [subsumption resolution 262,32401]
32401. ~cont(intersection(sK0,sK1),intersection(sK0,sK1)) (5:7) [backtracking split refutation 32387,32399,32400]
32400. $false {285} (1:0) [subsumption resolution 257,32399]
257. cont(intersection(sK0,sK1),sK0) (1:5) [resolution 173,213]
213. zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0) (0:9) [cnf transformation 65]
65. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 12]
12. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
173. ~zex(intersection(sK0,sK1)) (0:4) [cnf transformation 163]
163. ~zex(intersection(sK0,sK1)) & ~zex(intersection(sK0,intersection(sK0,sK1))) & ~p(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1))[skolemisation 96]
96. ? [X0,X1] : (~zex(intersection(X0,X1)) & ~zex(intersection(X0,intersection(X0,X1))) & ~p(intersection(X0,intersection(X0,X1)),intersection(X0,X1)))[flattening 95]
95. ? [X0,X1] : ((~zex(intersection(X0,X1)) & ~zex(intersection(X0,intersection(X0,X1)))) & ~p(intersection(X0,intersection(X0,X1)),intersection(X0,X1)))[ennf transformation 51]
51. ~! [X0,X1] : (~(~zex(intersection(X0,X1)) & ~zex(intersection(X0,intersection(X0,X1)))) | p(intersection(X0,intersection(X0,X1)),intersection(X0,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : (~(~zex(intersection(X0,X1)) & ~zex(intersection(X0,intersection(X0,X1)))) | p(intersection(X0,intersection(X0,X1)),intersection(X0,X1)))[negated conjecture]
32399. ~cont(intersection(sK0,sK1),sK0) {285} (5:5) [backtracking splitting component]
32387. ~cont(intersection(sK0,sK1),intersection(sK0,sK1)) | ~cont(intersection(sK0,sK1),sK0) (5:12) [resolution 8308,250]
250. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 157]
157. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 156]
156. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 92]
92. ! [X0,X1,X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[flattening 13]
13. ! [X0] : ! [X1] : ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[input]
8308. ~leq(intersection(sK0,sK1),intersection(sK0,intersection(sK0,sK1))) (4:9) [resolution 2571,210]
210. ~lt(X0,X1) | ~leq(X1,X0) (0:6) [cnf transformation 115]
115. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 114]
114. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[ennf transformation 63]
63. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[flattening 46]
46. ! [X0] : ! [X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[input]
2571. lt(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) (3:9) [backtracking split refutation 2549,2556,2570]
2570. $false {84} (3:0) [subsumption resolution 2556,2562]
2562. ~eqdim(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) (3:9) [backtracking split refutation 2546,2554,2555]
2555. $false {83} (0:0) [subsumption resolution 175,2554]
175. ~p(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) (0:9) [cnf transformation 163]
2554. p(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) {83} (3:9) [backtracking splitting component]
2546. ~eqdim(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) | p(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) (3:18) [resolution 1041,205]
205. ~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1) (0:9) [cnf transformation 111]
111. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 110]
110. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 61]
61. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 35]
35. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
1041. cont(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) (2:9) [resolution 266,238]
238. ~p(X2,intersection(X0,X1)) | cont(X2,X1) (0:8) [cnf transformation 138]
138. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 137]
137. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[ennf transformation 81]
81. ! [X0,X1,X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 14]
14. ! [X0] : ! [X1] : ! [X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[input]
266. p(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,intersection(sK0,sK1))) (1:11) [resolution 174,187]
187. zex(X0) | p(X0,X0) (0:5) [cnf transformation 2]
2. ! [X0] : (zex(X0) | p(X0,X0))[input]
174. ~zex(intersection(sK0,intersection(sK0,sK1))) (0:6) [cnf transformation 163]
2556. eqdim(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) {84} (3:9) [backtracking splitting component]
2549. lt(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) | eqdim(intersection(sK0,intersection(sK0,sK1)),intersection(sK0,sK1)) (3:18) [resolution 1041,218]
218. ~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1) (0:9) [cnf transformation 119]
119. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[ennf transformation 69]
69. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 36]
36. ! [X0] : ! [X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[input]
262. cont(intersection(sK0,sK1),intersection(sK0,sK1)) (1:7) [resolution 173,178]
178. zex(X0) | cont(X0,X0) (0:5) [cnf transformation 97]
97. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[ennf transformation 28]
28. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 3874
Passive clauses: 10515
Generated clauses: 32563
Final active clauses: 3697
Final passive clauses: 6180
Input formulas: 50
Initial clauses: 84

Duplicate literals: 198
Trivial inequalities: 7
Bw subsumption resolutions: 178
Fw demodulations: 1461
Condensations: 50

Simple tautologies: 181
Equational tautologies: 69
Forward subsumptions: 19644

Binary resolution: 28988
Factoring: 3
Forward superposition: 404
Backward superposition: 335

Backtracking splits: 285
Backtracking splits refuted: 120
Backtracking splits refuted at zero level: 10

Memory used: 7164KB
Time elapsed: 1.908 s
------------------------------
% Success in time 3.846 s
