Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal01.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 860
Passive clauses: 5153
Generated clauses: 10868
Final active clauses: 847
Final passive clauses: 4011
Input formulas: 53
Initial clauses: 106

Selected by SInE selection: 53
SInE iterations: 1

Duplicate literals: 76
Fw subsumption resolutions: 544

Simple tautologies: 56
Forward subsumptions: 4882

Binary resolution: 9845
Factoring: 38

Backtracking splits: 241
Backtracking splits refuted: 15
Backtracking splits refuted at zero level: 3

Memory used: 2046KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal01.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1410
Passive clauses: 9130
Generated clauses: 26355
Final active clauses: 1410
Final passive clauses: 7412
Input formulas: 53
Initial clauses: 105

Selected by SInE selection: 53
SInE iterations: 1

Duplicate literals: 331
Fw subsumption resolutions: 1389

Simple tautologies: 323
Equational tautologies: 11
Forward subsumptions: 15450

Binary resolution: 21385
Factoring: 135
Forward superposition: 2998
Backward superposition: 40

Unique components: 4

Memory used: 2814KB
Time elapsed: 1.401 s
------------------------------
% remaining time: 5981 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on theorems\conversions\codi_int_theorems_goal01.all
Refutation found. Thanks to Tanya!
3770. $false (2:0) [subsumption resolution 2073,3769]
3769. leq(intersection(sK1,sK0),intersection(sK0,sK1)) (3:7) [backtracking split refutation 1946,2075,3767]
3767. $false {90} (6:0) [subsumption resolution 3000,3765]
3765. ~cont(intersection(sK1,sK0),sK1) {90} (4:5) [backtracking split refutation 2971,2983,3764]
3764. $false {90, 101} (4:0) [subsumption resolution 2983,3762]
3762. cont(intersection(sK1,sK0),sK0) {90} (7:5) [resolution 3003,279]
279. ~p(X2,intersection(X0,X1)) | cont(X2,X1) (0:8) [cnf transformation 151]
151. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 150]
150. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[ennf transformation 90]
90. ! [X0,X1,X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 14]
14. ! [X0] : ! [X1] : ! [X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[input]
3003. p(intersection(sK1,sK0),intersection(sK1,sK0)) {90} (6:7) [resolution 2994,210]
210. zex(X0) | p(X0,X0) (0:5) [cnf transformation 2]
2. ! [X0] : (zex(X0) | p(X0,X0))[input]
2994. ~zex(intersection(sK1,sK0)) {90} (5:4) [backtracking split refutation 2989,2991,2992]
2992. $false {90, 103} (4:0) [subsumption resolution 2981,2991]
2981. ~zex(intersection(sK1,sK0)) {90} (4:4) [resolution 2075,255]
255. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 129]
129. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 75]
75. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[flattening 45]
45. ! [X0] : ! [X1] : (~zex(X0) | leq(X0,X1))[input]
2991. zex(intersection(sK1,sK0)) {103} (5:4) [backtracking splitting component]
2989. ~zex(sK3(intersection(sK1,sK0))) | zex(intersection(sK1,sK0)) {90} (5:9) [resolution 2970,208]
208. mindim(X0) | ~zex(sK3(X0)) | zex(X0) (0:7) [cnf transformation 180]
180. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & (zex(X0) | (~zex(sK3(X0)) & ~leq(X0,sK3(X0))) | mindim(X0)))[skolemisation 110]
110. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & (zex(X0) | ? [X2] : (~zex(X2) & ~leq(X0,X2)) | mindim(X0)))[flattening 109]
109. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & ((zex(X0) | ? [X2] : (~zex(X2) & ~leq(X0,X2))) | mindim(X0)))[ennf transformation 56]
56. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & (~(~zex(X0) & ! [X2] : (zex(X2) | leq(X0,X2))) | mindim(X0)))[rectify 47]
47. ! [X0] : ((~mindim(X0) | (~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1)))) & (~(~zex(X0) & ! [X1] : (zex(X1) | leq(X0,X1))) | mindim(X0)))[input]
2970. ~mindim(intersection(sK1,sK0)) {90} (4:4) [resolution 2075,350]
350. leq(X1,intersection(sK0,sK1)) | ~mindim(X1) (1:7) [resolution 195,207]
207. zex(X1) | leq(X0,X1) | ~mindim(X0) (0:7) [cnf transformation 180]
195. ~zex(intersection(sK0,sK1)) (0:4) [cnf transformation 178]
178. ~zex(intersection(sK0,sK1)) & ~zex(sK0) & ~zex(sK1) & ~p(intersection(sK0,sK1),intersection(sK1,sK0))[skolemisation 105]
105. ? [X0,X1] : (~zex(intersection(X0,X1)) & ~zex(X0) & ~zex(X1) & ~p(intersection(X0,X1),intersection(X1,X0)))[flattening 104]
104. ? [X0,X1] : ((~zex(intersection(X0,X1)) & ~zex(X0) & ~zex(X1)) & ~p(intersection(X0,X1),intersection(X1,X0)))[ennf transformation 54]
54. ~! [X0,X1] : (~(~zex(intersection(X0,X1)) & ~zex(X0) & ~zex(X1)) | p(intersection(X0,X1),intersection(X1,X0)))[flattening 1]
1. ~! [X0] : ! [X1] : (~(~zex(intersection(X0,X1)) & ~zex(X0) & ~zex(X1)) | p(intersection(X0,X1),intersection(X1,X0)))[negated conjecture]
2983. ~cont(intersection(sK1,sK0),sK0) {101} (4:5) [backtracking splitting component]
2971. ~cont(intersection(sK1,sK0),sK1) | ~cont(intersection(sK1,sK0),sK0) {90} (4:10) [resolution 2075,293]
293. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 174]
174. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 173]
173. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 103]
103. ! [X0,X1,X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[flattening 13]
13. ! [X0] : ! [X1] : ! [X2] : (~(cont(X2,X0) & cont(X2,X1)) | leq(X2,intersection(X0,X1)))[input]
3000. cont(intersection(sK1,sK0),sK1) {90} (6:5) [resolution 2994,254]
254. zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0) (0:9) [cnf transformation 74]
74. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[flattening 12]
12. ! [X0] : ! [X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[input]
2075. ~leq(intersection(sK1,sK0),intersection(sK0,sK1)) {90} (3:7) [backtracking splitting component]
1946. ~leq(intersection(sK0,sK1),intersection(sK1,sK0)) | ~leq(intersection(sK1,sK0),intersection(sK0,sK1)) (3:14) [resolution 1182,246]
246. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 124]
124. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 123]
123. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[ennf transformation 70]
70. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[flattening 52]
52. ! [X0] : ! [X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[input]
1182. ~eqdim(intersection(sK1,sK0),intersection(sK0,sK1)) (2:7) [resolution 892,258]
258. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 131]
131. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 77]
77. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[flattening 23]
23. ! [X0] : ! [X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[input]
892. ~eqdim(intersection(sK0,sK1),intersection(sK1,sK0)) (1:7) [backtracking split refutation 640,876,877]
877. $false {34} (1:0) [subsumption resolution 346,876]
346. cont(intersection(sK0,sK1),sK0) (1:5) [resolution 195,254]
876. ~cont(intersection(sK0,sK1),sK0) {34} (1:5) [backtracking splitting component]
640. ~eqdim(intersection(sK0,sK1),intersection(sK1,sK0)) | ~cont(intersection(sK0,sK1),sK0) (1:12) [backtracking split refutation 391,394,637]
637. $false {3} (1:0) [subsumption resolution 394,624]
624. cont(intersection(sK0,sK1),sK1) (2:5) [resolution 349,279]
349. p(intersection(sK0,sK1),intersection(sK0,sK1)) (1:7) [resolution 195,210]
394. ~cont(intersection(sK0,sK1),sK1) {3} (1:5) [backtracking splitting component]
391. ~eqdim(intersection(sK0,sK1),intersection(sK1,sK0)) | ~cont(intersection(sK0,sK1),sK0) | ~cont(intersection(sK0,sK1),sK1) (1:17) [resolution 198,277]
277. p(X2,intersection(X0,X1)) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:16) [cnf transformation 151]
198. ~p(intersection(sK0,sK1),intersection(sK1,sK0)) (0:7) [cnf transformation 178]
2073. ~leq(intersection(sK1,sK0),intersection(sK0,sK1)) (2:7) [backtracking split refutation 1180,1188,2072]
2072. $false {51} (1:0) [subsumption resolution 346,2068]
2068. ~cont(intersection(sK0,sK1),sK0) {51} (3:5) [backtracking split refutation 2052,2064,2065]
2065. $false {88} (7:0) [subsumption resolution 1270,2064]
1270. cont(intersection(sK0,sK1),sK1) (7:5) [backtracking split refutation 1262,1267,1268]
1268. $false {56} (2:0) [subsumption resolution 624,1267]
1267. ~cont(intersection(sK0,sK1),sK1) {56} (7:5) [backtracking splitting component]
1262. lt(intersection(sK0,sK1),sK1) | ~cont(intersection(sK0,sK1),sK1) {13} (7:10) [resolution 1252,261]
261. eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 134]
134. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[ennf transformation 80]
80. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 18]
18. ! [X0] : ! [X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[input]
1252. ~eqdim(intersection(sK0,sK1),sK1) {13} (6:5) [resolution 1200,258]
1200. ~eqdim(sK1,intersection(sK0,sK1)) {13} (5:5) [backtracking split refutation 1196,1198,1199]
1199. $false {53} (1:0) [subsumption resolution 305,1198]
305. cont(sK1,sK1) (1:3) [resolution 197,201]
201. zex(X0) | cont(X0,X0) (0:5) [cnf transformation 106]
106. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[ennf transformation 31]
31. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[input]
197. ~zex(sK1) (0:2) [cnf transformation 178]
1198. ~cont(sK1,sK1) {53} (5:3) [backtracking splitting component]
1196. ~eqdim(sK1,intersection(sK0,sK1)) | ~cont(sK1,sK1) {13} (5:8) [backtracking split refutation 1191,1194,1195]
1195. $false {13, 52} (4:0) [subsumption resolution 966,1194]
966. cont(sK1,sK0) {13} (4:3) [superposition 352,539]
539. sK1 = sK10(sK1,sK0) {13} (3:5) [backtracking splitting component]
352. cont(sK10(sK1,sK0),sK0) (2:5) [resolution 347,241]
241. ~c(X0,X1) | cont(sK10(X1,X0),X0) (0:8) [cnf transformation 189]
189. ! [X0,X1] : ((~c(X0,X1) | (cont(sK10(X1,X0),X0) & cont(sK10(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 122]
122. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 69]
69. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 68]
68. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 37]
37. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
347. c(sK0,sK1) (1:3) [resolution 195,212]
212. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 112]
112. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[ennf transformation 57]
57. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[flattening 11]
11. ! [X0] : ! [X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[input]
1194. ~cont(sK1,sK0) {52} (5:3) [backtracking splitting component]
1191. ~eqdim(sK1,intersection(sK0,sK1)) | ~cont(sK1,sK1) | ~cont(sK1,sK0) (5:11) [resolution 1141,277]
1141. ~p(sK1,intersection(sK0,sK1)) (4:5) [resolution 1091,247]
247. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 126]
126. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 125]
125. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 71]
71. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 38]
38. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
1091. ~cont(sK1,intersection(sK0,sK1)) (3:5) [backtracking split refutation 829,878,1090]
1090. $false {35} (2:0) [subsumption resolution 948,1082]
1082. ~leq(sK1,intersection(sK1,sK0)) {35} (4:5) [resolution 1074,251]
251. ~lt(X0,X1) | ~leq(X1,X0) (0:6) [cnf transformation 128]
128. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~leq(X0,X1) | leq(X1,X0) | lt(X0,X1)))[flattening 127]
127. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & ((~leq(X0,X1) | leq(X1,X0)) | lt(X0,X1)))[ennf transformation 72]
72. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[flattening 49]
49. ! [X0] : ! [X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~leq(X1,X0))) & (~(leq(X0,X1) & ~leq(X1,X0)) | lt(X0,X1)))[input]
1074. lt(intersection(sK1,sK0),sK1) {35} (3:5) [backtracking split refutation 920,925,1072]
1072. $false {35, 38} (2:0) [subsumption resolution 947,1071]
1071. leq(intersection(sK1,sK0),sK1) {38} (3:5) [backtracking split refutation 917,949,1069]
1069. $false {38, 41} (3:0) [subsumption resolution 925,1063]
1063. cont(intersection(sK1,sK0),sK1) {41} (6:5) [resolution 1057,254]
1057. ~zex(intersection(sK1,sK0)) {41} (5:4) [backtracking split refutation 1052,1054,1055]
1055. $false {41, 43} (4:0) [subsumption resolution 1045,1054]
1045. ~zex(intersection(sK1,sK0)) {41} (4:4) [resolution 949,255]
1054. zex(intersection(sK1,sK0)) {43} (5:4) [backtracking splitting component]
1052. ~zex(sK3(intersection(sK1,sK0))) | zex(intersection(sK1,sK0)) {41} (5:9) [resolution 1035,208]
1035. ~mindim(intersection(sK1,sK0)) {41} (4:4) [resolution 949,304]
304. leq(X1,sK1) | ~mindim(X1) (1:5) [resolution 197,207]
949. ~leq(intersection(sK1,sK0),sK1) {41} (3:5) [backtracking splitting component]
917. ~leq(sK1,intersection(sK1,sK0)) | ~leq(intersection(sK1,sK0),sK1) {35} (3:10) [resolution 906,246]
906. ~eqdim(intersection(sK1,sK0),sK1) {35} (2:5) [resolution 896,258]
896. ~eqdim(sK1,intersection(sK1,sK0)) {35} (1:5) [forward demodulation 892,878]
947. ~leq(intersection(sK1,sK0),sK1) {35} (2:5) [backtracking split refutation 904,912,946]
946. $false {35, 36} (1:0) [subsumption resolution 894,942]
942. ~cont(sK1,sK0) {36} (3:3) [backtracking split refutation 926,938,939]
939. $false {39} (1:0) [subsumption resolution 305,938]
938. ~cont(sK1,sK1) {39} (3:3) [backtracking splitting component]
926. ~cont(sK1,sK0) | ~cont(sK1,sK1) {36} (3:6) [resolution 912,293]
894. cont(sK1,sK0) {35} (1:3) [forward demodulation 346,878]
912. ~leq(sK1,intersection(sK1,sK0)) {36} (2:5) [backtracking splitting component]
904. ~leq(intersection(sK1,sK0),sK1) | ~leq(sK1,intersection(sK1,sK0)) {35} (2:10) [resolution 896,246]
925. ~cont(intersection(sK1,sK0),sK1) {38} (3:5) [backtracking splitting component]
920. lt(intersection(sK1,sK0),sK1) | ~cont(intersection(sK1,sK0),sK1) {35} (3:10) [resolution 906,261]
948. leq(sK1,intersection(sK1,sK0)) {35} (2:5) [backtracking split refutation 904,912,946]
878. intersection(sK0,sK1) = sK1 {35} (3:5) [backtracking splitting component]
829. ~cont(sK1,intersection(sK0,sK1)) | intersection(sK0,sK1) = sK1 (3:10) [resolution 624,268]
268. ~cont(X0,X1) | ~cont(X1,X0) | X0 = X1 (0:9) [cnf transformation 147]
147. ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1)[flattening 146]
146. ! [X0,X1] : ((~cont(X0,X1) | ~cont(X1,X0)) | X0 = X1)[ennf transformation 87]
87. ! [X0,X1] : (~(cont(X0,X1) & cont(X1,X0)) | X0 = X1)[flattening 32]
32. ! [X0] : ! [X1] : (~(cont(X0,X1) & cont(X1,X0)) | X0 = X1)[input]
2064. ~cont(intersection(sK0,sK1),sK1) {88} (3:5) [backtracking splitting component]
2052. ~cont(intersection(sK0,sK1),sK0) | ~cont(intersection(sK0,sK1),sK1) {51} (3:10) [resolution 1188,293]
1188. ~leq(intersection(sK0,sK1),intersection(sK1,sK0)) {51} (2:7) [backtracking splitting component]
1180. ~leq(intersection(sK1,sK0),intersection(sK0,sK1)) | ~leq(intersection(sK0,sK1),intersection(sK1,sK0)) (2:14) [resolution 892,246]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 711
Passive clauses: 1764
Generated clauses: 3689
Final active clauses: 454
Final passive clauses: 637
Input formulas: 53
Initial clauses: 103

Duplicate literals: 16
Bw subsumption resolutions: 93
Fw demodulations: 146
Condensations: 2

Simple tautologies: 26
Equational tautologies: 26
Forward subsumptions: 1489

Binary resolution: 2860
Forward superposition: 2
Backward superposition: 65

Backtracking splits: 123
Backtracking splits refuted: 83
Backtracking splits refuted at zero level: 10

Memory used: 895KB
Time elapsed: 0.168 s
------------------------------
% Success in time 2.098 s
