Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal17.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 516
Passive clauses: 1872
Generated clauses: 3616
Final active clauses: 516
Final passive clauses: 1294
Input formulas: 43
Initial clauses: 89

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 18
Fw subsumption resolutions: 37

Simple tautologies: 25
Forward subsumptions: 1702

Binary resolution: 3410
Factoring: 34

Backtracking splits: 20
Backtracking splits refuted: 4

Memory used: 767KB
Time elapsed: 0.427 s
------------------------------
% remaining time: 5992 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal17.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 899
Passive clauses: 4986
Generated clauses: 10609
Final active clauses: 899
Final passive clauses: 3860
Input formulas: 43
Initial clauses: 84

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 171
Fw subsumption resolutions: 924

Simple tautologies: 141
Equational tautologies: 7
Forward subsumptions: 4532

Binary resolution: 7425
Factoring: 111
Forward superposition: 1876
Backward superposition: 38

Unique components: 2

Memory used: 1791KB
Time elapsed: 1.421 s
------------------------------
% remaining time: 5976 next slice time: 32
lrs+11_14_bd=off:bs=off:bsr=on:cond=on:drc=off:fde=none:fsr=off:gsp=input_only:lcm=reverse:n=on:nwc=10:ptb=off:sfv=off:sgo=on:sio=off:sos=on:sp=reverse_arity:spl=backtracking:ssec=off:stl=81_32 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal17.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 2158
Passive clauses: 5313
Generated clauses: 19224
Final active clauses: 1681
Final passive clauses: 2356
Input formulas: 43
Initial clauses: 82

Unused predicate definitions: 2

Duplicate literals: 110
Bw subsumption resolutions: 230
Fw demodulations: 382
Condensations: 18

Simple tautologies: 85
Equational tautologies: 575
Forward subsumptions: 12087

Binary resolution: 16818
Factoring: 18
Forward superposition: 84
Backward superposition: 345

Backtracking splits: 426
Backtracking splits refuted: 187
Backtracking splits refuted at zero level: 13

Memory used: 3709KB
Time elapsed: 3.310 s
------------------------------
% remaining time: 5940 next slice time: 95
dis+2_4_bs=off:ep=on:n=on:nicw=on:nwc=1.5:ptb=off:sac=on:sio=off:spl=backtracking:ssec=off_95 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal17.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 2018
Passive clauses: 33935
Generated clauses: 73613
Final active clauses: 1980
Final passive clauses: 30539
Input formulas: 43
Initial clauses: 84

Unused predicate definitions: 2

Duplicate literals: 1251
Fw subsumption resolutions: 3788
Fw demodulations: 391
Bw demodulations: 1

Simple tautologies: 411
Equational tautologies: 15
Forward subsumptions: 35052

Binary resolution: 67564
Factoring: 38
Forward superposition: 578
Backward superposition: 26

Backtracking splits: 50
Backtracking splits refuted: 28
Backtracking splits refuted at zero level: 13

Memory used: 14328KB
Time elapsed: 9.613 s
------------------------------
% remaining time: 5842 next slice time: 4
lrs-1003_12_bms=on:bs=unit_only:bsr=on:cond=fast:drc=off:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=10:sfv=off:sgo=on:sio=off:sp=occurrence:spo=on:ssec=off:stl=81:updr=off_4 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal17.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 503
Passive clauses: 921
Generated clauses: 2207
Final active clauses: 501
Final passive clauses: 414
Input formulas: 43
Initial clauses: 88

Duplicate literals: 28
Fw subsumption resolutions: 31
Bw subsumption resolutions: 1
Condensations: 61

Simple tautologies: 29
Equational tautologies: 4
Forward subsumptions: 1110
Backward subsumptions: 12

Binary resolution: 1727
Factoring: 262

Split clauses: 7
Split components: 14
Unique components: 576
Introduced splitting names: 5

Memory used: 1279KB
Time elapsed: 0.501 s
------------------------------
% remaining time: 5836 next slice time: 22
lrs+1010_50_bs=off:drc=off:flr=on:lcm=predicate:n=on:nwc=1.2:sfv=off:sgo=on:stl=81_22 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal17.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1092
Passive clauses: 3685
Generated clauses: 18887
Final active clauses: 1092
Final passive clauses: 2593
Input formulas: 43
Initial clauses: 82

Unused predicate definitions: 2

Duplicate literals: 1137
Fw subsumption resolutions: 1558
Fw literal rewrites: 795

Simple tautologies: 418
Equational tautologies: 3
Forward subsumptions: 11311

Binary resolution: 14359
Factoring: 1
Forward superposition: 385
Backward superposition: 605
Self superposition: 2

Unique components: 2

Memory used: 2814KB
Time elapsed: 2.306 s
------------------------------
% remaining time: 5812 next slice time: 41
dis-1010_5_bd=off:bs=off:cond=fast:ep=on:fde=none:lcm=predicate:n=on:nicw=on:nwc=1.3:ptb=off:sac=on:sgo=on:sio=off:sos=on:spl=backtracking:ssec=off:updr=off_41 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal17.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 2526
Passive clauses: 10194
Generated clauses: 37555
Final active clauses: 2522
Final passive clauses: 7029
Input formulas: 43
Initial clauses: 88

Duplicate literals: 215
Fw subsumption resolutions: 659
Condensations: 77

Simple tautologies: 126
Forward subsumptions: 26975

Binary resolution: 36149
Factoring: 432

Backtracking splits: 22
Backtracking splits refuted: 4
Backtracking splits refuted at zero level: 4

Memory used: 5500KB
Time elapsed: 4.208 s
------------------------------
% remaining time: 5768 next slice time: 237
lrs+1_2:1_bs=off:ep=on:gsp=input_only:lcm=predicate:n=on:nicw=on:nwc=1.1:ptb=off:sac=on:sgo=on:sio=off:sos=on:sp=occurrence:spl=backtracking:spo=on:ssec=off:stl=81:updr=off_237 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal17.all
Refutation not found with incomplete strategy!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation not found, incomplete strategy

Active clauses: 88
Passive clauses: 2
Generated clauses: 2
Final active clauses: 88
Input formulas: 43
Initial clauses: 88

Memory used: 127KB
Time elapsed: 0.006 s
------------------------------
% remaining time: 5767 next slice time: 311
lrs+11_40_bd=off:bs=unit_only:bsr=unit_only:drc=off:flr=on:lcm=predicate:n=on:ptb=off:sac=on:sfv=off:spl=backtracking:ssec=off:stl=81:updr=off_311 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal17.all
Refutation found. Thanks to Tanya!
192068. $false (3:0) [subsumption resolution 192052,214]
214. ~zex(intersection(sK0,sK1)) (0:4) [cnf transformation 167]
167. ~zex(intersection(sK0,sK1)) & ~p(intersection(sK0,sK1),intersection(sK0,intersection(sK0,sK1)))[skolemisation 119]
119. ? [X0,X1] : (~zex(intersection(X0,X1)) & ~p(intersection(X0,X1),intersection(X0,intersection(X0,X1))))[ennf transformation 44]
44. ~! [X0,X1] : (~zex(intersection(X0,X1)) => p(intersection(X0,X1),intersection(X0,intersection(X0,X1))))[flattening 1]
1. ~! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => p(intersection(X0,X1),intersection(X0,intersection(X0,X1))))[negated conjecture]
192052. zex(intersection(sK0,sK1)) (3:4) [resolution 87542,215]
215. ~p(intersection(sK0,sK1),intersection(sK0,intersection(sK0,sK1))) (0:9) [cnf transformation 167]
87542. p(intersection(X4,X5),intersection(X4,intersection(X4,X5))) | zex(intersection(X4,X5)) (2:13) [subsumption resolution 87532,20600]
20600. eqdim(intersection(X7,X8),intersection(X7,intersection(X7,X8))) (3:9) [forward literal rewriting 20599,233]
233. ~eqdim(X0,X1) | eqdim(X1,X0) (0:6) [cnf transformation 130]
130. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 60]
60. ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))[flattening 59]
59. ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))[rectify 33]
33. ! [X67] : ! [X68] : (eqdim(X67,X68) => eqdim(X68,X67))[input]
20599. eqdim(intersection(X7,intersection(X7,X8)),intersection(X7,X8)) (3:9) [subsumption resolution 20588,8310]
8310. geq(X14,intersection(X13,X14)) (4:5) [subsumption resolution 8279,305]
305. geq(X1,X0) | ~lt(X0,X1) (0:6) [forward literal rewriting 280,245]
245. ~leq(X1,X0) | geq(X0,X1) (0:6) [cnf transformation 179]
179. ! [X0,X1] : ((~geq(X0,X1) | leq(X1,X0)) & (~leq(X1,X0) | geq(X0,X1)))[nnf transformation 79]
79. ! [X0,X1] : (geq(X0,X1) <=> leq(X1,X0))[flattening 78]
78. ! [X0] : ! [X1] : (geq(X0,X1) <=> leq(X1,X0))[rectify 5]
5. ! [X10] : ! [X11] : (geq(X10,X11) <=> leq(X11,X10))[input]
280. leq(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 207]
207. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | lt(X0,X1)))[flattening 206]
206. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | lt(X0,X1)))[nnf transformation 148]
148. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[ennf transformation 97]
97. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[flattening 96]
96. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[rectify 6]
6. ! [X12] : ! [X13] : (lt(X12,X13) <=> (leq(X12,X13) & ~eqdim(X12,X13)))[input]
8279. lt(intersection(X13,X14),X14) | geq(X14,intersection(X13,X14)) (4:10) [resolution 1564,302]
302. ~eqdim(X0,X1) | geq(X0,X1) (0:6) [forward literal rewriting 278,245]
278. leq(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 205]
205. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 204]
204. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[nnf transformation 95]
95. ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[flattening 94]
94. ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[rectify 41]
41. ! [X86] : ! [X87] : (eqdim(X86,X87) <=> (leq(X86,X87) & leq(X87,X86)))[input]
1564. eqdim(X10,intersection(X9,X10)) | lt(intersection(X9,X10),X10) (3:10) [forward literal rewriting 1563,233]
1563. lt(intersection(X9,X10),X10) | eqdim(intersection(X9,X10),X10) (3:10) [subsumption resolution 1557,555]
555. ~zex(X3) | eqdim(X3,X4) | lt(X3,X4) (2:8) [resolution 304,320]
320. geq(X3,X4) | ~zex(X4) (1:5) [resolution 245,228]
228. leq(X0,X1) | ~zex(X0) (0:5) [cnf transformation 124]
124. ! [X0,X1] : (~zex(X0) | leq(X0,X1))[ennf transformation 52]
52. ! [X0,X1] : (zex(X0) => leq(X0,X1))[flattening 51]
51. ! [X0] : ! [X1] : (zex(X0) => leq(X0,X1))[rectify 16]
16. ! [X33] : ! [X34] : (zex(X33) => leq(X33,X34))[input]
304. ~geq(X1,X0) | lt(X0,X1) | eqdim(X0,X1) (0:9) [forward literal rewriting 282,244]
244. leq(X1,X0) | ~geq(X0,X1) (0:6) [cnf transformation 179]
282. lt(X0,X1) | eqdim(X0,X1) | ~leq(X0,X1) (0:9) [cnf transformation 207]
1557. zex(intersection(X9,X10)) | lt(intersection(X9,X10),X10) | eqdim(intersection(X9,X10),X10) (3:14) [resolution 455,231]
231. ~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1) (0:9) [cnf transformation 127]
127. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 126]
126. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 56]
56. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 55]
55. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 28]
28. ! [X58] : ! [X59] : (cont(X58,X59) => (lt(X58,X59) | eqdim(X58,X59)))[input]
455. cont(intersection(X4,X5),X5) | zex(intersection(X4,X5)) (2:9) [resolution 451,294]
294. ~p(X2,intersection(X0,X1)) | cont(X2,X1) (0:8) [cnf transformation 210]
210. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 209]
209. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[nnf transformation 115]
115. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[flattening 114]
114. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) <=> p(X2,intersection(X0,X1)))[rectify 27]
27. ! [X55] : ! [X56] : ! [X57] : ((cont(X57,X55) & cont(X57,X56) & eqdim(X57,intersection(X55,X56))) <=> p(X57,intersection(X55,X56)))[input]
451. p(X0,X0) | zex(X0) (1:5) [forward literal rewriting 450,225]
225. cont(X0,X0) | zex(X0) (0:5) [cnf transformation 176]
176. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[nnf transformation 122]
122. ! [X0] : (~zex(X0) <=> cont(X0,X0))[ennf transformation 49]
49. ! [X0] : (~zex(X0) <=> cont(X0,X0))[rectify 17]
17. ! [X35] : (~zex(X35) <=> cont(X35,X35))[input]
450. p(X0,X0) | ~cont(X0,X0) (1:6) [resolution 276,217]
217. eqdim(X0,X0) (0:3) [cnf transformation 46]
46. ! [X0] : eqdim(X0,X0)[rectify 32]
32. ! [X66] : eqdim(X66,X66)[input]
276. ~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 203]
203. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 202]
202. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[nnf transformation 93]
93. ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[flattening 92]
92. ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[rectify 7]
7. ! [X14] : ! [X15] : (p(X14,X15) <=> (cont(X14,X15) & eqdim(X14,X15)))[input]
20588. ~geq(intersection(X7,X8),intersection(X7,intersection(X7,X8))) | eqdim(intersection(X7,intersection(X7,X8)),intersection(X7,X8)) (3:18) [resolution 20564,301]
301. ~geq(X0,X1) | ~geq(X1,X0) | eqdim(X0,X1) (0:9) [forward literal rewriting 300,244]
300. ~geq(X0,X1) | eqdim(X0,X1) | ~leq(X0,X1) (0:9) [forward literal rewriting 279,244]
279. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 205]
20564. geq(intersection(X4,intersection(X4,X5)),intersection(X4,X5)) (2:9) [subsumption resolution 20559,320]
20559. geq(intersection(X4,intersection(X4,X5)),intersection(X4,X5)) | zex(intersection(X4,X5)) (2:13) [resolution 3291,225]
3291. ~cont(intersection(X2,X3),X4) | geq(intersection(X2,X4),intersection(X2,X3)) (1:12) [subsumption resolution 3284,320]
3284. ~cont(intersection(X2,X3),X4) | geq(intersection(X2,X4),intersection(X2,X3)) | zex(intersection(X2,X3)) (1:16) [resolution 306,241]
241. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 143]
143. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 75]
75. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 74]
74. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 25]
25. ! [X50] : ! [X51] : (~zex(intersection(X50,X51)) => cont(intersection(X50,X51),X50))[input]
306. ~cont(X2,X0) | ~cont(X2,X1) | geq(intersection(X0,X1),X2) (0:11) [forward literal rewriting 288,245]
288. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 157]
157. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 156]
156. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 107]
107. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[flattening 106]
106. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[rectify 26]
26. ! [X52] : ! [X53] : ! [X54] : ((cont(X54,X52) & cont(X54,X53)) => leq(X54,intersection(X52,X53)))[input]
87532. ~eqdim(intersection(X4,X5),intersection(X4,intersection(X4,X5))) | p(intersection(X4,X5),intersection(X4,intersection(X4,X5))) | zex(intersection(X4,X5)) (2:22) [resolution 12496,225]
12496. ~cont(intersection(X2,X3),X4) | ~eqdim(intersection(X2,X3),intersection(X2,X4)) | p(intersection(X2,X3),intersection(X2,X4)) (1:19) [subsumption resolution 12485,230]
230. ~zex(X0) | ~cont(X0,X1) (0:5) [cnf transformation 125]
125. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 54]
54. ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))[flattening 53]
53. ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))[rectify 20]
20. ! [X39] : ! [X40] : (zex(X39) => (~cont(X40,X39) & ~cont(X39,X40)))[input]
12485. ~eqdim(intersection(X2,X3),intersection(X2,X4)) | ~cont(intersection(X2,X3),X4) | p(intersection(X2,X3),intersection(X2,X4)) | zex(intersection(X2,X3)) (1:23) [resolution 292,241]
292. ~cont(X2,X0) | ~eqdim(X2,intersection(X0,X1)) | ~cont(X2,X1) | p(X2,intersection(X0,X1)) (0:16) [cnf transformation 210]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 2995
Passive clauses: 17157
Generated clauses: 173309
Final active clauses: 2934
Final passive clauses: 7477
Input formulas: 43
Initial clauses: 86
Discarded non-redundant clauses: 27064

Duplicate literals: 764
Fw subsumption resolutions: 5756
Bw subsumption resolutions: 69
Fw demodulations: 14785
Fw literal rewrites: 13620

Simple tautologies: 539
Equational tautologies: 178
Forward subsumptions: 120421
Backward subsumptions: 429

Binary resolution: 131443
Factoring: 18
Forward superposition: 3390
Backward superposition: 3408
Self superposition: 2
Equality factoring: 2

Memory used: 18038KB
Time elapsed: 25.376 s
------------------------------
% Success in time 48.732 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 23:35:25 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal17.all.tptp"
Input read from 
============================ end of footer ===========================
