Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal16.all
Refutation found. Thanks to Tanya!
498. $false (4:0) [subsumption resolution 497,213]
213. zex(intersection(sK0,sK0)) (0:4) [cnf transformation 166]
166. zex(intersection(sK0,sK0)) & ~zex(sK0)[skolemisation 118]
118. ? [X0] : (zex(intersection(X0,X0)) & ~zex(X0))[ennf transformation 1]
1. ~! [X0] : (zex(intersection(X0,X0)) => zex(X0))[negated conjecture]
497. ~zex(intersection(sK0,sK0)) (4:4) [resolution 493,283]
283. ~c(X0,X1) | ~zex(intersection(X0,X1)) (0:7) [cnf transformation 207]
207. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[nnf transformation 148]
148. ! [X0,X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[ennf transformation 98]
98. ! [X0,X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[flattening 97]
97. ! [X0] : ! [X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[rectify 24]
24. ! [X48] : ! [X49] : (~c(X48,X49) <=> zex(intersection(X48,X49)))[input]
493. c(sK0,sK0) (3:3) [resolution 470,309]
309. cont(sK0,sK0) (1:3) [resolution 224,214]
214. ~zex(sK0) (0:2) [cnf transformation 166]
224. zex(X0) | cont(X0,X0) (0:5) [cnf transformation 175]
175. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[nnf transformation 121]
121. ! [X0] : (~zex(X0) <=> cont(X0,X0))[ennf transformation 17]
17. ! [X0] : (~zex(X0) <=> cont(X0,X0))[input]
470. ~cont(sK0,X0) | c(X0,sK0) (2:6) [resolution 272,309]
272. ~cont(X3,X1) | ~cont(X3,X0) | c(X0,X1) (0:9) [cnf transformation 200]
200. ! [X0,X1] : ((~c(X0,X1) | (cont(sK10(X1,X0),X0) & cont(sK10(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 199]
199. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[rectify 198]
198. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | c(X0,X1)))[nnf transformation 90]
90. ! [X0,X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))[flattening 89]
89. ! [X0] : ! [X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))[rectify 2]
2. ! [X1] : ! [X2] : (c(X1,X2) <=> ? [X3] : (cont(X3,X1) & cont(X3,X2)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 125
Passive clauses: 216
Generated clauses: 280
Final active clauses: 125
Final passive clauses: 80
Input formulas: 43
Initial clauses: 89

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 3
Fw subsumption resolutions: 3

Simple tautologies: 7
Forward subsumptions: 61

Binary resolution: 179
Factoring: 6

Memory used: 255KB
Time elapsed: 0.039 s
------------------------------
% Success in time 0.105 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 23:34:32 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal16.all.tptp"
Input read from 
============================ end of footer ===========================
