Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal03.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 746
Passive clauses: 3581
Generated clauses: 12169
Final active clauses: 746
Final passive clauses: 2681
Input formulas: 44
Initial clauses: 92

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 68
Fw subsumption resolutions: 112

Simple tautologies: 38
Forward subsumptions: 8477

Binary resolution: 11934
Factoring: 53

Backtracking splits: 47

Memory used: 1279KB
Time elapsed: 0.403 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal03.all
Refutation found. Thanks to Tanya!
616. $false (4:0) [subsumption resolution 615,275]
275. c(sK0,sK1) (1:3) [resolution 176,163]
163. ~zex(intersection(sK0,sK1)) (0:4) [cnf transformation 146]
146. zex(intersection(sK1,sK0)) & ~zex(intersection(sK0,sK1))[skolemisation 87]
87. ? [X0,X1] : (zex(intersection(X1,X0)) & ~zex(intersection(X0,X1)))[ennf transformation 45]
45. ~! [X0,X1] : (~zex(intersection(X1,X0)) | zex(intersection(X0,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : (~zex(intersection(X1,X0)) | zex(intersection(X0,X1)))[negated conjecture]
176. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 94]
94. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[ennf transformation 48]
48. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[flattening 2]
2. ! [X0] : ! [X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[input]
615. ~c(sK0,sK1) (4:3) [duplicate literal removal 607]
607. ~c(sK0,sK1) | ~c(sK0,sK1) (4:6) [resolution 496,205]
205. cont(sK10(X1,X0),X0) | ~c(X0,X1) (0:8) [cnf transformation 157]
157. ! [X0,X1] : ((~c(X0,X1) | (cont(sK10(X1,X0),X0) & cont(sK10(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 104]
104. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 60]
60. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 59]
59. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 28]
28. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
496. ~cont(sK10(sK1,X5),sK0) | ~c(X5,sK1) (3:8) [resolution 485,206]
206. cont(sK10(X1,X0),X1) | ~c(X0,X1) (0:8) [cnf transformation 157]
485. ~cont(X0,sK1) | ~cont(X0,sK0) (2:6) [resolution 207,276]
276. ~c(sK1,sK0) (1:3) [resolution 177,162]
162. zex(intersection(sK1,sK0)) (0:4) [cnf transformation 146]
177. ~zex(intersection(X0,X1)) | ~c(X0,X1) (0:7) [cnf transformation 94]
207. c(X0,X1) | ~cont(X3,X1) | ~cont(X3,X0) (0:9) [cnf transformation 157]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 110
Passive clauses: 229
Generated clauses: 447
Final active clauses: 110
Final passive clauses: 100
Input formulas: 44
Initial clauses: 91

Selected by SInE selection: 44
SInE iterations: 1

Duplicate literals: 29
Fw subsumption resolutions: 29

Simple tautologies: 29
Equational tautologies: 1
Forward subsumptions: 143

Binary resolution: 147
Factoring: 1
Forward superposition: 108
Backward superposition: 44

Unique components: 2

Memory used: 255KB
Time elapsed: 0.019 s
------------------------------
% Success in time 0.513 s
