Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal13.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 561
Passive clauses: 2143
Generated clauses: 4368
Final active clauses: 561
Final passive clauses: 1511
Input formulas: 43
Initial clauses: 90

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 32
Fw subsumption resolutions: 38

Simple tautologies: 29
Forward subsumptions: 2184

Binary resolution: 4160
Factoring: 40

Backtracking splits: 10

Memory used: 895KB
Time elapsed: 0.402 s
------------------------------
% remaining time: 5994 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal13.all
Refutation found. Thanks to Tanya!
917. $false (6:0) [subsumption resolution 912,213]
213. ~zex(sK0) (0:2) [cnf transformation 168]
168. ~zex(sK0) & ~zex(intersection(sK0,sK0)) & ~p(intersection(sK0,sK0),sK0)[skolemisation 120]
120. ? [X0] : (~zex(X0) & ~zex(intersection(X0,X0)) & ~p(intersection(X0,X0),X0))[flattening 119]
119. ? [X0] : ((~zex(X0) & ~zex(intersection(X0,X0))) & ~p(intersection(X0,X0),X0))[ennf transformation 1]
1. ~! [X0] : ((~zex(X0) & ~zex(intersection(X0,X0))) => p(intersection(X0,X0),X0))[negated conjecture]
912. zex(sK0) (6:2) [resolution 911,225]
225. cont(X0,X0) | zex(X0) (0:5) [cnf transformation 177]
177. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[nnf transformation 123]
123. ! [X0] : (~zex(X0) <=> cont(X0,X0))[ennf transformation 48]
48. ! [X0] : (~zex(X0) <=> cont(X0,X0))[rectify 17]
17. ! [X33] : (~zex(X33) <=> cont(X33,X33))[input]
911. ~cont(sK0,sK0) (5:3) [subsumption resolution 904,214]
214. ~zex(intersection(sK0,sK0)) (0:4) [cnf transformation 168]
904. ~cont(sK0,sK0) | zex(intersection(sK0,sK0)) (5:7) [resolution 903,241]
241. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 144]
144. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 75]
75. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 74]
74. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 25]
25. ! [X50] : ! [X51] : (~zex(intersection(X50,X51)) => cont(intersection(X50,X51),X50))[input]
903. ~cont(intersection(sK0,sK0),sK0) | ~cont(sK0,sK0) (4:8) [duplicate literal removal 899]
899. ~cont(sK0,sK0) | ~cont(sK0,sK0) | ~cont(intersection(sK0,sK0),sK0) (4:11) [resolution 284,664]
664. ~leq(sK0,intersection(sK0,sK0)) | ~cont(intersection(sK0,sK0),sK0) (3:10) [subsumption resolution 655,514]
514. ~lt(X1,X0) | ~cont(X0,X1) (2:6) [subsumption resolution 511,235]
235. ~lt(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 133]
133. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[ennf transformation 63]
63. ! [X0,X1] : (lt(X0,X1) => ~lt(X1,X0))[flattening 62]
62. ! [X0] : ! [X1] : (lt(X0,X1) => ~lt(X1,X0))[rectify 30]
30. ! [X61] : ! [X62] : (lt(X61,X62) => ~lt(X62,X61))[input]
511. lt(X0,X1) | ~cont(X0,X1) | ~lt(X1,X0) (2:9) [resolution 231,309]
309. ~eqdim(X2,X1) | ~lt(X1,X2) (1:6) [resolution 234,233]
233. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 131]
131. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 59]
59. ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))[flattening 58]
58. ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))[rectify 33]
33. ! [X67] : ! [X68] : (eqdim(X67,X68) => eqdim(X68,X67))[input]
234. ~eqdim(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 132]
132. ! [X0,X1] : (~lt(X0,X1) | ~eqdim(X0,X1))[ennf transformation 61]
61. ! [X0,X1] : (lt(X0,X1) => ~eqdim(X0,X1))[flattening 60]
60. ! [X0] : ! [X1] : (lt(X0,X1) => ~eqdim(X0,X1))[rectify 35]
35. ! [X72] : ! [X73] : (lt(X72,X73) => ~eqdim(X72,X73))[input]
231. eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 128]
128. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 127]
127. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 55]
55. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 54]
54. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 28]
28. ! [X58] : ! [X59] : (cont(X58,X59) => (lt(X58,X59) | eqdim(X58,X59)))[input]
655. ~cont(intersection(sK0,sK0),sK0) | lt(sK0,intersection(sK0,sK0)) | ~leq(sK0,intersection(sK0,sK0)) (3:15) [resolution 644,278]
278. eqdim(X0,X1) | lt(X0,X1) | ~leq(X0,X1) (0:9) [cnf transformation 206]
206. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | lt(X0,X1)))[flattening 205]
205. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | lt(X0,X1)))[nnf transformation 149]
149. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[ennf transformation 97]
97. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[flattening 96]
96. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[rectify 6]
6. ! [X11] : ! [X12] : (lt(X11,X12) <=> (leq(X11,X12) & ~eqdim(X11,X12)))[input]
644. ~eqdim(sK0,intersection(sK0,sK0)) | ~cont(intersection(sK0,sK0),sK0) (2:10) [resolution 577,233]
577. ~eqdim(intersection(sK0,sK0),sK0) | ~cont(intersection(sK0,sK0),sK0) (1:10) [resolution 272,215]
215. ~p(intersection(sK0,sK0),sK0) (0:5) [cnf transformation 168]
272. p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 202]
202. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 201]
201. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[nnf transformation 93]
93. ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[flattening 92]
92. ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[rectify 7]
7. ! [X13] : ! [X14] : (p(X13,X14) <=> (cont(X13,X14) & eqdim(X13,X14)))[input]
284. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 158]
158. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 157]
157. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 107]
107. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[flattening 106]
106. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[rectify 26]
26. ! [X52] : ! [X53] : ! [X54] : ((cont(X54,X52) & cont(X54,X53)) => leq(X54,intersection(X52,X53)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 147
Passive clauses: 328
Generated clauses: 699
Final active clauses: 147
Final passive clauses: 168
Input formulas: 43
Initial clauses: 85

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 28
Fw subsumption resolutions: 58

Simple tautologies: 25
Equational tautologies: 1
Forward subsumptions: 267

Binary resolution: 325
Factoring: 3
Forward superposition: 170
Backward superposition: 30

Unique components: 3

Memory used: 255KB
Time elapsed: 0.032 s
------------------------------
% Success in time 0.782 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 23:34:11 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal13.all.tptp"
Input read from 
============================ end of footer ===========================
