Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2235
Passive clauses: 22752
Generated clauses: 37509
Input formulas: 69
Initial clauses: 132

Selected by SInE selection: 69
SInE iterations: 1

Duplicate literals: 134
Fw subsumption resolutions: 2734

Simple tautologies: 71
Forward subsumptions: 10591

Binary resolution: 32503
Factoring: 26

Backtracking splits: 1523
Backtracking splits refuted: 141
Backtracking splits refuted at zero level: 37

Memory used: 7419KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f2260,plain,(
  $false),
  inference(subsumption_resolution,[],[f2255,f353])).
fof(f353,plain,(
  ~zex(sK0)),
  inference(cnf_transformation,[],[f288])).
fof(f288,plain,(
  eqdim(sK0,sK1) & ~zex(sK0) & ~zex(sK1) & (cont(sum(sK0,sK1),sK2) & ~cont(sK0,sK2))),
  inference(skolemisation,[status(esa)],[f199])).
fof(f199,plain,(
  ? [X0,X1] : (eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & ? [X2] : (cont(sum(X0,X1),X2) & ~cont(X0,X2)))),
  inference(flattening,[],[f198])).
fof(f198,plain,(
  ? [X0,X1] : ((eqdim(X0,X1) & ~zex(X0) & ~zex(X1)) & ? [X2] : (cont(sum(X0,X1),X2) & ~cont(X0,X2)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : ((eqdim(X0,X1) & ~zex(X0) & ~zex(X1)) => ! [X2] : (cont(sum(X0,X1),X2) => cont(X0,X2)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1] : ((eqdim(X0,X1) & ~zex(X0) & ~zex(X1)) => ! [X2] : (cont(sum(X0,X1),X2) => cont(X0,X2)))),
  inference(rectify,[],[f12])).
fof(f12,negated_conjecture,(
  ~! [X28] : ! [X29] : ((eqdim(X28,X29) & ~zex(X28) & ~zex(X29)) => ! [X30] : (cont(sum(X28,X29),X30) => cont(X28,X30)))),
  file(unknown,unknown)).
fof(f2255,plain,(
  zex(sK0)),
  inference(resolution,[],[f2254,f370])).
fof(f370,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X163] : (~zex(X163) <=> cont(X163,X163))),
  file(unknown,unknown)).
fof(f2254,plain,(
  ~cont(sK0,sK0)),
  inference(subsumption_resolution,[],[f2248,f352])).
fof(f352,plain,(
  eqdim(sK0,sK1)),
  inference(cnf_transformation,[],[f288])).
fof(f2248,plain,(
  ~cont(sK0,sK0) | ~eqdim(sK0,sK1)),
  inference(resolution,[],[f2231,f991])).
fof(f991,plain,(
  ( ! [X6,X5] : (eqdim(X5,X6) | ~eqdim(X6,X5)) )),
  inference(subsumption_resolution,[],[f984,f445])).
fof(f445,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f341])).
fof(f341,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X105] : ! [X106] : (eqdim(X105,X106) <=> (~less(X105,X106) & ~less(X106,X105)))),
  file(unknown,unknown)).
fof(f984,plain,(
  ( ! [X6,X5] : (eqdim(X5,X6) | less(X5,X6) | ~eqdim(X6,X5)) )),
  inference(resolution,[],[f446,f444])).
fof(f444,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f342])).
fof(f446,plain,(
  ( ! [X0,X1] : (less(X1,X0) | eqdim(X0,X1) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f342])).
fof(f2231,plain,(
  ~eqdim(sK1,sK0) | ~cont(sK0,sK0)),
  inference(resolution,[],[f2215,f434])).
fof(f434,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f331])).
fof(f331,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X107] : ! [X108] : (leq(X107,X108) <=> (less(X107,X108) | eqdim(X107,X108)))),
  file(unknown,unknown)).
fof(f2215,plain,(
  ~leq(sK1,sK0) | ~cont(sK0,sK0)),
  inference(resolution,[],[f1357,f355])).
fof(f355,plain,(
  cont(sum(sK0,sK1),sK2)),
  inference(cnf_transformation,[],[f288])).
fof(f1357,plain,(
  ( ! [X6,X7] : (~cont(sum(X7,X6),sK2) | ~leq(X6,X7) | ~cont(sK0,X7)) )),
  inference(superposition,[],[f1282,f373])).
fof(f373,plain,(
  ( ! [X0,X1] : (sum(X0,X1) = sum(X1,X0)) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : sum(X0,X1) = sum(X1,X0)),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X31] : ! [X32] : sum(X31,X32) = sum(X32,X31)),
  file(unknown,unknown)).
fof(f1282,plain,(
  ( ! [X14,X13] : (~cont(sum(X14,X13),sK2) | ~leq(X14,X13) | ~cont(sK0,X13)) )),
  inference(resolution,[],[f461,f1039])).
fof(f1039,plain,(
  ( ! [X12] : (~cont(sK0,X12) | ~cont(X12,sK2)) )),
  inference(resolution,[],[f472,f356])).
fof(f356,plain,(
  ~cont(sK0,sK2)),
  inference(cnf_transformation,[],[f288])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))),
  inference(ennf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X166] : ! [X167] : ! [X168] : ((cont(X166,X167) & cont(X167,X168)) => cont(X166,X168))),
  file(unknown,unknown)).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (cont(X2,sum(X0,X1)) | ~cont(X2,X1) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X1) | cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X1)) | cont(X2,sum(X0,X1)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & cont(X2,X1)) => cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & cont(X2,X1)) => cont(X2,sum(X0,X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X35] : ! [X36] : ! [X37] : ((leq(X35,X36) & cont(X37,X36)) => cont(X37,sum(X35,X36)))),
  file(unknown,unknown)).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 300
Passive clauses: 799
Generated clauses: 1901
Final active clauses: 300
Final passive clauses: 462
Input formulas: 69
Initial clauses: 131

Selected by SInE selection: 69
SInE iterations: 1

Duplicate literals: 52
Fw subsumption resolutions: 153

Simple tautologies: 84
Equational tautologies: 1
Forward subsumptions: 846

Binary resolution: 710
Factoring: 9
Forward superposition: 724
Backward superposition: 124

Unique components: 5

Memory used: 511KB
Time elapsed: 0.023 s
------------------------------
% Success in time 0.469 s
