Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_down_sum_theorems_goal18.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1313
Passive clauses: 10483
Generated clauses: 15042
Final active clauses: 1267
Final passive clauses: 8869
Input formulas: 108
Initial clauses: 165

Selected by SInE selection: 108
SInE iterations: 1

Duplicate literals: 47
Fw subsumption resolutions: 328

Simple tautologies: 49
Forward subsumptions: 3887

Binary resolution: 14095
Factoring: 11

Backtracking splits: 391
Backtracking splits refuted: 20
Backtracking splits refuted at zero level: 3

Memory used: 3198KB
Time elapsed: 0.402 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_down_sum_theorems_goal18.all
Refutation found. Thanks to Tanya!
13365. $false (5:0) [subsumption resolution 13359,372]
372. ~zex(sK0) (0:2) [cnf transformation 351]
351. eqdim(sK0,sK1) & ~zex(sK0) & ~zex(sK1) & cont(sum(sK0,sK1),sK2) & ~zex(sum(sK0,sK1)) & ~cont(sK0,sK2)[skolemisation 207]
207. ? [X0,X1,X2] : (eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & cont(sum(X0,X1),X2) & ~zex(sum(X0,X1)) & ~cont(X0,X2))[flattening 206]
206. ? [X0,X1,X2] : ((eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & cont(sum(X0,X1),X2) & ~zex(sum(X0,X1))) & ~cont(X0,X2))[ennf transformation 109]
109. ~! [X0,X1,X2] : (~(eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & cont(sum(X0,X1),X2) & ~zex(sum(X0,X1))) | cont(X0,X2))[flattening 1]
1. ~! [X0] : ! [X1] : ! [X2] : (~(eqdim(X0,X1) & ~zex(X0) & ~zex(X1) & cont(sum(X0,X1),X2) & ~zex(sum(X0,X1))) | cont(X0,X2))[negated conjecture]
13359. zex(sK0) (5:2) [resolution 13308,379]
379. cont(X0,X0) | zex(X0) (0:5) [cnf transformation 208]
208. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[ennf transformation 79]
79. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[input]
13308. ~cont(sK0,sK0) (4:3) [subsumption resolution 13297,371]
371. eqdim(sK0,sK1) (0:3) [cnf transformation 351]
13297. ~cont(sK0,sK0) | ~eqdim(sK0,sK1) (4:6) [resolution 12898,417]
417. leq(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 234]
234. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 233]
233. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[ennf transformation 122]
122. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[flattening 97]
97. ! [X0] : ! [X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~(leq(X0,X1) & leq(X1,X0)) | eqdim(X0,X1)))[input]
12898. ~leq(sK1,sK0) | ~cont(sK0,sK0) (3:6) [resolution 5012,374]
374. cont(sum(sK0,sK1),sK2) (0:5) [cnf transformation 351]
5012. ~cont(sum(X29,X30),sK2) | ~leq(X30,X29) | ~cont(sK0,X29) (2:11) [resolution 1686,1280]
1280. ~cont(sK0,X12) | ~cont(X12,sK2) (1:6) [resolution 520,376]
376. ~cont(sK0,sK2) (0:3) [cnf transformation 351]
520. cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1) (0:9) [cnf transformation 333]
333. ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))[flattening 332]
332. ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))[ennf transformation 197]
197. ! [X0,X1,X2] : (~(cont(X0,X1) & cont(X1,X2)) | cont(X0,X2))[flattening 81]
81. ! [X0] : ! [X1] : ! [X2] : (~(cont(X0,X1) & cont(X1,X2)) | cont(X0,X2))[input]
1686. cont(X14,sum(X13,X12)) | ~cont(X14,X13) | ~leq(X12,X13) (1:11) [superposition 513,401]
401. sum(X0,X1) = sum(X1,X0) (0:7) [cnf transformation 114]
114. ! [X0,X1] : sum(X0,X1) = sum(X1,X0)[flattening 11]
11. ! [X0] : ! [X1] : sum(X0,X1) = sum(X1,X0)[input]
513. cont(X2,sum(X0,X1)) | ~cont(X2,X1) | ~leq(X0,X1) (0:11) [cnf transformation 319]
319. ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,X1) | cont(X2,sum(X0,X1)))[flattening 318]
318. ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,X1)) | cont(X2,sum(X0,X1)))[ennf transformation 190]
190. ! [X0,X1,X2] : (~(leq(X0,X1) & cont(X2,X1)) | cont(X2,sum(X0,X1)))[flattening 13]
13. ! [X0] : ! [X1] : ! [X2] : (~(leq(X0,X1) & cont(X2,X1)) | cont(X2,sum(X0,X1)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 926
Passive clauses: 5688
Generated clauses: 12980
Final active clauses: 926
Final passive clauses: 4589
Input formulas: 108
Initial clauses: 164

Selected by SInE selection: 108
SInE iterations: 1

Duplicate literals: 219
Fw subsumption resolutions: 885

Simple tautologies: 604
Equational tautologies: 7
Forward subsumptions: 5749

Binary resolution: 8482
Factoring: 39
Forward superposition: 3059
Backward superposition: 137

Unique components: 6

Memory used: 2302KB
Time elapsed: 0.508 s
------------------------------
% Success in time 0.967 s
