Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 1199 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:nwc=2.0:nicw=on:ptb=off:ssec=off:ss=included:st=5.0:sio=off:spl=backtracking:updr=off_1 on unknown
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit
Termination phase: Saturation

Active clauses: 2096
Passive clauses: 20999
Generated clauses: 34003
Input formulas: 38
Initial clauses: 88

Selected by SInE selection: 38
SInE iterations: 1

Duplicate literals: 199
Fw subsumption resolutions: 1687

Simple tautologies: 103
Forward subsumptions: 10255

Binary resolution: 30688
Factoring: 105

Backtracking splits: 933
Backtracking splits refuted: 193
Backtracking splits refuted at zero level: 43

Memory used: 6652KB
Time elapsed: 0.400 s
------------------------------
% remaining time: 1195 next slice time: 13
dis-1010_2_bs=off:ep=on:nwc=1.5:sswn=on:sswsr=on:ss=included:st=1.5:sgo=on:sp=occurrence_9 on unknown
Refutation found. Thanks to Tanya!
% SZS status Theorem for unknown
% SZS output start Proof for unknown
fof(f2528,plain,(
  $false),
  inference(subsumption_resolution,[],[f2511,f2479])).
fof(f2479,plain,(
  ~eqdim(sK0,sK1)),
  inference(resolution,[],[f2477,f268])).
fof(f268,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | less(X0,X1) | eqdim(X0,X1)) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : ((~leq(X0,X1) | (less(X0,X1) | eqdim(X0,X1))) & ((~less(X0,X1) & ~eqdim(X0,X1)) | leq(X0,X1)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1] : (leq(X0,X1) <=> (less(X0,X1) | eqdim(X0,X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X43] : ! [X44] : (leq(X43,X44) <=> (less(X43,X44) | eqdim(X43,X44)))),
  file(unknown,unknown)).
fof(f2477,plain,(
  ~leq(sK0,sK1)),
  inference(subsumption_resolution,[],[f2466,f202])).
fof(f202,plain,(
  ~zex(difference(sK0,sK1))),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ~zex(difference(sK0,sK1)) & ~p(difference(sK0,sK1),sK0)),
  inference(skolemisation,[status(esa)],[f106])).
fof(f106,plain,(
  ? [X0,X1] : (~zex(difference(X0,X1)) & ~p(difference(X0,X1),X0))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  inference(flattening,[],[f1])).
fof(f1,negated_conjecture,(
  ~! [X0] : ! [X1] : (~zex(difference(X0,X1)) => p(difference(X0,X1),X0))),
  file(unknown,unknown)).
fof(f2466,plain,(
  ~leq(sK0,sK1) | zex(difference(sK0,sK1))),
  inference(resolution,[],[f2456,f230])).
fof(f230,plain,(
  ( ! [X0,X1] : (eqdim(X0,difference(X0,X1)) | zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) | eqdim(X0,difference(X0,X1)))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : (~zex(difference(X0,X1)) => eqdim(X0,difference(X0,X1)))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1] : (~zex(difference(X0,X1)) => eqdim(X0,difference(X0,X1)))),
  inference(rectify,[],[f2])).
fof(f2,axiom,(
  ! [X2] : ! [X3] : (~zex(difference(X2,X3)) => eqdim(X2,difference(X2,X3)))),
  file(unknown,unknown)).
fof(f2456,plain,(
  ~eqdim(sK0,difference(sK0,sK1)) | ~leq(sK0,sK1)),
  inference(resolution,[],[f2448,f628])).
fof(f628,plain,(
  ( ! [X10,X9] : (eqdim(X9,X10) | ~eqdim(X10,X9)) )),
  inference(subsumption_resolution,[],[f623,f264])).
fof(f264,plain,(
  ( ! [X0,X1] : (~less(X1,X0) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & (less(X0,X1) | less(X1,X0) | eqdim(X0,X1)))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((~eqdim(X0,X1) | (~less(X0,X1) & ~less(X1,X0))) & ((less(X0,X1) | less(X1,X0)) | eqdim(X0,X1)))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (~less(X0,X1) & ~less(X1,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X41] : ! [X42] : (eqdim(X41,X42) <=> (~less(X41,X42) & ~less(X42,X41)))),
  file(unknown,unknown)).
fof(f623,plain,(
  ( ! [X10,X9] : (eqdim(X9,X10) | less(X9,X10) | ~eqdim(X10,X9)) )),
  inference(resolution,[],[f265,f263])).
fof(f263,plain,(
  ( ! [X0,X1] : (~less(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f265,plain,(
  ( ! [X0,X1] : (less(X1,X0) | eqdim(X0,X1) | less(X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f2448,plain,(
  ~eqdim(difference(sK0,sK1),sK0) | ~leq(sK0,sK1)),
  inference(subsumption_resolution,[],[f2431,f202])).
fof(f2431,plain,(
  ~leq(sK0,sK1) | zex(difference(sK0,sK1)) | ~eqdim(difference(sK0,sK1),sK0)),
  inference(resolution,[],[f889,f604])).
fof(f604,plain,(
  ~cont(difference(sK0,sK1),sK0) | ~eqdim(difference(sK0,sK1),sK0)),
  inference(resolution,[],[f262,f203])).
fof(f203,plain,(
  ~p(difference(sK0,sK1),sK0)),
  inference(cnf_transformation,[],[f147])).
fof(f262,plain,(
  ( ! [X0,X1] : (p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X52] : ! [X53] : (p(X52,X53) <=> (cont(X52,X53) & eqdim(X52,X53)))),
  file(unknown,unknown)).
fof(f889,plain,(
  ( ! [X0,X1] : (cont(difference(X0,X1),X0) | ~leq(X0,X1) | zex(difference(X0,X1))) )),
  inference(resolution,[],[f277,f216])).
fof(f216,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))),
  inference(nnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (~zex(X0) <=> cont(X0,X0))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X81] : (~zex(X81) <=> cont(X81,X81))),
  file(unknown,unknown)).
fof(f277,plain,(
  ( ! [X2,X0,X1] : (~cont(X2,difference(X0,X1)) | cont(X2,X0) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : (~leq(X0,X1) | ~cont(X2,difference(X0,X1)) | cont(X2,X0))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : ((~leq(X0,X1) | ~cont(X2,difference(X0,X1))) | cont(X2,X0))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2] : ((leq(X0,X1) & cont(X2,difference(X0,X1))) => cont(X2,X0))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ! [X1] : ! [X2] : ((leq(X0,X1) & cont(X2,difference(X0,X1))) => cont(X2,X0))),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X9] : ! [X10] : ! [X11] : ((leq(X9,X10) & cont(X11,difference(X9,X10))) => cont(X11,X9))),
  file(unknown,unknown)).
fof(f2511,plain,(
  eqdim(sK0,sK1)),
  inference(resolution,[],[f2480,f1424])).
fof(f1424,plain,(
  less(sK0,sK1) | eqdim(sK0,sK1)),
  inference(resolution,[],[f1419,f265])).
fof(f1419,plain,(
  ~less(sK1,sK0)),
  inference(subsumption_resolution,[],[f1414,f315])).
fof(f315,plain,(
  ~zex(sK0)),
  inference(resolution,[],[f273,f202])).
fof(f273,plain,(
  ( ! [X0,X1] : (zex(difference(X0,X1)) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : ((~zex(difference(X0,X1)) | (zex(X0) | cont(X0,X1))) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1] : (zex(difference(X0,X1)) <=> (zex(X0) | cont(X0,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X15] : ! [X16] : (zex(difference(X15,X16)) <=> (zex(X15) | cont(X15,X16)))),
  file(unknown,unknown)).
fof(f1414,plain,(
  ~less(sK1,sK0) | zex(sK0)),
  inference(resolution,[],[f1408,f216])).
fof(f1408,plain,(
  ~cont(sK0,sK0) | ~less(sK1,sK0)),
  inference(subsumption_resolution,[],[f1407,f627])).
fof(f627,plain,(
  ( ! [X0] : (eqdim(X0,X0)) )),
  inference(subsumption_resolution,[],[f618,f218])).
fof(f218,plain,(
  ( ! [X0] : (~less(X0,X0)) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : ~less(X0,X0)),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X30] : ~less(X30,X30)),
  file(unknown,unknown)).
fof(f618,plain,(
  ( ! [X0] : (eqdim(X0,X0) | less(X0,X0)) )),
  inference(resolution,[],[f265,f218])).
fof(f1407,plain,(
  ~cont(sK0,sK0) | ~eqdim(sK0,sK0) | ~less(sK1,sK0)),
  inference(duplicate_literal_removal,[],[f1406])).
fof(f1406,plain,(
  ~eqdim(sK0,sK0) | ~cont(sK0,sK0) | ~less(sK1,sK0) | ~less(sK1,sK0)),
  inference(superposition,[],[f1259,f221])).
fof(f221,plain,(
  ( ! [X0,X1] : (difference(X0,X1) = X0 | ~less(X1,X0)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : (~less(X1,X0) | difference(X0,X1) = X0)),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1] : (less(X1,X0) => difference(X0,X1) = X0)),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X4] : ! [X5] : (less(X5,X4) => difference(X4,X5) = X4)),
  file(unknown,unknown)).
fof(f1259,plain,(
  ~eqdim(difference(sK0,sK1),sK0) | ~cont(sK0,sK0) | ~less(sK1,sK0)),
  inference(superposition,[],[f604,f221])).
fof(f2480,plain,(
  ~less(sK0,sK1)),
  inference(resolution,[],[f2477,f267])).
fof(f267,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~less(X0,X1)) )),
  inference(cnf_transformation,[],[f192])).
% SZS output end Proof for unknown
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 345
Passive clauses: 1002
Generated clauses: 2321
Final active clauses: 345
Final passive clauses: 632
Input formulas: 38
Initial clauses: 87

Selected by SInE selection: 38
SInE iterations: 1

Duplicate literals: 62
Fw subsumption resolutions: 174

Simple tautologies: 71
Equational tautologies: 5
Forward subsumptions: 1018

Binary resolution: 1190
Factoring: 11
Forward superposition: 728
Backward superposition: 72

Unique components: 2

Memory used: 511KB
Time elapsed: 0.047 s
------------------------------
% Success in time 0.564 s
