Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 11999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_down_theoremsEP-E2.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1180
Passive clauses: 8045
Generated clauses: 16179
Final active clauses: 1101
Final passive clauses: 5319
Input formulas: 69
Initial clauses: 132

Selected by SInE selection: 69
SInE iterations: 1

Duplicate literals: 104
Fw subsumption resolutions: 895

Simple tautologies: 62
Forward subsumptions: 6998

Binary resolution: 14641
Factoring: 48

Backtracking splits: 254
Backtracking splits refuted: 83
Backtracking splits refuted at zero level: 37

Memory used: 2942KB
Time elapsed: 0.404 s
------------------------------
% remaining time: 11995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_down_theoremsEP-E2.all
Refutation found. Thanks to Tanya!
13733. $false (8:0) [subsumption resolution 13720,255]
255. eqdim(sK0,sK1) (0:3) [cnf transformation 232]
232. ~zex(sK0) & ~zex(sK1) & ~p(sK1,sK0) & eqdim(sK0,sK1) & po(sK0,sK1) & ~pp(sK1,sK0) & ! [X2] : (~p(X2,sK1) | po(X2,sK0))[skolemisation 138]
138. ? [X0,X1] : (~zex(X0) & ~zex(X1) & ~p(X1,X0) & eqdim(X0,X1) & po(X0,X1) & ~pp(X1,X0) & ! [X2] : (~p(X2,X1) | po(X2,X0)))[flattening 137]
137. ? [X0,X1] : ((~zex(X0) & ~zex(X1) & ~p(X1,X0) & eqdim(X0,X1) & po(X0,X1) & ~pp(X1,X0)) & ! [X2] : (~p(X2,X1) | po(X2,X0)))[ennf transformation 70]
70. ~! [X0,X1] : (~(~zex(X0) & ~zex(X1) & ~p(X1,X0) & eqdim(X0,X1) & po(X0,X1) & ~pp(X1,X0)) | ? [X2] : (p(X2,X1) & ~po(X2,X0)))[flattening 1]
1. ~! [X0] : ! [X1] : (~(~zex(X0) & ~zex(X1) & ~p(X1,X0) & eqdim(X0,X1) & po(X0,X1) & ~pp(X1,X0)) | ? [X2] : (p(X2,X1) & ~po(X2,X0)))[negated conjecture]
13720. ~eqdim(sK0,sK1) (8:3) [resolution 13615,332]
332. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 170]
170. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 101]
101. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[flattening 39]
39. ! [X0] : ! [X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[input]
13615. ~eqdim(sK1,sK0) (7:3) [resolution 13558,770]
770. ~cont(sK1,sK0) | ~eqdim(sK1,sK0) (1:6) [resolution 314,254]
254. ~p(sK1,sK0) (0:3) [cnf transformation 232]
314. p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 163]
163. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 162]
162. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 89]
89. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 54]
54. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
13558. cont(sK1,sK0) (6:3) [subsumption resolution 13548,253]
253. ~zex(sK1) (0:2) [cnf transformation 232]
13548. zex(sK1) | cont(sK1,sK0) (6:5) [resolution 13533,315]
315. ~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1) (0:9) [cnf transformation 164]
164. ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & ((~zex(X0) & ~cont(X0,X1)) | zex(difference(X0,X1))))[ennf transformation 90]
90. ! [X0,X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & (~(zex(X0) | cont(X0,X1)) | zex(difference(X0,X1))))[flattening 11]
11. ! [X0] : ! [X1] : ((~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) & (~(zex(X0) | cont(X0,X1)) | zex(difference(X0,X1))))[input]
13533. zex(difference(sK1,sK0)) (5:4) [resolution 13525,324]
324. p(difference(X0,X1),X0) | zex(difference(X0,X1)) (0:9) [cnf transformation 95]
95. ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))[flattening 12]
12. ! [X0] : ! [X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))[input]
13525. ~p(difference(X0,sK0),sK1) (4:5) [resolution 13083,258]
258. po(X2,sK0) | ~p(X2,sK1) (0:6) [cnf transformation 232]
13083. ~po(difference(X0,X1),X1) (3:5) [duplicate literal removal 13069]
13069. ~po(difference(X0,X1),X1) | ~po(difference(X0,X1),X1) (3:10) [resolution 1675,304]
304. p(sK10(X1,X0),X1) | ~po(X0,X1) (0:8) [cnf transformation 243]
243. ! [X0,X1] : ((~po(X0,X1) | (p(sK10(X1,X0),X0) & p(sK10(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[skolemisation 157]
157. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[ennf transformation 86]
86. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[flattening 85]
85. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[rectify 31]
31. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X2] : (p(X2,X0) & p(X2,X1)) | po(X0,X1)))[input]
1675. ~p(sK10(X12,difference(X13,X14)),X14) | ~po(difference(X13,X14),X12) (2:12) [resolution 1172,303]
303. p(sK10(X1,X0),X0) | ~po(X0,X1) (0:8) [cnf transformation 243]
1172. ~p(X0,difference(X1,X2)) | ~p(X0,X2) (1:8) [resolution 1171,305]
305. po(X0,X1) | ~p(X3,X1) | ~p(X3,X0) (0:9) [cnf transformation 243]
1171. ~po(X1,difference(X0,X1)) (0:5) [subsumption resolution 1170,614]
614. ~po(X0,X1) | ~zex(X1) (2:5) [resolution 304,411]
411. ~p(X3,X4) | ~zex(X4) (1:5) [resolution 312,327]
327. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 167]
167. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 98]
98. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[flattening 50]
50. ! [X0] : ! [X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[input]
312. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 163]
1170. ~po(X1,difference(X0,X1)) | zex(difference(X0,X1)) (0:9) [subsumption resolution 351,600]
600. ~po(X0,X1) | ~zex(X0) (2:5) [resolution 303,411]
351. ~po(X1,difference(X0,X1)) | zex(difference(X0,X1)) | zex(X1) (0:11) [cnf transformation 196]
196. ! [X0,X1] : (zex(X1) | zex(difference(X0,X1)) | ~po(X1,difference(X0,X1)))[flattening 195]
195. ! [X0,X1] : ((zex(X1) | zex(difference(X0,X1))) | ~po(X1,difference(X0,X1)))[ennf transformation 117]
117. ! [X0,X1] : (~(~zex(X1) & ~zex(difference(X0,X1))) | ~po(X1,difference(X0,X1)))[flattening 15]
15. ! [X0] : ! [X1] : (~(~zex(X1) & ~zex(difference(X0,X1))) | ~po(X1,difference(X0,X1)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 1111
Passive clauses: 5556
Generated clauses: 13472
Final active clauses: 1111
Final passive clauses: 4260
Input formulas: 69
Initial clauses: 131

Selected by SInE selection: 69
SInE iterations: 1

Duplicate literals: 257
Fw subsumption resolutions: 787

Simple tautologies: 482
Equational tautologies: 6
Forward subsumptions: 6579

Binary resolution: 9373
Factoring: 50
Forward superposition: 2795
Backward superposition: 102

Unique components: 7

Memory used: 2174KB
Time elapsed: 0.436 s
------------------------------
% Success in time 0.894 s
