Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal13.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 555
Passive clauses: 2119
Generated clauses: 4224
Final active clauses: 548
Final passive clauses: 1462
Input formulas: 43
Initial clauses: 90

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 51
Fw subsumption resolutions: 120

Simple tautologies: 42
Forward subsumptions: 1915

Binary resolution: 3859
Factoring: 18

Backtracking splits: 61
Backtracking splits refuted: 9
Backtracking splits refuted at zero level: 1

Memory used: 895KB
Time elapsed: 0.407 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal13.all
Refutation found. Thanks to Tanya!
1020. $false (8:0) [subsumption resolution 1015,212]
212. ~zex(sK0) (0:2) [cnf transformation 167]
167. ~zex(sK0) & ~zex(intersection(sK0,sK0)) & ~p(intersection(sK0,sK0),sK0)[skolemisation 119]
119. ? [X0] : (~zex(X0) & ~zex(intersection(X0,X0)) & ~p(intersection(X0,X0),X0))[flattening 118]
118. ? [X0] : ((~zex(X0) & ~zex(intersection(X0,X0))) & ~p(intersection(X0,X0),X0))[ennf transformation 1]
1. ~! [X0] : ((~zex(X0) & ~zex(intersection(X0,X0))) => p(intersection(X0,X0),X0))[negated conjecture]
1015. zex(sK0) (8:2) [resolution 1007,224]
224. cont(X0,X0) | zex(X0) (0:5) [cnf transformation 176]
176. ! [X0] : ((zex(X0) | cont(X0,X0)) & (~cont(X0,X0) | ~zex(X0)))[nnf transformation 122]
122. ! [X0] : (~zex(X0) <=> cont(X0,X0))[ennf transformation 48]
48. ! [X0] : (~zex(X0) <=> cont(X0,X0))[rectify 23]
23. ! [X45] : (~zex(X45) <=> cont(X45,X45))[input]
1007. ~cont(sK0,sK0) (7:3) [subsumption resolution 1000,213]
213. ~zex(intersection(sK0,sK0)) (0:4) [cnf transformation 167]
1000. ~cont(sK0,sK0) | zex(intersection(sK0,sK0)) (7:7) [resolution 966,240]
240. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 143]
143. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 75]
75. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 74]
74. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 30]
30. ! [X59] : ! [X60] : (~zex(intersection(X59,X60)) => cont(intersection(X59,X60),X59))[input]
966. ~cont(intersection(sK0,sK0),sK0) | ~cont(sK0,sK0) (6:8) [resolution 643,906]
906. ~lt(intersection(sK0,sK0),sK0) | ~cont(sK0,sK0) (5:8) [resolution 886,275]
275. leq(X0,X1) | ~lt(X0,X1) (0:6) [cnf transformation 205]
205. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & (~leq(X0,X1) | eqdim(X0,X1) | lt(X0,X1)))[flattening 204]
204. ! [X0,X1] : ((~lt(X0,X1) | (leq(X0,X1) & ~eqdim(X0,X1))) & ((~leq(X0,X1) | eqdim(X0,X1)) | lt(X0,X1)))[nnf transformation 148]
148. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[ennf transformation 97]
97. ! [X0,X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[flattening 96]
96. ! [X0] : ! [X1] : (lt(X0,X1) <=> (leq(X0,X1) & ~eqdim(X0,X1)))[rectify 18]
18. ! [X34] : ! [X35] : (lt(X34,X35) <=> (leq(X34,X35) & ~eqdim(X34,X35)))[input]
886. ~leq(intersection(sK0,sK0),sK0) | ~cont(sK0,sK0) (4:8) [duplicate literal removal 882]
882. ~cont(sK0,sK0) | ~cont(sK0,sK0) | ~leq(intersection(sK0,sK0),sK0) (4:11) [resolution 283,642]
642. ~leq(sK0,intersection(sK0,sK0)) | ~leq(intersection(sK0,sK0),sK0) (3:10) [resolution 633,274]
274. eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1) (0:9) [cnf transformation 203]
203. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & (~leq(X0,X1) | ~leq(X1,X0) | eqdim(X0,X1)))[flattening 202]
202. ! [X0,X1] : ((~eqdim(X0,X1) | (leq(X0,X1) & leq(X1,X0))) & ((~leq(X0,X1) | ~leq(X1,X0)) | eqdim(X0,X1)))[nnf transformation 95]
95. ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[flattening 94]
94. ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X0,X1) & leq(X1,X0)))[rectify 28]
28. ! [X55] : ! [X56] : (eqdim(X55,X56) <=> (leq(X55,X56) & leq(X56,X55)))[input]
633. ~eqdim(intersection(sK0,sK0),sK0) (2:5) [subsumption resolution 627,213]
627. ~eqdim(intersection(sK0,sK0),sK0) | zex(intersection(sK0,sK0)) (2:9) [resolution 570,240]
570. ~cont(intersection(sK0,sK0),sK0) | ~eqdim(intersection(sK0,sK0),sK0) (1:10) [resolution 271,214]
214. ~p(intersection(sK0,sK0),sK0) (0:5) [cnf transformation 167]
271. p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 201]
201. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 200]
200. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[nnf transformation 93]
93. ! [X0,X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[flattening 92]
92. ! [X0] : ! [X1] : (p(X0,X1) <=> (cont(X0,X1) & eqdim(X0,X1)))[rectify 19]
19. ! [X36] : ! [X37] : (p(X36,X37) <=> (cont(X36,X37) & eqdim(X36,X37)))[input]
283. leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0) (0:11) [cnf transformation 157]
157. ! [X0,X1,X2] : (~cont(X2,X0) | ~cont(X2,X1) | leq(X2,intersection(X0,X1)))[flattening 156]
156. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1)) | leq(X2,intersection(X0,X1)))[ennf transformation 107]
107. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[flattening 106]
106. ! [X0] : ! [X1] : ! [X2] : ((cont(X2,X0) & cont(X2,X1)) => leq(X2,intersection(X0,X1)))[rectify 31]
31. ! [X61] : ! [X62] : ! [X63] : ((cont(X63,X61) & cont(X63,X62)) => leq(X63,intersection(X61,X62)))[input]
643. lt(intersection(sK0,sK0),sK0) | ~cont(intersection(sK0,sK0),sK0) (3:10) [resolution 633,230]
230. eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 127]
127. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 126]
126. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 55]
55. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 54]
54. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 34]
34. ! [X70] : ! [X71] : (cont(X70,X71) => (lt(X70,X71) | eqdim(X70,X71)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 164
Passive clauses: 362
Generated clauses: 803
Final active clauses: 164
Final passive clauses: 185
Input formulas: 43
Initial clauses: 85

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 27
Fw subsumption resolutions: 65

Simple tautologies: 26
Equational tautologies: 1
Forward subsumptions: 330

Binary resolution: 380
Factoring: 2
Forward superposition: 212
Backward superposition: 32

Unique components: 3

Memory used: 255KB
Time elapsed: 0.062 s
------------------------------
% Success in time 0.633 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 00:58:42 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal13.all.tptp"
Input read from 
============================ end of footer ===========================
