% Starting sequential portfolio solving with schedule "casc"
% spawned child 12839 with time: 4 (total remaining time 2999)
% WARNING: value z3 for option sas not known
% dis+1010_4_afp=10000:afq=1.2:anc=none:irw=on:lma=on:nm=64:nwc=10:sas=z3:sac=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal13.all
% Refutation found. Thanks to Tanya!
% SZS status Theorem for codi_down_sum_theorems_goal13.all
% SZS output start Proof for codi_down_sum_theorems_goal13.all
fof(f1,conjecture,(
  ! [X0] : ! [X1] : ! [X2] : (p(difference(X2,X0),X2) | ~(~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal13.all.tptp',conjecture130)).
fof(f2,negated_conjecture,(
  ~! [X0] : ! [X1] : ! [X2] : (p(difference(X2,X0),X2) | ~(~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)))),
  inference(negated_conjecture,[],[f1])).
fof(f12,axiom,(
  ! [X0] : ! [X1] : ((zex(difference(X0,X1)) | ~(cont(X0,X1) | zex(X0))) & (cont(X0,X1) | zex(X0) | ~zex(difference(X0,X1))))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal13.all.tptp',axiom15200)).
fof(f21,axiom,(
  ! [X0] : ! [X1] : (p(difference(X0,X1),X0) | zex(difference(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal13.all.tptp',axiom64980)).
fof(f88,axiom,(
  ! [X0] : ! [X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal13.all.tptp',axiom89540)).
fof(f118,plain,(
  ~! [X0,X1,X2] : (p(difference(X2,X0),X2) | ~(~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)))),
  inference(flattening,[],[f2])).
fof(f154,plain,(
  ! [X0,X1] : ((zex(difference(X0,X1)) | ~(cont(X0,X1) | zex(X0))) & (cont(X0,X1) | zex(X0) | ~zex(difference(X0,X1))))),
  inference(flattening,[],[f12])).
fof(f158,plain,(
  ! [X0,X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  inference(flattening,[],[f88])).
fof(f180,plain,(
  ! [X0,X1] : (p(difference(X0,X1),X0) | zex(difference(X0,X1)))),
  inference(flattening,[],[f21])).
fof(f274,plain,(
  ? [X0,X1,X2] : (~p(difference(X2,X0),X2) & (~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f118])).
fof(f275,plain,(
  ? [X0,X1,X2] : (~p(difference(X2,X0),X2) & ~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1))),
  inference(flattening,[],[f274])).
fof(f305,plain,(
  ! [X0,X1] : ((zex(difference(X0,X1)) | (~cont(X0,X1) & ~zex(X0))) & (cont(X0,X1) | zex(X0) | ~zex(difference(X0,X1))))),
  inference(ennf_transformation,[],[f154])).
fof(f389,plain,(
  ? [X0,X1,X2] : (~p(difference(X2,X0),X2) & ~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)) => (~p(difference(sK2,sK0),sK2) & ~zex(sK1) & ~zex(sK0) & ~cont(sK2,sK1) & ~cont(sK2,sK0) & cont(sK2,sum(sK0,sK1)) & eqdim(sK0,sK1))),
  introduced(choice_axiom,[])).
fof(f390,plain,(
  ~p(difference(sK2,sK0),sK2) & ~zex(sK1) & ~zex(sK0) & ~cont(sK2,sK1) & ~cont(sK2,sK0) & cont(sK2,sum(sK0,sK1)) & eqdim(sK0,sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f275,f389])).
fof(f442,plain,(
  cont(sK2,sum(sK0,sK1))),
  inference(cnf_transformation,[],[f390])).
fof(f443,plain,(
  ~cont(sK2,sK0)),
  inference(cnf_transformation,[],[f390])).
fof(f447,plain,(
  ~p(difference(sK2,sK0),sK2)),
  inference(cnf_transformation,[],[f390])).
fof(f525,plain,(
  ( ! [X0,X1] : (~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f305])).
fof(f534,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f158])).
fof(f556,plain,(
  ( ! [X0,X1] : (p(difference(X0,X1),X0) | zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f180])).
fof(f636,plain,(
  ~zex(sK2)),
  inference(resolution,[],[f534,f442])).
fof(f1128,plain,(
  zex(difference(sK2,sK0))),
  inference(resolution,[],[f556,f447])).
fof(f1138,plain,(
  zex(sK2) | cont(sK2,sK0)),
  inference(resolution,[],[f1128,f525])).
fof(f1141,plain,(
  cont(sK2,sK0)),
  inference(subsumption_resolution,[],[f1138,f636])).
fof(f1142,plain,(
  $false),
  inference(subsumption_resolution,[],[f1141,f443])).
% SZS output end Proof for codi_down_sum_theorems_goal13.all
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.064 s
% ------------------------------
% ------------------------------
% Success in time 0.168 s
