Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_int_theorems_goal24.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 875
Passive clauses: 4714
Generated clauses: 11585
Final active clauses: 859
Final passive clauses: 3543
Input formulas: 50
Initial clauses: 86

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 77
Fw subsumption resolutions: 495

Simple tautologies: 61
Forward subsumptions: 6305

Binary resolution: 10715
Factoring: 28

Backtracking splits: 157
Backtracking splits refuted: 12

Memory used: 2046KB
Time elapsed: 0.405 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_int_theorems_goal24.all
Refutation found. Thanks to Tanya!
746. $false (8:0) [subsumption resolution 745,442]
442. c(sK1,sK0) (3:3) [resolution 440,187]
187. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 102]
102. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[ennf transformation 54]
54. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[flattening 11]
11. ! [X0] : ! [X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[input]
440. ~zex(intersection(sK1,sK0)) (2:4) [resolution 434,304]
304. c(intersection(sK1,sK0),sK0) (1:5) [resolution 187,173]
173. ~zex(intersection(intersection(sK1,sK0),sK0)) (0:6) [cnf transformation 162]
162. zex(intersection(sK0,intersection(sK0,sK1))) & ~zex(intersection(intersection(sK1,sK0),sK0))[skolemisation 95]
95. ? [X0,X1] : (zex(intersection(X0,intersection(X0,X1))) & ~zex(intersection(intersection(X1,X0),X0)))[ennf transformation 51]
51. ~! [X0,X1] : (~zex(intersection(X0,intersection(X0,X1))) | zex(intersection(intersection(X1,X0),X0)))[flattening 1]
1. ~! [X0] : ! [X1] : (~zex(intersection(X0,intersection(X0,X1))) | zex(intersection(intersection(X1,X0),X0)))[negated conjecture]
434. ~c(X2,X3) | ~zex(X2) (1:5) [resolution 198,213]
213. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 116]
116. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 67]
67. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[flattening 31]
31. ! [X0] : ! [X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[input]
198. cont(sK5(X1,X0),X0) | ~c(X0,X1) (0:8) [cnf transformation 166]
166. ! [X0,X1] : ((~c(X0,X1) | (cont(sK5(X1,X0),X0) & cont(sK5(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 108]
108. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[ennf transformation 60]
60. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[flattening 59]
59. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X3] : (cont(X3,X0) & cont(X3,X1)) | c(X0,X1)))[rectify 34]
34. ! [X0] : ! [X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (~? [X2] : (cont(X2,X0) & cont(X2,X1)) | c(X0,X1)))[input]
745. ~c(sK1,sK0) (8:3) [duplicate literal removal 737]
737. ~c(sK1,sK0) | ~c(sK1,sK0) (8:6) [resolution 560,199]
199. cont(sK5(X1,X0),X1) | ~c(X0,X1) (0:8) [cnf transformation 166]
560. ~cont(sK5(X1,sK1),sK0) | ~c(sK1,X1) (7:8) [resolution 539,198]
539. ~cont(X1,sK1) | ~cont(X1,sK0) (6:6) [resolution 529,200]
200. c(X0,X1) | ~cont(X3,X1) | ~cont(X3,X0) (0:9) [cnf transformation 166]
529. ~c(sK0,sK1) (5:3) [resolution 520,188]
188. ~zex(intersection(X0,X1)) | ~c(X0,X1) (0:7) [cnf transformation 102]
520. zex(intersection(sK0,sK1)) (4:4) [resolution 506,185]
185. p(X0,X0) | zex(X0) (0:5) [cnf transformation 2]
2. ! [X0] : (zex(X0) | p(X0,X0))[input]
506. ~p(X0,intersection(sK0,sK1)) (3:5) [subsumption resolution 496,235]
235. ~p(X2,intersection(X0,X1)) | cont(X2,X0) (0:8) [cnf transformation 137]
137. ! [X0,X1,X2] : ((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1)) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 136]
136. ! [X0,X1,X2] : (((~cont(X2,X0) | ~cont(X2,X1) | ~eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[ennf transformation 81]
81. ! [X0,X1,X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[flattening 14]
14. ! [X0] : ! [X1] : ! [X2] : ((~(cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1))) | p(X2,intersection(X0,X1))) & (~p(X2,intersection(X0,X1)) | (cont(X2,X0) & cont(X2,X1) & eqdim(X2,intersection(X0,X1)))))[input]
496. ~cont(X0,sK0) | ~p(X0,intersection(sK0,sK1)) (3:8) [resolution 492,201]
201. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 110]
110. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 109]
109. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 61]
61. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 35]
35. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
492. ~cont(X0,intersection(sK0,sK1)) | ~cont(X0,sK0) (2:8) [resolution 200,305]
305. ~c(sK0,intersection(sK0,sK1)) (1:5) [resolution 188,172]
172. zex(intersection(sK0,intersection(sK0,sK1))) (0:6) [cnf transformation 162]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 121
Passive clauses: 278
Generated clauses: 562
Final active clauses: 121
Final passive clauses: 137
Input formulas: 50
Initial clauses: 85

Selected by SInE selection: 50
SInE iterations: 1

Duplicate literals: 21
Fw subsumption resolutions: 31

Simple tautologies: 32
Equational tautologies: 1
Forward subsumptions: 211

Binary resolution: 213
Factoring: 2
Forward superposition: 158
Backward superposition: 52

Unique components: 2

Memory used: 255KB
Time elapsed: 0.025 s
------------------------------
% Success in time 0.501 s
