Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal04.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 657
Passive clauses: 2961
Generated clauses: 6883
Final active clauses: 655
Final passive clauses: 2185
Input formulas: 43
Initial clauses: 89

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 48
Fw subsumption resolutions: 53

Simple tautologies: 38
Forward subsumptions: 3819

Binary resolution: 6642
Factoring: 44

Backtracking splits: 80
Backtracking splits refuted: 6

Memory used: 1151KB
Time elapsed: 0.411 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal04.all
Refutation found. Thanks to Tanya!
663. $false (4:0) [subsumption resolution 662,326]
326. c(sK1,sK0) (1:3) [resolution 277,213]
213. ~zex(intersection(sK1,sK0)) (0:4) [cnf transformation 167]
167. zex(intersection(sK0,sK1)) & ~zex(intersection(sK1,sK0))[skolemisation 119]
119. ? [X0,X1] : (zex(intersection(X0,X1)) & ~zex(intersection(X1,X0)))[ennf transformation 44]
44. ~! [X0,X1] : (zex(intersection(X0,X1)) => zex(intersection(X1,X0)))[flattening 1]
1. ~! [X0] : ! [X1] : (zex(intersection(X0,X1)) => zex(intersection(X1,X0)))[negated conjecture]
277. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 206]
206. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[nnf transformation 149]
149. ! [X0,X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[ennf transformation 99]
99. ! [X0,X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[flattening 98]
98. ! [X0] : ! [X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[rectify 24]
24. ! [X48] : ! [X49] : (~c(X48,X49) <=> zex(intersection(X48,X49)))[input]
662. ~c(sK1,sK0) (4:3) [duplicate literal removal 654]
654. ~c(sK1,sK0) | ~c(sK1,sK0) (4:6) [resolution 557,266]
266. cont(sK11(X1,X0),X1) | ~c(X0,X1) (0:8) [cnf transformation 199]
199. ! [X0,X1] : ((~c(X0,X1) | (cont(sK11(X1,X0),X0) & cont(sK11(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 198]
198. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[rectify 197]
197. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | c(X0,X1)))[nnf transformation 91]
91. ! [X0,X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))[flattening 90]
90. ! [X0] : ! [X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))[rectify 2]
2. ! [X2] : ! [X3] : (c(X2,X3) <=> ? [X4] : (cont(X4,X2) & cont(X4,X3)))[input]
557. ~cont(sK11(X5,sK1),sK0) | ~c(sK1,X5) (3:8) [resolution 546,265]
265. cont(sK11(X1,X0),X0) | ~c(X0,X1) (0:8) [cnf transformation 199]
546. ~cont(X0,sK1) | ~cont(X0,sK0) (2:6) [resolution 267,327]
327. ~c(sK0,sK1) (1:3) [resolution 278,212]
212. zex(intersection(sK0,sK1)) (0:4) [cnf transformation 167]
278. ~zex(intersection(X0,X1)) | ~c(X0,X1) (0:7) [cnf transformation 206]
267. c(X0,X1) | ~cont(X3,X1) | ~cont(X3,X0) (0:9) [cnf transformation 199]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 108
Passive clauses: 230
Generated clauses: 446
Final active clauses: 108
Final passive clauses: 106
Input formulas: 43
Initial clauses: 84

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 32
Fw subsumption resolutions: 27

Simple tautologies: 25
Equational tautologies: 1
Forward subsumptions: 142

Binary resolution: 151
Factoring: 3
Forward superposition: 112
Backward superposition: 40

Unique components: 2

Memory used: 255KB
Time elapsed: 0.021 s
------------------------------
% Success in time 0.509 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 23:32:54 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal04.all.tptp"
Input read from 
============================ end of footer ===========================
