Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 11999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on theorems\conversions\codi_down_theoremsEP-E1.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 1239
Passive clauses: 8485
Generated clauses: 17724
Final active clauses: 1134
Final passive clauses: 5122
Input formulas: 65
Initial clauses: 119

Selected by SInE selection: 65
SInE iterations: 1

Duplicate literals: 177
Fw subsumption resolutions: 917

Simple tautologies: 66
Forward subsumptions: 7911

Binary resolution: 15959
Factoring: 48

Backtracking splits: 332
Backtracking splits refuted: 122
Backtracking splits refuted at zero level: 71

Memory used: 2942KB
Time elapsed: 0.405 s
------------------------------
% remaining time: 11995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on theorems\conversions\codi_down_theoremsEP-E1.all
Refutation found. Thanks to Tanya!
10295. $false (6:0) [subsumption resolution 10286,232]
232. pp(sK1,sK0) (0:3) [cnf transformation 216]
216. pp(sK1,sK0) & ! [X2] : (~p(X2,sK0) | po(X2,sK1))[skolemisation 129]
129. ? [X0,X1] : (pp(X1,X0) & ! [X2] : (~p(X2,X0) | po(X2,X1)))[ennf transformation 66]
66. ~! [X0,X1] : (~pp(X1,X0) | ? [X2] : (p(X2,X0) & ~po(X2,X1)))[flattening 1]
1. ~! [X0] : ! [X1] : (~pp(X1,X0) | ? [X2] : (p(X2,X0) & ~po(X2,X1)))[negated conjecture]
10286. ~pp(sK1,sK0) (6:3) [resolution 10260,582]
582. ~zex(difference(X3,X2)) | ~pp(X2,X3) (3:7) [resolution 304,370]
370. ~pp(X0,X1) | ~zex(X0) (2:5) [resolution 367,278]
278. p(X0,X1) | ~pp(X0,X1) (0:6) [cnf transformation 148]
148. ! [X0,X1] : ((~pp(X0,X1) | (p(X0,X1) & X0 != X1)) & (~p(X0,X1) | X0 = X1 | pp(X0,X1)))[flattening 147]
147. ! [X0,X1] : ((~pp(X0,X1) | (p(X0,X1) & X0 != X1)) & ((~p(X0,X1) | X0 = X1) | pp(X0,X1)))[ennf transformation 82]
82. ! [X0,X1] : ((~pp(X0,X1) | (p(X0,X1) & ~X0 = X1)) & (~(p(X0,X1) & ~X0 = X1) | pp(X0,X1)))[flattening 26]
26. ! [X0] : ! [X1] : ((~pp(X0,X1) | (p(X0,X1) & ~X0 = X1)) & (~(p(X0,X1) & ~X0 = X1) | pp(X0,X1)))[input]
367. ~p(X1,X2) | ~zex(X1) (1:5) [resolution 284,303]
303. ~cont(X0,X1) | ~zex(X0) (0:5) [cnf transformation 158]
158. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 94]
94. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[flattening 46]
46. ! [X0] : ! [X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[input]
284. cont(X0,X1) | ~p(X0,X1) (0:6) [cnf transformation 152]
152. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~cont(X0,X1) | ~eqdim(X0,X1) | p(X0,X1)))[flattening 151]
151. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & ((~cont(X0,X1) | ~eqdim(X0,X1)) | p(X0,X1)))[ennf transformation 84]
84. ! [X0,X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[flattening 50]
50. ! [X0] : ! [X1] : ((~p(X0,X1) | (cont(X0,X1) & eqdim(X0,X1))) & (~(cont(X0,X1) & eqdim(X0,X1)) | p(X0,X1)))[input]
304. pp(difference(X0,X1),X0) | ~pp(X1,X0) (0:8) [cnf transformation 159]
159. ! [X0,X1] : (~pp(X1,X0) | pp(difference(X0,X1),X0))[ennf transformation 95]
95. ! [X0,X1] : (~pp(X1,X0) | pp(difference(X0,X1),X0))[flattening 10]
10. ! [X0] : ! [X1] : (~pp(X1,X0) | pp(difference(X0,X1),X0))[input]
10260. zex(difference(sK0,sK1)) (5:4) [resolution 10252,299]
299. p(difference(X0,X1),X0) | zex(difference(X0,X1)) (0:9) [cnf transformation 91]
91. ! [X0,X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))[flattening 8]
8. ! [X0] : ! [X1] : (zex(difference(X0,X1)) | p(difference(X0,X1),X0))[input]
10252. ~p(difference(X0,sK1),sK0) (4:5) [resolution 10250,233]
233. po(X2,sK1) | ~p(X2,sK0) (0:6) [cnf transformation 216]
10250. ~po(difference(X8,X9),X9) (3:5) [duplicate literal removal 10241]
10241. ~po(difference(X8,X9),X9) | ~po(difference(X8,X9),X9) (3:10) [resolution 1385,276]
276. p(sK10(X1,X0),X1) | ~po(X0,X1) (0:8) [cnf transformation 227]
227. ! [X0,X1] : ((~po(X0,X1) | (p(sK10(X1,X0),X0) & p(sK10(X1,X0),X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[skolemisation 146]
146. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (! [X3] : (~p(X3,X0) | ~p(X3,X1)) | po(X0,X1)))[ennf transformation 81]
81. ! [X0,X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[flattening 80]
80. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X3] : (p(X3,X0) & p(X3,X1)) | po(X0,X1)))[rectify 27]
27. ! [X0] : ! [X1] : ((~po(X0,X1) | ? [X2] : (p(X2,X0) & p(X2,X1))) & (~? [X2] : (p(X2,X0) & p(X2,X1)) | po(X0,X1)))[input]
1385. ~p(sK10(X12,difference(X13,X14)),X14) | ~po(difference(X13,X14),X12) (2:12) [resolution 1098,275]
275. p(sK10(X1,X0),X0) | ~po(X0,X1) (0:8) [cnf transformation 227]
1098. ~p(X0,difference(X1,X2)) | ~p(X0,X2) (1:8) [resolution 1097,277]
277. po(X0,X1) | ~p(X3,X1) | ~p(X3,X0) (0:9) [cnf transformation 227]
1097. ~po(X1,difference(X0,X1)) (0:5) [subsumption resolution 1096,572]
572. ~po(X0,X1) | ~zex(X1) (2:5) [resolution 276,368]
368. ~p(X3,X4) | ~zex(X4) (1:5) [resolution 284,302]
302. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 158]
1096. ~po(X1,difference(X0,X1)) | zex(difference(X0,X1)) (0:9) [subsumption resolution 318,559]
559. ~po(X0,X1) | ~zex(X0) (2:5) [resolution 275,368]
318. ~po(X1,difference(X0,X1)) | zex(difference(X0,X1)) | zex(X1) (0:11) [cnf transformation 180]
180. ! [X0,X1] : (zex(X1) | zex(difference(X0,X1)) | ~po(X1,difference(X0,X1)))[flattening 179]
179. ! [X0,X1] : ((zex(X1) | zex(difference(X0,X1))) | ~po(X1,difference(X0,X1)))[ennf transformation 109]
109. ! [X0,X1] : (~(~zex(X1) & ~zex(difference(X0,X1))) | ~po(X1,difference(X0,X1)))[flattening 11]
11. ! [X0] : ! [X1] : (~(~zex(X1) & ~zex(difference(X0,X1))) | ~po(X1,difference(X0,X1)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 956
Passive clauses: 4368
Generated clauses: 10054
Final active clauses: 956
Final passive clauses: 3246
Input formulas: 65
Initial clauses: 118

Selected by SInE selection: 65
SInE iterations: 1

Duplicate literals: 167
Fw subsumption resolutions: 627

Simple tautologies: 316
Equational tautologies: 8
Forward subsumptions: 4734

Binary resolution: 7106
Factoring: 45
Forward superposition: 1943
Backward superposition: 54

Unique components: 2

Memory used: 1663KB
Time elapsed: 0.408 s
------------------------------
% Success in time 0.873 s
