% Starting sequential portfolio solving with schedule "casc"
% spawned child 12821 with time: 4 (total remaining time 2999)
% WARNING: value z3 for option sas not known
% dis+1010_4_afp=10000:afq=1.2:anc=none:irw=on:lma=on:nm=64:nwc=10:sas=z3:sac=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6908
% Time elapsed: 0.500 s
% spawned child 12822 with time: 10 (total remaining time 2994)
% ott+11_8_amm=off:anc=none:bsr=on:cond=on:irw=on:nm=2:nwc=1.1:ss=axioms:st=2.0:sac=on_8 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8699
% Time elapsed: 1.100 s
% ------------------------------
% ------------------------------
% spawned child 12823 with time: 4 (total remaining time 2985)
% dis+1011_10_aac=none:add=large:afp=10000:afq=1.1:amm=off:anc=none:fsr=off:gs=on:gsem=off:irw=on:nm=6:newcnf=on:nwc=2.5:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 11257
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 12824 with time: 6 (total remaining time 2979)
% dis+11_28_av=off:fsr=off:irw=on:lcm=predicate:nm=2:newcnf=on:nwc=5:sp=occurrence:urr=on:updr=off_4 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 4221
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% spawned child 12825 with time: 6 (total remaining time 2972)
% lrs+1011_5_add=large:afp=1000:afq=1.2:anc=none:fsr=off:irw=on:lma=on:nm=64:newcnf=on:nwc=1:stl=30:sd=3:ss=axioms:st=2.0:sos=on:sac=on:sp=reverse_arity:urr=on:updr=off_4 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7419
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% spawned child 12826 with time: 4 (total remaining time 2965)
% lrs+1011_3:1_add=off:afr=on:afp=10000:afq=1.1:amm=off:bce=on:cond=on:ep=R:fsr=off:nm=16:nwc=1:stl=30:sos=all:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8827
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 12827 with time: 23 (total remaining time 2960)
% lrs+1011_3:2_aac=none:afp=10000:afq=1.2:amm=off:anc=all:cond=on:fde=none:gs=on:gsem=on:irw=on:lma=on:nm=32:newcnf=on:nwc=3:nicw=on:stl=30:sac=on:sp=reverse_arity:urr=ec_only:updr=off:uhcvi=on_21 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 18421
% Time elapsed: 2.400 s
% ------------------------------
% ------------------------------
% spawned child 12828 with time: 4 (total remaining time 2939)
% lrs+10_6_aac=none:acc=model:add=large:afr=on:afp=10000:afq=1.2:amm=off:anc=none:bd=off:ccuc=small_ones:irw=on:lcm=reverse:nm=0:nwc=1:nicw=on:stl=30:sos=on:sp=reverse_arity:updr=off_2 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 5884
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 12829 with time: 7 (total remaining time 2934)
% dis+11_5_av=off:bd=off:bs=unit_only:bsr=on:cond=on:lcm=reverse:nm=0:nwc=1.2_5 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 2814
% Time elapsed: 0.800 s
% ------------------------------
% ------------------------------
% spawned child 12830 with time: 4 (total remaining time 2927)
% dis+11_32_av=off:ep=RST:fsr=off:lwlo=on:nm=6:nwc=1.1:sd=5:ss=axioms:st=5.0:sp=reverse_arity:uhcvi=on_2 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 1791
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 12831 with time: 7 (total remaining time 2922)
% dis+1003_64_add=off:afr=on:afp=100000:afq=1.1:anc=none:cond=fast:fde=none:irw=on:nm=6:newcnf=on:nwc=1.3:uhcvi=on_5 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7419
% Time elapsed: 0.800 s
% ------------------------------
% ------------------------------
% spawned child 12832 with time: 5 (total remaining time 2914)
% dis+11_4_afr=on:afp=1000:afq=1.4:amm=off:anc=none:lcm=reverse:lma=on:lwlo=on:nm=6:newcnf=on:nwc=1:sd=2:ss=axioms:st=2.0:sp=reverse_arity_3 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7036
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% spawned child 12833 with time: 87 (total remaining time 2909)
% ott+1002_128_av=off:bd=off:bs=on:bsr=on:cond=on:fsr=off:nm=6:newcnf=on:nwc=1:sp=reverse_arity:updr=off_82 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit

% Memory used [KB]: 6396
% Time elapsed: 8.827 s
% ------------------------------
% ------------------------------
% spawned child 12834 with time: 8 (total remaining time 2837)
% lrs+1002_8_add=large:afp=40000:afq=1.0:amm=off:anc=none:cond=on:gs=on:irw=on:nm=16:newcnf=on:nwc=1:stl=30:sos=on:sp=reverse_arity:updr=off_6 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit

% Memory used [KB]: 9850
% Time elapsed: 1.783 s
% ------------------------------
% ------------------------------
% spawned child 12835 with time: 54 (total remaining time 2819)
% lrs+1011_3:1_av=off:cond=on:irw=on:lma=on:nm=16:nwc=1:stl=30:sos=all:updr=off_51 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 12025
% Time elapsed: 5.500 s
% ------------------------------
% ------------------------------
% spawned child 12836 with time: 7 (total remaining time 2771)
% lrs+1_3:2_aac=none:add=large:anc=all_dependent:bce=on:cond=fast:ep=RST:fsr=off:lma=on:nm=2:nwc=1:stl=30:sos=on:sp=occurrence:urr=on:updr=off:uhcvi=on_5 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 5500
% Time elapsed: 0.869 s
% ------------------------------
% ------------------------------
% spawned child 12837 with time: 11 (total remaining time 2762)
% ott+1_8_av=off:bd=off:bs=on:cond=on:gsp=input_only:gs=on:gsem=off:irw=on:lcm=predicate:lwlo=on:nwc=1:sos=on_10 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 10618
% Time elapsed: 1.200 s
% ------------------------------
% ------------------------------
% spawned child 12840 with time: 7 (total remaining time 2751)
% dis+11_5_afp=40000:afq=1.4:anc=none:br=off:cond=on:fsr=off:irw=on:lma=on:nm=4:nwc=1:sd=3:ss=axioms:st=2.0:urr=on:updr=off_5 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7803
% Time elapsed: 0.800 s
% ------------------------------
% ------------------------------
% spawned child 12844 with time: 4 (total remaining time 2743)
% dis+1010_10_afr=on:afp=100000:afq=2.0:amm=sco:anc=none:ep=RS:fde=unused:gs=on:nwc=1:sos=on:sac=on:sp=occurrence_2 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 11129
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 12845 with time: 9 (total remaining time 2738)
% lrs+1011_8:1_afr=on:afp=1000:afq=2.0:br=off:gsp=input_only:gs=on:nm=16:nwc=1:stl=30:sos=all:sp=occurrence:urr=on_7 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 14711
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% spawned child 12846 with time: 9 (total remaining time 2729)
% lrs+1003_2_acc=on:afp=4000:afq=2.0:amm=sco:bd=off:ccuc=small_ones:fsr=off:fde=unused:gsp=input_only:nm=64:newcnf=on:nwc=2.5:nicw=on:stl=30:urr=ec_only_7 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8827
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% spawned child 12848 with time: 4 (total remaining time 2719)
% ott+4_40_av=off:bce=on:cond=fast:fde=none:nm=0:nwc=1:sos=all:updr=off_2 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 2686
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% spawned child 12849 with time: 5 (total remaining time 2713)
% dis-3_5_av=off:cond=on:gs=on:gsem=off:lcm=predicate:nm=0:nwc=1:sos=on:sp=reverse_arity_3 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7931
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% spawned child 12850 with time: 56 (total remaining time 2707)
% lrs+1011_14_av=off:fsr=off:irw=on:nwc=1:stl=30:sos=on:sp=occurrence:urr=ec_only:updr=off_53 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 13304
% Time elapsed: 5.700 s
% ------------------------------
% ------------------------------
% spawned child 12857 with time: 5 (total remaining time 2656)
% dis+11_10_av=off:lma=on:nm=64:nwc=5:sp=reverse_arity_3 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit

% Memory used [KB]: 767
% Time elapsed: 0.724 s
% ------------------------------
% ------------------------------
% spawned child 12858 with time: 13 (total remaining time 2648)
% lrs+1011_4:1_av=off:fsr=off:irw=on:nwc=1:stl=30:sd=4:ss=axioms:st=1.5:sp=reverse_arity_12 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 6140
% Time elapsed: 1.400 s
% ------------------------------
% ------------------------------
% spawned child 12864 with time: 245 (total remaining time 2634)
% lrs+10_2:3_afp=1000:afq=1.1:amm=sco:anc=none:er=known:lcm=reverse:lwlo=on:nm=64:newcnf=on:nwc=2:stl=30:sd=5:ss=axioms:sos=theory:sac=on:sp=occurrence_233 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 37355
% Time elapsed: 24.600 s
% ------------------------------
% ------------------------------
% spawned child 12867 with time: 7 (total remaining time 2424)
% lrs+1011_3:2_av=off:er=known:lma=on:nm=2:newcnf=on:nwc=2:stl=30:sd=2:ss=axioms:st=3.0:urr=on:updr=off_5 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit

% Memory used [KB]: 767
% Time elapsed: 3.708 s
% ------------------------------
% ------------------------------
% spawned child 12868 with time: 38 (total remaining time 2386)
% lrs+1011_2:3_av=off:gs=on:gsem=off:nwc=1.5:stl=30:sos=theory:sp=occurrence:urr=ec_only:updr=off_36 on codi_down_sum_theorems_goal12.all
% Time limit reached!
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 20468
% Time elapsed: 3.900 s
% ------------------------------
% ------------------------------
% spawned child 12879 with time: 62 (total remaining time 2350)
% lrs+11_4_av=off:gsp=input_only:irw=on:lma=on:nm=0:nwc=1.2:stl=30:sd=2:ss=axioms:sp=reverse_arity:urr=on:updr=off_59 on codi_down_sum_theorems_goal12.all
% Refutation found. Thanks to Tanya!
% SZS status Theorem for codi_down_sum_theorems_goal12.all
% SZS output start Proof for codi_down_sum_theorems_goal12.all
fof(f1,conjecture,(
  ! [X0] : ! [X1] : ! [X2] : (p(intersection(X2,X0),X2) | ~(~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',conjecture120)).
fof(f2,negated_conjecture,(
  ~! [X0] : ! [X1] : ! [X2] : (p(intersection(X2,X0),X2) | ~(~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)))),
  inference(negated_conjecture,[],[f1])).
fof(f6,axiom,(
  ! [X0] : ! [X1] : ! [X3] : (cont(difference(X3,X0),X1) | ~(~cont(X3,X0) & cont(X3,sum(X0,X1))))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom40)).
fof(f53,axiom,(
  ! [X0] : (p(X0,intersection(X0,X0)) | ~(~zex(intersection(X0,X0)) & ~zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom87070)).
fof(f55,axiom,(
  ! [X0] : (zex(X0) | ~zex(intersection(X0,X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom87090)).
fof(f66,axiom,(
  ! [X0] : ! [X1] : (cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom87200)).
fof(f67,axiom,(
  ! [X0] : ! [X1] : ! [X3] : (leq(X3,intersection(X0,X1)) | ~(cont(X3,X1) & cont(X3,X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom87210)).
fof(f77,axiom,(
  ! [X0] : ! [X1] : (eqdim(X1,X0) | ~eqdim(X0,X1))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom89430)).
fof(f84,axiom,(
  ! [X0] : ! [X1] : (lt(X0,X1) | ~(~zex(X1) & zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom89500)).
fof(f85,axiom,(
  ! [X0] : ((~zex(X0) | ~cont(X0,X0)) & (cont(X0,X0) | zex(X0)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom89510)).
fof(f86,axiom,(
  ! [X0] : ! [X1] : (X0 = X1 | ~(cont(X1,X0) & cont(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom89520)).
fof(f87,axiom,(
  ! [X0] : ! [X1] : ! [X3] : (cont(X0,X3) | ~(cont(X1,X3) & cont(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom89530)).
fof(f88,axiom,(
  ! [X0] : ! [X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom89540)).
fof(f92,axiom,(
  ! [X0] : ! [X1] : ((p(X0,X1) | ~(eqdim(X0,X1) & cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom88840)).
fof(f93,axiom,(
  ! [X0] : ! [X1] : (eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom88850)).
fof(f97,axiom,(
  ! [X0] : ! [X1] : ! [X3] : (leq(X0,X3) | ~(leq(X1,X3) & leq(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom89120)).
fof(f106,axiom,(
  ! [X0] : ! [X1] : ((eqdim(X0,X1) | ~(leq(X1,X0) & leq(X0,X1))) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom89310)).
fof(f108,axiom,(
  ! [X0] : ! [X1] : (eqdim(X0,difference(X0,X1)) | zex(difference(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom31660)).
fof(f110,axiom,(
  ! [X0] : ! [X1] : ! [X3] : (cont(X3,difference(X0,X1)) | ~(lt(intersection(X3,X1),X3) & cont(X3,X0) & leq(X0,X1)))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom31680)).
fof(f113,axiom,(
  ! [X0] : ! [X1] : ((zex(difference(X0,X1)) | ~(cont(X0,X1) | zex(X0))) & (cont(X0,X1) | zex(X0) | ~zex(difference(X0,X1))))),
  file('Dropbox/Experiments/codi_down_theoremsT6-T10_goals/codi_down_sum_theorems_goal12.all.tptp',axiom31710)).
fof(f118,plain,(
  ~! [X0,X1,X2] : (p(intersection(X2,X0),X2) | ~(~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)))),
  inference(flattening,[],[f2])).
fof(f138,plain,(
  ! [X0] : ! [X1] : ! [X2] : (cont(X0,X2) | ~(cont(X1,X2) & cont(X0,X1)))),
  inference(rectify,[],[f87])).
fof(f139,plain,(
  ! [X0,X1,X2] : (cont(X0,X2) | ~(cont(X1,X2) & cont(X0,X1)))),
  inference(flattening,[],[f138])).
fof(f140,plain,(
  ! [X0,X1] : (cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)))),
  inference(flattening,[],[f66])).
fof(f141,plain,(
  ! [X0,X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0))),
  inference(flattening,[],[f88])).
fof(f142,plain,(
  ! [X0,X1] : ((zex(difference(X0,X1)) | ~(cont(X0,X1) | zex(X0))) & (cont(X0,X1) | zex(X0) | ~zex(difference(X0,X1))))),
  inference(flattening,[],[f113])).
fof(f144,plain,(
  ! [X0] : ! [X1] : ! [X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  inference(rectify,[],[f6])).
fof(f145,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | ~(~cont(X2,X0) & cont(X2,sum(X0,X1))))),
  inference(flattening,[],[f144])).
fof(f160,plain,(
  ! [X0,X1] : (eqdim(X0,difference(X0,X1)) | zex(difference(X0,X1)))),
  inference(flattening,[],[f108])).
fof(f162,plain,(
  ! [X0,X1] : (eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1))),
  inference(flattening,[],[f93])).
fof(f165,plain,(
  ! [X0,X1] : (eqdim(X1,X0) | ~eqdim(X0,X1))),
  inference(flattening,[],[f77])).
fof(f166,plain,(
  ! [X0,X1] : ((p(X0,X1) | ~(eqdim(X0,X1) & cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  inference(flattening,[],[f92])).
fof(f186,plain,(
  ! [X0] : ! [X1] : ! [X2] : (leq(X2,intersection(X0,X1)) | ~(cont(X2,X1) & cont(X2,X0)))),
  inference(rectify,[],[f67])).
fof(f187,plain,(
  ! [X0,X1,X2] : (leq(X2,intersection(X0,X1)) | ~(cont(X2,X1) & cont(X2,X0)))),
  inference(flattening,[],[f186])).
fof(f196,plain,(
  ! [X0] : ! [X1] : ! [X2] : (cont(X2,difference(X0,X1)) | ~(lt(intersection(X2,X1),X2) & cont(X2,X0) & leq(X0,X1)))),
  inference(rectify,[],[f110])).
fof(f197,plain,(
  ! [X0,X1,X2] : (cont(X2,difference(X0,X1)) | ~(lt(intersection(X2,X1),X2) & cont(X2,X0) & leq(X0,X1)))),
  inference(flattening,[],[f196])).
fof(f200,plain,(
  ! [X0] : ! [X1] : ! [X2] : (leq(X0,X2) | ~(leq(X1,X2) & leq(X0,X1)))),
  inference(rectify,[],[f97])).
fof(f201,plain,(
  ! [X0,X1,X2] : (leq(X0,X2) | ~(leq(X1,X2) & leq(X0,X1)))),
  inference(flattening,[],[f200])).
fof(f205,plain,(
  ! [X0,X1] : ((eqdim(X0,X1) | ~(leq(X1,X0) & leq(X0,X1))) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1)))),
  inference(flattening,[],[f106])).
fof(f210,plain,(
  ! [X0,X1] : (X0 = X1 | ~(cont(X1,X0) & cont(X0,X1)))),
  inference(flattening,[],[f86])).
fof(f219,plain,(
  ! [X0,X1] : (lt(X0,X1) | ~(~zex(X1) & zex(X0)))),
  inference(flattening,[],[f84])).
fof(f223,plain,(
  ? [X0,X1,X2] : (~p(intersection(X2,X0),X2) & (~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f118])).
fof(f224,plain,(
  ? [X0,X1,X2] : (~p(intersection(X2,X0),X2) & ~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1))),
  inference(flattening,[],[f223])).
fof(f229,plain,(
  ! [X0] : (p(X0,intersection(X0,X0)) | (zex(intersection(X0,X0)) | zex(X0)))),
  inference(ennf_transformation,[],[f53])).
fof(f230,plain,(
  ! [X0] : (p(X0,intersection(X0,X0)) | zex(intersection(X0,X0)) | zex(X0))),
  inference(flattening,[],[f229])).
fof(f231,plain,(
  ! [X0,X1,X2] : (cont(X0,X2) | (~cont(X1,X2) | ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f139])).
fof(f232,plain,(
  ! [X0,X1,X2] : (cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1))),
  inference(flattening,[],[f231])).
fof(f233,plain,(
  ! [X0,X1] : ((zex(difference(X0,X1)) | (~cont(X0,X1) & ~zex(X0))) & (cont(X0,X1) | zex(X0) | ~zex(difference(X0,X1))))),
  inference(ennf_transformation,[],[f142])).
fof(f235,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | (cont(X2,X0) | ~cont(X2,sum(X0,X1))))),
  inference(ennf_transformation,[],[f145])).
fof(f236,plain,(
  ! [X0,X1,X2] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1)))),
  inference(flattening,[],[f235])).
fof(f245,plain,(
  ! [X0,X1] : ((p(X0,X1) | (~eqdim(X0,X1) | ~cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  inference(ennf_transformation,[],[f166])).
fof(f246,plain,(
  ! [X0,X1] : ((p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1)) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1)))),
  inference(flattening,[],[f245])).
fof(f259,plain,(
  ! [X0,X1,X2] : (leq(X2,intersection(X0,X1)) | (~cont(X2,X1) | ~cont(X2,X0)))),
  inference(ennf_transformation,[],[f187])).
fof(f260,plain,(
  ! [X0,X1,X2] : (leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0))),
  inference(flattening,[],[f259])).
fof(f269,plain,(
  ! [X0,X1,X2] : (cont(X2,difference(X0,X1)) | (~lt(intersection(X2,X1),X2) | ~cont(X2,X0) | ~leq(X0,X1)))),
  inference(ennf_transformation,[],[f197])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cont(X2,difference(X0,X1)) | ~lt(intersection(X2,X1),X2) | ~cont(X2,X0) | ~leq(X0,X1))),
  inference(flattening,[],[f269])).
fof(f273,plain,(
  ! [X0,X1,X2] : (leq(X0,X2) | (~leq(X1,X2) | ~leq(X0,X1)))),
  inference(ennf_transformation,[],[f201])).
fof(f274,plain,(
  ! [X0,X1,X2] : (leq(X0,X2) | ~leq(X1,X2) | ~leq(X0,X1))),
  inference(flattening,[],[f273])).
fof(f277,plain,(
  ! [X0,X1] : ((eqdim(X0,X1) | (~leq(X1,X0) | ~leq(X0,X1))) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1)))),
  inference(ennf_transformation,[],[f205])).
fof(f278,plain,(
  ! [X0,X1] : ((eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1)) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1)))),
  inference(flattening,[],[f277])).
fof(f281,plain,(
  ! [X0,X1] : (X0 = X1 | (~cont(X1,X0) | ~cont(X0,X1)))),
  inference(ennf_transformation,[],[f210])).
fof(f282,plain,(
  ! [X0,X1] : (X0 = X1 | ~cont(X1,X0) | ~cont(X0,X1))),
  inference(flattening,[],[f281])).
fof(f291,plain,(
  ! [X0,X1] : (lt(X0,X1) | (zex(X1) | ~zex(X0)))),
  inference(ennf_transformation,[],[f219])).
fof(f292,plain,(
  ! [X0,X1] : (lt(X0,X1) | zex(X1) | ~zex(X0))),
  inference(flattening,[],[f291])).
fof(f295,plain,(
  ? [X0,X1,X2] : (~p(intersection(X2,X0),X2) & ~zex(X1) & ~zex(X0) & ~cont(X2,X1) & ~cont(X2,X0) & cont(X2,sum(X0,X1)) & eqdim(X0,X1)) => (~p(intersection(sK2,sK0),sK2) & ~zex(sK1) & ~zex(sK0) & ~cont(sK2,sK1) & ~cont(sK2,sK0) & cont(sK2,sum(sK0,sK1)) & eqdim(sK0,sK1))),
  introduced(choice_axiom,[])).
fof(f296,plain,(
  ~p(intersection(sK2,sK0),sK2) & ~zex(sK1) & ~zex(sK0) & ~cont(sK2,sK1) & ~cont(sK2,sK0) & cont(sK2,sum(sK0,sK1)) & eqdim(sK0,sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f224,f295])).
fof(f313,plain,(
  eqdim(sK0,sK1)),
  inference(cnf_transformation,[],[f296])).
fof(f314,plain,(
  cont(sK2,sum(sK0,sK1))),
  inference(cnf_transformation,[],[f296])).
fof(f315,plain,(
  ~cont(sK2,sK0)),
  inference(cnf_transformation,[],[f296])).
fof(f316,plain,(
  ~cont(sK2,sK1)),
  inference(cnf_transformation,[],[f296])).
fof(f318,plain,(
  ~zex(sK1)),
  inference(cnf_transformation,[],[f296])).
fof(f319,plain,(
  ~p(intersection(sK2,sK0),sK2)),
  inference(cnf_transformation,[],[f296])).
fof(f325,plain,(
  ( ! [X0] : (p(X0,intersection(X0,X0)) | zex(intersection(X0,X0)) | zex(X0)) )),
  inference(cnf_transformation,[],[f230])).
fof(f335,plain,(
  ( ! [X0] : (~zex(intersection(X0,X0)) | zex(X0)) )),
  inference(cnf_transformation,[],[f55])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2)) )),
  inference(cnf_transformation,[],[f232])).
fof(f339,plain,(
  ( ! [X0,X1] : (cont(intersection(X0,X1),X0) | zex(intersection(X0,X1))) )),
  inference(cnf_transformation,[],[f140])).
fof(f341,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~zex(X0)) )),
  inference(cnf_transformation,[],[f141])).
fof(f342,plain,(
  ( ! [X0,X1] : (~zex(difference(X0,X1)) | zex(X0) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f348,plain,(
  ( ! [X0] : (cont(X0,X0) | zex(X0)) )),
  inference(cnf_transformation,[],[f85])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (cont(difference(X2,X0),X1) | cont(X2,X0) | ~cont(X2,sum(X0,X1))) )),
  inference(cnf_transformation,[],[f236])).
fof(f362,plain,(
  ( ! [X0,X1] : (eqdim(X0,difference(X0,X1)) | zex(difference(X0,X1))) )),
  inference(cnf_transformation,[],[f160])).
fof(f364,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f162])).
fof(f367,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0)) )),
  inference(cnf_transformation,[],[f165])).
fof(f368,plain,(
  ( ! [X0,X1] : (~p(X0,X1) | cont(X0,X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f370,plain,(
  ( ! [X0,X1] : (~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (leq(X2,intersection(X0,X1)) | ~cont(X2,X1) | ~cont(X2,X0)) )),
  inference(cnf_transformation,[],[f260])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (~lt(intersection(X2,X1),X2) | cont(X2,difference(X0,X1)) | ~cont(X2,X0) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f270])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (~leq(X0,X1) | ~leq(X1,X2) | leq(X0,X2)) )),
  inference(cnf_transformation,[],[f274])).
fof(f396,plain,(
  ( ! [X0,X1] : (leq(X0,X1) | ~eqdim(X0,X1)) )),
  inference(cnf_transformation,[],[f278])).
fof(f408,plain,(
  ( ! [X0,X1] : (~cont(X0,X1) | ~cont(X1,X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f282])).
fof(f418,plain,(
  ( ! [X0,X1] : (~zex(X0) | zex(X1) | lt(X0,X1)) )),
  inference(cnf_transformation,[],[f292])).
fof(f432,plain,(
  ~zex(sK2)),
  inference(unit_resulting_resolution,[],[f314,f341])).
fof(f437,plain,(
  cont(sK2,sK2)),
  inference(unit_resulting_resolution,[],[f432,f348])).
fof(f480,plain,(
  ~zex(intersection(sK1,sK1))),
  inference(unit_resulting_resolution,[],[f318,f335])).
fof(f520,plain,(
  eqdim(sK1,sK0)),
  inference(unit_resulting_resolution,[],[f313,f367])).
fof(f565,plain,(
  leq(sK1,sK0)),
  inference(unit_resulting_resolution,[],[f520,f396])).
fof(f1019,plain,(
  cont(intersection(sK1,sK1),sK1)),
  inference(unit_resulting_resolution,[],[f480,f339])).
fof(f1118,plain,(
  ~zex(difference(sK2,sK0))),
  inference(unit_resulting_resolution,[],[f315,f432,f342])).
fof(f1379,plain,(
  eqdim(sK2,difference(sK2,sK0))),
  inference(unit_resulting_resolution,[],[f1118,f362])).
fof(f1473,plain,(
  leq(sK2,difference(sK2,sK0))),
  inference(unit_resulting_resolution,[],[f1379,f396])).
fof(f2194,plain,(
  p(sK1,intersection(sK1,sK1))),
  inference(unit_resulting_resolution,[],[f318,f480,f325])).
fof(f2220,plain,(
  cont(sK1,intersection(sK1,sK1))),
  inference(unit_resulting_resolution,[],[f2194,f368])).
fof(f2636,plain,(
  intersection(sK1,sK1) = sK1),
  inference(unit_resulting_resolution,[],[f1019,f2220,f408])).
fof(f3463,plain,(
  cont(difference(sK2,sK0),sK1)),
  inference(unit_resulting_resolution,[],[f315,f314,f350])).
fof(f3505,plain,(
  ~cont(sK2,difference(sK2,sK0))),
  inference(unit_resulting_resolution,[],[f316,f3463,f338])).
fof(f3517,plain,(
  leq(difference(sK2,sK0),intersection(sK1,sK1))),
  inference(unit_resulting_resolution,[],[f3463,f3463,f383])).
fof(f3523,plain,(
  leq(difference(sK2,sK0),sK1)),
  inference(forward_demodulation,[],[f3517,f2636])).
fof(f3552,plain,(
  leq(sK2,sK1)),
  inference(unit_resulting_resolution,[],[f1473,f3523,f390])).
fof(f3566,plain,(
  leq(sK2,sK0)),
  inference(unit_resulting_resolution,[],[f565,f3552,f390])).
fof(f4195,plain,(
  ~lt(intersection(sK2,sK0),sK2)),
  inference(unit_resulting_resolution,[],[f3566,f437,f3505,f388])).
fof(f4828,plain,(
  ~zex(intersection(sK2,sK0))),
  inference(unit_resulting_resolution,[],[f432,f4195,f418])).
fof(f4835,plain,(
  cont(intersection(sK2,sK0),sK2)),
  inference(unit_resulting_resolution,[],[f4828,f339])).
fof(f5129,plain,(
  eqdim(intersection(sK2,sK0),sK2)),
  inference(unit_resulting_resolution,[],[f4195,f4835,f364])).
fof(f5130,plain,(
  ~eqdim(intersection(sK2,sK0),sK2)),
  inference(unit_resulting_resolution,[],[f319,f4835,f370])).
fof(f5140,plain,(
  $false),
  inference(subsumption_resolution,[],[f5129,f5130])).
% SZS output end Proof for codi_down_sum_theorems_goal12.all
% ------------------------------
% Version: Vampire 4.2.2 (commit e1949dd on 2017-12-14 18:39:21 +0000)
% Termination reason: Refutation

% Memory used [KB]: 2046
% Time elapsed: 0.752 s
% ------------------------------
% ------------------------------
% Success in time 65.78 s
