Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal10.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 674
Passive clauses: 2524
Generated clauses: 5511
Final active clauses: 652
Final passive clauses: 1730
Input formulas: 43
Initial clauses: 90

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 21
Fw subsumption resolutions: 137

Simple tautologies: 77
Forward subsumptions: 2717

Binary resolution: 5115
Factoring: 23

Backtracking splits: 89
Backtracking splits refuted: 16
Backtracking splits refuted at zero level: 5

Memory used: 1023KB
Time elapsed: 0.401 s
------------------------------
% remaining time: 5995 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal10.all
Refutation found. Thanks to Tanya!
900. $false (5:0) [subsumption resolution 892,436]
436. ~zex(sK1) (3:2) [subsumption resolution 430,330]
330. zex(sK0) | ~zex(sK1) (2:4) [resolution 240,311]
311. ~lt(sK1,sK0) (1:3) [resolution 236,215]
215. lt(sK0,sK1) (0:3) [cnf transformation 169]
169. inc(sK0,sK1) & lt(sK0,sK1) & ~lt(intersection(sK0,sK1),sK1)[skolemisation 121]
121. ? [X0,X1] : (inc(X0,X1) & lt(X0,X1) & ~lt(intersection(X0,X1),X1))[flattening 120]
120. ? [X0,X1] : ((inc(X0,X1) & lt(X0,X1)) & ~lt(intersection(X0,X1),X1))[ennf transformation 44]
44. ~! [X0,X1] : ((inc(X0,X1) & lt(X0,X1)) => lt(intersection(X0,X1),X1))[flattening 1]
1. ~! [X0] : ! [X1] : ((inc(X0,X1) & lt(X0,X1)) => lt(intersection(X0,X1),X1))[negated conjecture]
236. ~lt(X1,X0) | ~lt(X0,X1) (0:6) [cnf transformation 134]
134. ! [X0,X1] : (~lt(X0,X1) | ~lt(X1,X0))[ennf transformation 64]
64. ! [X0,X1] : (lt(X0,X1) => ~lt(X1,X0))[flattening 63]
63. ! [X0] : ! [X1] : (lt(X0,X1) => ~lt(X1,X0))[rectify 26]
26. ! [X53] : ! [X54] : (lt(X53,X54) => ~lt(X54,X53))[input]
240. lt(X0,X1) | zex(X1) | ~zex(X0) (0:7) [cnf transformation 142]
142. ! [X0,X1] : (~zex(X0) | zex(X1) | lt(X0,X1))[flattening 141]
141. ! [X0,X1] : ((~zex(X0) | zex(X1)) | lt(X0,X1))[ennf transformation 72]
72. ! [X0,X1] : ((zex(X0) & ~zex(X1)) => lt(X0,X1))[flattening 71]
71. ! [X0] : ! [X1] : ((zex(X0) & ~zex(X1)) => lt(X0,X1))[rectify 36]
36. ! [X76] : ! [X77] : ((zex(X76) & ~zex(X77)) => lt(X76,X77))[input]
430. ~zex(sK1) | ~zex(sK0) (3:4) [resolution 350,228]
228. ~lt(X0,X0) (0:3) [cnf transformation 125]
125. ! [X0] : ~lt(X0,X0)[ennf transformation 50]
50. ! [X0] : ~lt(X0,X0)[rectify 25]
25. ! [X52] : ~lt(X52,X52)[input]
350. lt(X0,sK1) | ~zex(X0) | ~zex(sK0) (2:7) [superposition 215,333]
333. X0 = X1 | ~zex(X1) | ~zex(X0) (1:7) [resolution 297,304]
304. ~sQ1_eqProxy(X0,X1) | X0 = X1 (0:6) [equality proxy axiom]
297. sQ1_eqProxy(X0,X1) | ~zex(X1) | ~zex(X0) (0:7) [equality proxy replacement 237]
237. X0 = X1 | ~zex(X1) | ~zex(X0) (0:7) [cnf transformation 136]
136. ! [X0,X1] : (~zex(X0) | ~zex(X1) | X0 = X1)[flattening 135]
135. ! [X0,X1] : ((~zex(X0) | ~zex(X1)) | X0 = X1)[ennf transformation 66]
66. ! [X0,X1] : ((zex(X0) & zex(X1)) => X0 = X1)[flattening 65]
65. ! [X0] : ! [X1] : ((zex(X0) & zex(X1)) => X0 = X1)[rectify 40]
40. ! [X84] : ! [X85] : ((zex(X84) & zex(X85)) => X84 = X85)[input]
892. zex(sK1) (5:2) [resolution 889,329]
329. ~zex(intersection(sK0,sK1)) | zex(sK1) (1:6) [resolution 240,216]
216. ~lt(intersection(sK0,sK1),sK1) (0:5) [cnf transformation 169]
889. zex(intersection(sK0,sK1)) (4:4) [subsumption resolution 879,215]
879. ~lt(sK0,sK1) | zex(intersection(sK0,sK1)) (4:7) [resolution 777,242]
242. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 145]
145. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 76]
76. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 75]
75. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 21]
21. ! [X42] : ! [X43] : (~zex(intersection(X42,X43)) => cont(intersection(X42,X43),X42))[input]
777. ~cont(intersection(sK0,sK1),X4) | ~lt(X4,sK1) (3:8) [subsumption resolution 767,690]
690. ~lt(intersection(sK0,sK1),X16) | ~lt(X16,sK1) (1:8) [resolution 288,216]
288. lt(X0,X2) | ~lt(X1,X2) | ~lt(X0,X1) (0:9) [cnf transformation 165]
165. ! [X0,X1,X2] : (~lt(X0,X1) | ~lt(X1,X2) | lt(X0,X2))[flattening 164]
164. ! [X0,X1,X2] : ((~lt(X0,X1) | ~lt(X1,X2)) | lt(X0,X2))[ennf transformation 114]
114. ! [X0,X1,X2] : ((lt(X0,X1) & lt(X1,X2)) => lt(X0,X2))[flattening 113]
113. ! [X0] : ! [X1] : ! [X2] : ((lt(X0,X1) & lt(X1,X2)) => lt(X0,X2))[rectify 27]
27. ! [X55] : ! [X56] : ! [X57] : ((lt(X55,X56) & lt(X56,X57)) => lt(X55,X57))[input]
767. ~lt(X4,sK1) | lt(intersection(sK0,sK1),X4) | ~cont(intersection(sK0,sK1),X4) (3:13) [resolution 710,232]
232. eqdim(X0,X1) | lt(X0,X1) | ~cont(X0,X1) (0:9) [cnf transformation 129]
129. ! [X0,X1] : (~cont(X0,X1) | lt(X0,X1) | eqdim(X0,X1))[flattening 128]
128. ! [X0,X1] : (~cont(X0,X1) | (lt(X0,X1) | eqdim(X0,X1)))[ennf transformation 56]
56. ! [X0,X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[flattening 55]
55. ! [X0] : ! [X1] : (cont(X0,X1) => (lt(X0,X1) | eqdim(X0,X1)))[rectify 24]
24. ! [X50] : ! [X51] : (cont(X50,X51) => (lt(X50,X51) | eqdim(X50,X51)))[input]
710. ~eqdim(intersection(sK0,sK1),X6) | ~lt(X6,sK1) (2:8) [resolution 608,234]
234. eqdim(X1,X0) | ~eqdim(X0,X1) (0:6) [cnf transformation 132]
132. ! [X0,X1] : (~eqdim(X0,X1) | eqdim(X1,X0))[ennf transformation 60]
60. ! [X0,X1] : (eqdim(X0,X1) => eqdim(X1,X0))[flattening 59]
59. ! [X0] : ! [X1] : (eqdim(X0,X1) => eqdim(X1,X0))[rectify 29]
29. ! [X59] : ! [X60] : (eqdim(X59,X60) => eqdim(X60,X59))[input]
608. ~eqdim(X16,intersection(sK0,sK1)) | ~lt(X16,sK1) (1:8) [resolution 282,216]
282. lt(X1,X2) | ~lt(X0,X2) | ~eqdim(X0,X1) (0:9) [cnf transformation 153]
153. ! [X0,X1,X2] : (~eqdim(X0,X1) | ~lt(X0,X2) | lt(X1,X2))[flattening 152]
152. ! [X0,X1,X2] : (~eqdim(X0,X1) | (~lt(X0,X2) | lt(X1,X2)))[ennf transformation 102]
102. ! [X0,X1,X2] : (eqdim(X0,X1) => (lt(X0,X2) => lt(X1,X2)))[flattening 101]
101. ! [X0] : ! [X1] : ! [X2] : (eqdim(X0,X1) => (lt(X0,X2) => lt(X1,X2)))[rectify 33]
33. ! [X69] : ! [X70] : ! [X71] : (eqdim(X69,X70) => (lt(X69,X71) => lt(X70,X71)))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 137
Passive clauses: 362
Generated clauses: 681
Final active clauses: 137
Final passive clauses: 199
Input formulas: 43
Initial clauses: 85

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 29
Fw subsumption resolutions: 48

Simple tautologies: 33
Equational tautologies: 1
Forward subsumptions: 229

Binary resolution: 297
Factoring: 1
Forward superposition: 162
Backward superposition: 60

Unique components: 3

Memory used: 255KB
Time elapsed: 0.034 s
------------------------------
% Success in time 0.583 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 23:33:47 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal10.all.tptp"
Input read from 
============================ end of footer ===========================
