Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 5999 next slice time: 3
dis-4_5_bd=off:bs=off:ep=RST:fde=none:lcm=predicate:n=on:nicw=on:nwc=2:ptb=off:sio=off:spl=backtracking:ss=included:ssec=off:st=5:updr=off_3 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal21.all
Time limit reached!
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Time limit

Active clauses: 523
Passive clauses: 1654
Generated clauses: 3748
Final active clauses: 523
Final passive clauses: 1064
Input formulas: 43
Initial clauses: 89

Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 33
Fw subsumption resolutions: 26

Simple tautologies: 27
Forward subsumptions: 2063

Binary resolution: 3555
Factoring: 36

Backtracking splits: 13

Memory used: 767KB
Time elapsed: 0.402 s
------------------------------
% remaining time: 5993 next slice time: 13
dis-1010_2_bs=off:ep=on:n=on:nwc=1.5:sgo=on:sp=occurrence:ss=included:sswn=on:sswsr=on:st=1.5_13 on C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal21.all
Refutation found. Thanks to Tanya!
2646. $false (7:0) [resolution 2645,326]
326. c(y,intersection(sK0,sK0)) (1:5) [resolution 277,213]
213. ~zex(intersection(y,intersection(sK0,sK0))) (0:6) [cnf transformation 167]
167. zex(intersection(sK0,y)) & ~zex(intersection(y,intersection(sK0,sK0)))[skolemisation 119]
119. ? [X0] : (zex(intersection(X0,y)) & ~zex(intersection(y,intersection(X0,X0))))[ennf transformation 1]
1. ~! [X0] : (zex(intersection(X0,y)) => zex(intersection(y,intersection(X0,X0))))[negated conjecture]
277. zex(intersection(X0,X1)) | c(X0,X1) (0:7) [cnf transformation 206]
206. ! [X0,X1] : ((c(X0,X1) | zex(intersection(X0,X1))) & (~zex(intersection(X0,X1)) | ~c(X0,X1)))[nnf transformation 149]
149. ! [X0,X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[ennf transformation 99]
99. ! [X0,X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[flattening 98]
98. ! [X0] : ! [X1] : (~c(X0,X1) <=> zex(intersection(X0,X1)))[rectify 20]
20. ! [X41] : ! [X42] : (~c(X41,X42) <=> zex(intersection(X41,X42)))[input]
2645. ~c(y,intersection(sK0,X3)) (6:5) [subsumption resolution 2624,501]
501. ~c(X2,X3) | ~zex(X3) (1:5) [resolution 266,227]
227. ~cont(X1,X0) | ~zex(X0) (0:5) [cnf transformation 125]
125. ! [X0,X1] : (~zex(X0) | (~cont(X1,X0) & ~cont(X0,X1)))[ennf transformation 53]
53. ! [X0,X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))[flattening 52]
52. ! [X0] : ! [X1] : (zex(X0) => (~cont(X1,X0) & ~cont(X0,X1)))[rectify 16]
16. ! [X32] : ! [X33] : (zex(X32) => (~cont(X33,X32) & ~cont(X32,X33)))[input]
266. cont(sK10(X1,X0),X1) | ~c(X0,X1) (0:8) [cnf transformation 199]
199. ! [X0,X1] : ((~c(X0,X1) | (cont(sK10(X1,X0),X0) & cont(sK10(X1,X0),X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[skolemisation 198]
198. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X3] : (~cont(X3,X0) | ~cont(X3,X1)) | c(X0,X1)))[rectify 197]
197. ! [X0,X1] : ((~c(X0,X1) | ? [X2] : (cont(X2,X0) & cont(X2,X1))) & (! [X2] : (~cont(X2,X0) | ~cont(X2,X1)) | c(X0,X1)))[nnf transformation 91]
91. ! [X0,X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))[flattening 90]
90. ! [X0] : ! [X1] : (c(X0,X1) <=> ? [X2] : (cont(X2,X0) & cont(X2,X1)))[rectify 2]
2. ! [X1] : ! [X2] : (c(X1,X2) <=> ? [X3] : (cont(X3,X1) & cont(X3,X2)))[input]
2624. ~c(y,intersection(sK0,X3)) | zex(intersection(sK0,X3)) (6:9) [resolution 2607,239]
239. cont(intersection(X0,X1),X0) | zex(intersection(X0,X1)) (0:9) [cnf transformation 143]
143. ! [X0,X1] : (zex(intersection(X0,X1)) | cont(intersection(X0,X1),X0))[ennf transformation 75]
75. ! [X0,X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[flattening 74]
74. ! [X0] : ! [X1] : (~zex(intersection(X0,X1)) => cont(intersection(X0,X1),X0))[rectify 21]
21. ! [X43] : ! [X44] : (~zex(intersection(X43,X44)) => cont(intersection(X43,X44),X43))[input]
2607. ~cont(X1,sK0) | ~c(y,X1) (5:6) [duplicate literal removal 2599]
2599. ~cont(X1,sK0) | ~c(y,X1) | ~c(y,X1) (5:9) [resolution 963,266]
963. ~cont(sK10(X16,y),X17) | ~cont(X17,sK0) | ~c(y,X16) (4:11) [resolution 675,265]
265. cont(sK10(X1,X0),X0) | ~c(X0,X1) (0:8) [cnf transformation 199]
675. ~cont(X9,y) | ~cont(X9,X8) | ~cont(X8,sK0) (3:9) [resolution 281,557]
557. ~cont(X0,sK0) | ~cont(X0,y) (2:6) [resolution 267,327]
327. ~c(sK0,y) (1:3) [resolution 278,212]
212. zex(intersection(sK0,y)) (0:4) [cnf transformation 167]
278. ~zex(intersection(X0,X1)) | ~c(X0,X1) (0:7) [cnf transformation 206]
267. c(X0,X1) | ~cont(X3,X1) | ~cont(X3,X0) (0:9) [cnf transformation 199]
281. cont(X0,X2) | ~cont(X1,X2) | ~cont(X0,X1) (0:9) [cnf transformation 155]
155. ! [X0,X1,X2] : (~cont(X0,X1) | ~cont(X1,X2) | cont(X0,X2))[flattening 154]
154. ! [X0,X1,X2] : ((~cont(X0,X1) | ~cont(X1,X2)) | cont(X0,X2))[ennf transformation 105]
105. ! [X0,X1,X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))[flattening 104]
104. ! [X0] : ! [X1] : ! [X2] : ((cont(X0,X1) & cont(X1,X2)) => cont(X0,X2))[rectify 15]
15. ! [X29] : ! [X30] : ! [X31] : ((cont(X29,X30) & cont(X30,X31)) => cont(X29,X31))[input]
------------------------------
Version: Vampire 0.6 (revision 903)
Termination reason: Refutation

Active clauses: 354
Passive clauses: 1107
Generated clauses: 2438
Final active clauses: 354
Final passive clauses: 704
Input formulas: 43
Initial clauses: 84

Unused predicate definitions: 2
Selected by SInE selection: 43
SInE iterations: 1

Duplicate literals: 104
Fw subsumption resolutions: 110

Simple tautologies: 44
Equational tautologies: 3
Forward subsumptions: 1117

Binary resolution: 1157
Factoring: 24
Forward superposition: 922
Backward superposition: 44

Unique components: 2

Memory used: 511KB
Time elapsed: 0.314 s
------------------------------
% Success in time 1.157 s
============================= C:\Reasoning\vampire_win ================================
execution finished: Sun Jun 24 23:39:21 2018
total CPU time used: 0
The command was "C:\Reasoning\vampire_win --mode casc --proof tptp -t 600 --input_file C:\Users\torsten\Documents\github\colore\ontologies\multidim_mereotopology_codi\theorems\conversions\codi_int_theorems_goal21.all.tptp"
Input read from 
============================ end of footer ===========================
