% Refutation found. Thanks to Tanya!
% SZS status Theorem for 24apr2018_eppsubset_T12
% SZS output start Proof for 24apr2018_eppsubset_T12
6. ! [X15] : (~zex(X15) <=> cont(X15,X15)) [input sos60]
9. ! [X21] : ! [X22] : (zex(X21) => (~cont(X21,X22) & ~cont(X22,X21))) [input sos90]
14. ! [X31] : ! [X32] : (eqdim(X31,X32) <=> (leq(X32,X31) & leq(X31,X32))) [input sos140]
15. ! [X33] : leq(X33,X33) [input sos150]
25. ! [X53] : ! [X54] : (p(X53,X54) <=> (eqdim(X53,X54) & cont(X53,X54))) [input sos250]
31. ! [X72] : ! [X73] : ! [X74] : (path(X72,X73,X74) => (cont(X74,X72) & cont(X73,X72))) [input sos320]
36. ! [X100] : ! [X101] : (connects(X100,X101) <=> ? [X102] : ? [X103] : ? [X104] : (path(X101,X103,X104) & path(X100,X102,X103))) [input sos370]
42. ! [X135] : ! [X136] : ! [X137] : ! [X138] : ! [X139] : ((p(X136,X135) & path(X136,X137,X139) & path(X135,X137,X138)) => ~? [X140] : (connects(X140,X136) & p(X140,X135))) [input sos430]
48. ! [X176] : ! [X28] : ! [X27] : (path(X176,X28,X27) => X27 != X28) [input t120]
49. ~! [X176] : ! [X28] : ! [X27] : (path(X176,X28,X27) => X27 != X28) [negated conjecture 48]
58. ! [X0] : (~zex(X0) <=> cont(X0,X0)) [rectify 6]
63. ! [X0] : ! [X1] : (zex(X0) => (~cont(X0,X1) & ~cont(X1,X0))) [rectify 9]
64. ! [X0,X1] : (zex(X0) => (~cont(X0,X1) & ~cont(X1,X0))) [flattening 63]
71. ! [X0] : ! [X1] : (eqdim(X0,X1) <=> (leq(X1,X0) & leq(X0,X1))) [rectify 14]
72. ! [X0,X1] : (eqdim(X0,X1) <=> (leq(X1,X0) & leq(X0,X1))) [flattening 71]
73. ! [X0] : leq(X0,X0) [rectify 15]
91. ! [X0] : ! [X1] : (p(X0,X1) <=> (eqdim(X0,X1) & cont(X0,X1))) [rectify 25]
92. ! [X0,X1] : (p(X0,X1) <=> (eqdim(X0,X1) & cont(X0,X1))) [flattening 91]
103. ! [X0] : ! [X1] : ! [X2] : (path(X0,X1,X2) => (cont(X2,X0) & cont(X1,X0))) [rectify 31]
104. ! [X0,X1,X2] : (path(X0,X1,X2) => (cont(X2,X0) & cont(X1,X0))) [flattening 103]
113. ! [X0] : ! [X1] : (connects(X0,X1) <=> ? [X2] : ? [X3] : ? [X4] : (path(X1,X3,X4) & path(X0,X2,X3))) [rectify 36]
114. ! [X0,X1] : (connects(X0,X1) <=> ? [X2,X3,X4] : (path(X1,X3,X4) & path(X0,X2,X3))) [flattening 113]
125. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ((p(X1,X0) & path(X1,X2,X4) & path(X0,X2,X3)) => ~? [X5] : (connects(X5,X1) & p(X5,X0))) [rectify 42]
126. ! [X0,X1,X2,X3,X4] : ((p(X1,X0) & path(X1,X2,X4) & path(X0,X2,X3)) => ~? [X5] : (connects(X5,X1) & p(X5,X0))) [flattening 125]
137. ~! [X0] : ! [X1] : ! [X2] : (path(X0,X1,X2) => X1 != X2) [rectify 49]
138. ~! [X0,X1,X2] : (path(X0,X1,X2) => X1 != X2) [flattening 137]
153. ! [X0,X1] : ((~cont(X0,X1) & ~cont(X1,X0)) | ~zex(X0)) [ennf transformation 64]
170. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X1,X0)) | ~path(X0,X1,X2)) [ennf transformation 104]
188. ! [X0,X1,X2,X3,X4] : (! [X5] : (~connects(X5,X1) | ~p(X5,X0)) | (~p(X1,X0) | ~path(X1,X2,X4) | ~path(X0,X2,X3))) [ennf transformation 126]
189. ! [X0,X1,X2,X3,X4] : (! [X5] : (~connects(X5,X1) | ~p(X5,X0)) | ~p(X1,X0) | ~path(X1,X2,X4) | ~path(X0,X2,X3)) [flattening 188]
198. ? [X0,X1,X2] : (X1 = X2 & path(X0,X1,X2)) [ennf transformation 138]
202. ! [X0] : ((~zex(X0) | ~cont(X0,X0)) & (cont(X0,X0) | zex(X0))) [nnf transformation 58]
205. ! [X0,X1] : ((eqdim(X0,X1) | (~leq(X1,X0) | ~leq(X0,X1))) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1))) [nnf transformation 72]
206. ! [X0,X1] : ((eqdim(X0,X1) | ~leq(X1,X0) | ~leq(X0,X1)) & ((leq(X1,X0) & leq(X0,X1)) | ~eqdim(X0,X1))) [flattening 205]
209. ! [X0,X1] : ((p(X0,X1) | (~eqdim(X0,X1) | ~cont(X0,X1))) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1))) [nnf transformation 92]
210. ! [X0,X1] : ((p(X0,X1) | ~eqdim(X0,X1) | ~cont(X0,X1)) & ((eqdim(X0,X1) & cont(X0,X1)) | ~p(X0,X1))) [flattening 209]
213. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & (? [X2,X3,X4] : (path(X1,X3,X4) & path(X0,X2,X3)) | ~connects(X0,X1))) [nnf transformation 114]
214. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & (? [X5,X6,X7] : (path(X1,X6,X7) & path(X0,X5,X6)) | ~connects(X0,X1))) [rectify 213]
215. ! [X1,X0] : (? [X5,X6,X7] : (path(X1,X6,X7) & path(X0,X5,X6)) => (path(X1,sK6(X0,X1),sK7(X0,X1)) & path(X0,sK5(X0,X1),sK6(X0,X1)))) [choice axiom]
216. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & ((path(X1,sK6(X0,X1),sK7(X0,X1)) & path(X0,sK5(X0,X1),sK6(X0,X1))) | ~connects(X0,X1))) [skolemisation 214,215]
231. ? [X0,X1,X2] : (X1 = X2 & path(X0,X1,X2)) => (sK25 = sK26 & path(sK24,sK25,sK26)) [choice axiom]
232. sK25 = sK26 & path(sK24,sK25,sK26) [skolemisation 198,231]
241. cont(X0,X0) | zex(X0) [cnf transformation 202]
245. ~cont(X1,X0) | ~zex(X0) [cnf transformation 153]
255. ~leq(X1,X0) | eqdim(X0,X1) | ~leq(X0,X1) [cnf transformation 206]
256. leq(X0,X0) [cnf transformation 73]
266. ~eqdim(X0,X1) | p(X0,X1) | ~cont(X0,X1) [cnf transformation 210]
272. ~path(X0,X1,X2) | cont(X1,X0) [cnf transformation 170]
282. ~path(X1,X3,X4) | connects(X0,X1) | ~path(X0,X2,X3) [cnf transformation 216]
303. ~path(X1,X2,X4) | ~p(X5,X0) | ~p(X1,X0) | ~connects(X5,X1) | ~path(X0,X2,X3) [cnf transformation 189]
313. path(sK24,sK25,sK26) [cnf transformation 232]
314. sK25 = sK26 [cnf transformation 232]
327. path(sK24,sK25,sK25) [forward demodulation 313,314]
335. cont(sK25,sK24) [resolution 272,327]
338. ~zex(sK24) [resolution 335,245]
385. eqdim(X0,X0) | ~leq(X0,X0) [resolution 255,256]
388. eqdim(X0,X0) [subsumption resolution 385,256]
415. ~cont(X0,X0) | p(X0,X0) [resolution 266,388]
508. ~path(X0,X1,sK25) | connects(X0,sK24) [resolution 282,327]
667. ~path(X1,sK25,X2) | ~p(sK24,X1) | ~connects(X0,sK24) | ~p(X0,X1) [resolution 303,327]
787. p(X0,X0) | zex(X0) [resolution 415,241]
788. connects(sK24,sK24) [resolution 508,327]
1272. ~p(sK24,sK24) | ~connects(X0,sK24) | ~p(X0,sK24) [resolution 667,327]
1274. 68 <=> ! [X0] : (~connects(X0,sK24) | ~p(X0,sK24)) [avatar definition]
1275. ~connects(X0,sK24) | ~p(X0,sK24) <- {68} [avatar component clause 1274]
1277. 70 <=> p(sK24,sK24) [avatar definition]
1278. p(sK24,sK24) <- {70} [avatar component clause 1277]
1280. 71 <=> ~p(sK24,sK24) [avatar definition]
1281. ~p(sK24,sK24) <- {71} [avatar component clause 1280]
1282. 68 | ~71 [avatar split clause 1272,1280,1274]
1283. zex(sK24) <- {71} [resolution 1281,787]
1284. $false <- {71} [subsumption resolution 1283,338]
1285. 71 [avatar contradiction clause 1284]
1340. ~p(sK24,sK24) <- {68} [resolution 1275,788]
1342. $false <- {68, 70} [subsumption resolution 1340,1278]
1343. ~68 | ~70 [avatar contradiction clause 1342]
1345. $false [avatar sat refutation 1282,1285,1343]
% SZS output end Proof for 24apr2018_eppsubset_T12
% ------------------------------
% Version: Vampire 4.2.2 (commit c6486af3 on 2018-01-23 13:53:44 +0000)
% Termination reason: Refutation

% Memory used [KB]: 5756
% Time elapsed: 0.217 s
% ------------------------------
----  Runtime statistics ----
clauses created: 1023
clauses deleted: 190
ssat_new_components: 36
ssat_non_splittable_sat_clauses: 1
ssat_reused_components: 8
ssat_sat_clauses: 27
total_frozen: 90
total_unfrozen: 12
unworthy child removed: 78
-----------------------------
% ------------------------------
