% Refutation found. Thanks to Tanya!
% SZS status Theorem for 24apr2018_full_T07
% SZS output start Proof for 24apr2018_full_T07
2. ! [X2] : ! [X3] : ! [X4] : (arc(X2,X3,X4) => (node(X4) & node(X3) & edge(X2))) [input sos20]
4. ! [X6] : (edge(X6) => ? [X7] : ? [X8] : arc(X6,X7,X8)) [input sos40]
22. ! [X49] : ! [X50] : (pp(X49,X50) <=> (X49 != X50 & p(X49,X50))) [input sos220]
43. ! [X95] : ! [X96] : ! [X97] : (path(X95,X96,X97) => (node(X97) & node(X96) & curve(X95))) [input sos430]
45. ! [X99] : ! [X100] : ! [X101] : (path(X99,X100,X101) => X100 != X101) [input sos450]
46. ! [X102] : ! [X103] : ! [X104] : (arc(X102,X103,X104) => path(X102,X103,X104)) [input sos460]
47. ! [X105] : ! [X106] : ! [X107] : ! [X108] : ((p(X108,X105) & arc(X105,X106,X107)) => X105 = X108) [input sos470]
48. ! [X109] : ! [X110] : ! [X111] : (path(X109,X110,X111) => (cont(X111,X109) & cont(X110,X109))) [input sos480]
49. ! [X112] : ! [X113] : ! [X114] : ! [X115] : ((cont(X112,X113) & path(X113,X114,X115) & node(X112)) => ? [X116] : ? [X117] : (p(X116,X113) & (path(X116,X117,X112) | path(X116,X112,X117)))) [input sos490]
50. ! [X118] : ! [X119] : ! [X120] : ! [X121] : ! [X122] : ((path(X118,X121,X122) & path(X118,X119,X120)) => (X120 = X122 & X119 = X121)) [input sos500]
53. ! [X137] : ! [X138] : (connects(X137,X138) <=> ? [X139] : ? [X140] : ? [X141] : (path(X138,X140,X141) & path(X137,X139,X140))) [input sos530]
58. ! [X164] : ! [X165] : ! [X166] : ! [X167] : ! [X168] : ! [X169] : ((pp(X164,X167) & X166 != X169 & X165 != X168 & path(X164,X165,X166) & path(X167,X168,X169)) => ? [X170] : ? [X171] : (pp(X171,X167) & pp(X170,X167) & connects(X164,X171) & connects(X170,X164))) [input sos580]
63. ! [X201] : ! [X202] : ! [X203] : ((~arc(X201,X202,X203) & path(X201,X202,X203)) => ? [X204] : ? [X205] : ? [X206] : (pp(X206,X201) & pp(X204,X201) & connects(X204,X206) & arc(X204,X202,X205))) [input sos630]
65. ! [X17] : ! [X213] : ! [X214] : (path(X17,X213,X214) => ? [X215] : ? [X18] : ? [X216] : ? [X217] : (p(X18,X17) & path(X18,X217,X214) & p(X215,X17) & path(X215,X213,X216))) [input t70]
66. ~! [X17] : ! [X213] : ! [X214] : (path(X17,X213,X214) => ? [X215] : ? [X18] : ? [X216] : ? [X217] : (p(X18,X17) & path(X18,X217,X214) & p(X215,X17) & path(X215,X213,X216))) [negated conjecture 65]
68. ! [X0] : ! [X1] : ! [X2] : (arc(X0,X1,X2) => (node(X2) & node(X1) & edge(X0))) [rectify 2]
69. ! [X0,X1,X2] : (arc(X0,X1,X2) => (node(X2) & node(X1) & edge(X0))) [flattening 68]
71. ! [X0] : (edge(X0) => ? [X1] : ? [X2] : arc(X0,X1,X2)) [rectify 4]
72. ! [X0] : (edge(X0) => ? [X1,X2] : arc(X0,X1,X2)) [flattening 71]
101. ! [X0] : ! [X1] : (pp(X0,X1) <=> (X0 != X1 & p(X0,X1))) [rectify 22]
102. ! [X0,X1] : (pp(X0,X1) <=> (X0 != X1 & p(X0,X1))) [flattening 101]
138. ! [X0] : ! [X1] : ! [X2] : (path(X0,X1,X2) => (node(X2) & node(X1) & curve(X0))) [rectify 43]
139. ! [X0,X1,X2] : (path(X0,X1,X2) => (node(X2) & node(X1) & curve(X0))) [flattening 138]
141. ! [X0] : ! [X1] : ! [X2] : (path(X0,X1,X2) => X1 != X2) [rectify 45]
142. ! [X0,X1,X2] : (path(X0,X1,X2) => X1 != X2) [flattening 141]
143. ! [X0] : ! [X1] : ! [X2] : (arc(X0,X1,X2) => path(X0,X1,X2)) [rectify 46]
144. ! [X0,X1,X2] : (arc(X0,X1,X2) => path(X0,X1,X2)) [flattening 143]
145. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((p(X3,X0) & arc(X0,X1,X2)) => X0 = X3) [rectify 47]
146. ! [X0,X1,X2,X3] : ((p(X3,X0) & arc(X0,X1,X2)) => X0 = X3) [flattening 145]
147. ! [X0] : ! [X1] : ! [X2] : (path(X0,X1,X2) => (cont(X2,X0) & cont(X1,X0))) [rectify 48]
148. ! [X0,X1,X2] : (path(X0,X1,X2) => (cont(X2,X0) & cont(X1,X0))) [flattening 147]
149. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((cont(X0,X1) & path(X1,X2,X3) & node(X0)) => ? [X4] : ? [X5] : (p(X4,X1) & (path(X4,X5,X0) | path(X4,X0,X5)))) [rectify 49]
150. ! [X0,X1,X2,X3] : ((cont(X0,X1) & path(X1,X2,X3) & node(X0)) => ? [X4,X5] : (p(X4,X1) & (path(X4,X5,X0) | path(X4,X0,X5)))) [flattening 149]
151. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ((path(X0,X3,X4) & path(X0,X1,X2)) => (X2 = X4 & X1 = X3)) [rectify 50]
152. ! [X0,X1,X2,X3,X4] : ((path(X0,X3,X4) & path(X0,X1,X2)) => (X2 = X4 & X1 = X3)) [flattening 151]
157. ! [X0] : ! [X1] : (connects(X0,X1) <=> ? [X2] : ? [X3] : ? [X4] : (path(X1,X3,X4) & path(X0,X2,X3))) [rectify 53]
158. ! [X0,X1] : (connects(X0,X1) <=> ? [X2,X3,X4] : (path(X1,X3,X4) & path(X0,X2,X3))) [flattening 157]
167. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ((pp(X0,X3) & X2 != X5 & X1 != X4 & path(X0,X1,X2) & path(X3,X4,X5)) => ? [X6] : ? [X7] : (pp(X7,X3) & pp(X6,X3) & connects(X0,X7) & connects(X6,X0))) [rectify 58]
168. ! [X0,X1,X2,X3,X4,X5] : ((pp(X0,X3) & X2 != X5 & X1 != X4 & path(X0,X1,X2) & path(X3,X4,X5)) => ? [X6,X7] : (pp(X7,X3) & pp(X6,X3) & connects(X0,X7) & connects(X6,X0))) [flattening 167]
177. ! [X0] : ! [X1] : ! [X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3] : ? [X4] : ? [X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4))) [rectify 63]
178. ! [X0,X1,X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4))) [flattening 177]
181. ~! [X0] : ! [X1] : ! [X2] : (path(X0,X1,X2) => ? [X3] : ? [X4] : ? [X5] : ? [X6] : (p(X4,X0) & path(X4,X6,X2) & p(X3,X0) & path(X3,X1,X5))) [rectify 66]
182. ~! [X0,X1,X2] : (path(X0,X1,X2) => ? [X3,X4,X5,X6] : (p(X4,X0) & path(X4,X6,X2) & p(X3,X0) & path(X3,X1,X5))) [flattening 181]
186. ! [X0,X1,X2] : ((node(X2) & node(X1) & edge(X0)) | ~arc(X0,X1,X2)) [ennf transformation 69]
188. ! [X0] : (? [X1,X2] : arc(X0,X1,X2) | ~edge(X0)) [ennf transformation 72]
220. ! [X0,X1,X2] : ((node(X2) & node(X1) & curve(X0)) | ~path(X0,X1,X2)) [ennf transformation 139]
222. ! [X0,X1,X2] : (X1 != X2 | ~path(X0,X1,X2)) [ennf transformation 142]
223. ! [X0,X1,X2] : (path(X0,X1,X2) | ~arc(X0,X1,X2)) [ennf transformation 144]
224. ! [X0,X1,X2,X3] : (X0 = X3 | (~p(X3,X0) | ~arc(X0,X1,X2))) [ennf transformation 146]
225. ! [X0,X1,X2,X3] : (X0 = X3 | ~p(X3,X0) | ~arc(X0,X1,X2)) [flattening 224]
226. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X1,X0)) | ~path(X0,X1,X2)) [ennf transformation 148]
227. ! [X0,X1,X2,X3] : (? [X4,X5] : (p(X4,X1) & (path(X4,X5,X0) | path(X4,X0,X5))) | (~cont(X0,X1) | ~path(X1,X2,X3) | ~node(X0))) [ennf transformation 150]
228. ! [X0,X1,X2,X3] : (? [X4,X5] : (p(X4,X1) & (path(X4,X5,X0) | path(X4,X0,X5))) | ~cont(X0,X1) | ~path(X1,X2,X3) | ~node(X0)) [flattening 227]
229. ! [X0,X1,X2,X3,X4] : ((X2 = X4 & X1 = X3) | (~path(X0,X3,X4) | ~path(X0,X1,X2))) [ennf transformation 152]
230. ! [X0,X1,X2,X3,X4] : ((X2 = X4 & X1 = X3) | ~path(X0,X3,X4) | ~path(X0,X1,X2)) [flattening 229]
242. ! [X0,X1,X2,X3,X4,X5] : (? [X6,X7] : (pp(X7,X3) & pp(X6,X3) & connects(X0,X7) & connects(X6,X0)) | (~pp(X0,X3) | X2 = X5 | X1 = X4 | ~path(X0,X1,X2) | ~path(X3,X4,X5))) [ennf transformation 168]
243. ! [X0,X1,X2,X3,X4,X5] : (? [X6,X7] : (pp(X7,X3) & pp(X6,X3) & connects(X0,X7) & connects(X6,X0)) | ~pp(X0,X3) | X2 = X5 | X1 = X4 | ~path(X0,X1,X2) | ~path(X3,X4,X5)) [flattening 242]
252. ! [X0,X1,X2] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) | (arc(X0,X1,X2) | ~path(X0,X1,X2))) [ennf transformation 178]
253. ! [X0,X1,X2] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) | arc(X0,X1,X2) | ~path(X0,X1,X2)) [flattening 252]
254. ? [X0,X1,X2] : (! [X3,X4,X5,X6] : (~p(X4,X0) | ~path(X4,X6,X2) | ~p(X3,X0) | ~path(X3,X1,X5)) & path(X0,X1,X2)) [ennf transformation 182]
255. ! [X0] : (? [X1,X2] : arc(X0,X1,X2) => arc(X0,sK0(X0),sK1(X0))) [choice axiom]
256. ! [X0] : (arc(X0,sK0(X0),sK1(X0)) | ~edge(X0)) [skolemisation 188,255]
261. ! [X0,X1] : ((pp(X0,X1) | (X0 = X1 | ~p(X0,X1))) & ((X0 != X1 & p(X0,X1)) | ~pp(X0,X1))) [nnf transformation 102]
262. ! [X0,X1] : ((pp(X0,X1) | X0 = X1 | ~p(X0,X1)) & ((X0 != X1 & p(X0,X1)) | ~pp(X0,X1))) [flattening 261]
285. ! [X1,X0] : (? [X4,X5] : (p(X4,X1) & (path(X4,X5,X0) | path(X4,X0,X5))) => (p(sK7(X0,X1),X1) & (path(sK7(X0,X1),sK8(X0,X1),X0) | path(sK7(X0,X1),X0,sK8(X0,X1))))) [choice axiom]
286. ! [X0,X1,X2,X3] : ((p(sK7(X0,X1),X1) & (path(sK7(X0,X1),sK8(X0,X1),X0) | path(sK7(X0,X1),X0,sK8(X0,X1)))) | ~cont(X0,X1) | ~path(X1,X2,X3) | ~node(X0)) [skolemisation 228,285]
287. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & (? [X2,X3,X4] : (path(X1,X3,X4) & path(X0,X2,X3)) | ~connects(X0,X1))) [nnf transformation 158]
288. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & (? [X5,X6,X7] : (path(X1,X6,X7) & path(X0,X5,X6)) | ~connects(X0,X1))) [rectify 287]
289. ! [X1,X0] : (? [X5,X6,X7] : (path(X1,X6,X7) & path(X0,X5,X6)) => (path(X1,sK10(X0,X1),sK11(X0,X1)) & path(X0,sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
290. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & ((path(X1,sK10(X0,X1),sK11(X0,X1)) & path(X0,sK9(X0,X1),sK10(X0,X1))) | ~connects(X0,X1))) [skolemisation 288,289]
299. ! [X3,X0] : (? [X6,X7] : (pp(X7,X3) & pp(X6,X3) & connects(X0,X7) & connects(X6,X0)) => (pp(sK18(X0,X3),X3) & pp(sK17(X0,X3),X3) & connects(X0,sK18(X0,X3)) & connects(sK17(X0,X3),X0))) [choice axiom]
300. ! [X0,X1,X2,X3,X4,X5] : ((pp(sK18(X0,X3),X3) & pp(sK17(X0,X3),X3) & connects(X0,sK18(X0,X3)) & connects(sK17(X0,X3),X0)) | ~pp(X0,X3) | X2 = X5 | X1 = X4 | ~path(X0,X1,X2) | ~path(X3,X4,X5)) [skolemisation 243,299]
303. ! [X1,X0] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) => (pp(sK24(X0,X1),X0) & pp(sK22(X0,X1),X0) & connects(sK22(X0,X1),sK24(X0,X1)) & arc(sK22(X0,X1),X1,sK23(X0,X1)))) [choice axiom]
304. ! [X0,X1,X2] : ((pp(sK24(X0,X1),X0) & pp(sK22(X0,X1),X0) & connects(sK22(X0,X1),sK24(X0,X1)) & arc(sK22(X0,X1),X1,sK23(X0,X1))) | arc(X0,X1,X2) | ~path(X0,X1,X2)) [skolemisation 253,303]
305. ? [X0,X1,X2] : (! [X3,X4,X5,X6] : (~p(X4,X0) | ~path(X4,X6,X2) | ~p(X3,X0) | ~path(X3,X1,X5)) & path(X0,X1,X2)) => (! [X6,X5,X4,X3] : (~p(X4,sK25) | ~path(X4,X6,sK27) | ~p(X3,sK25) | ~path(X3,sK26,X5)) & path(sK25,sK26,sK27)) [choice axiom]
306. ! [X3,X4,X5,X6] : (~p(X4,sK25) | ~path(X4,X6,sK27) | ~p(X3,sK25) | ~path(X3,sK26,X5)) & path(sK25,sK26,sK27) [skolemisation 254,305]
308. ~arc(X0,X1,X2) | edge(X0) [cnf transformation 186]
312. arc(X0,sK0(X0),sK1(X0)) | ~edge(X0) [cnf transformation 256]
331. ~pp(X0,X1) | p(X0,X1) [cnf transformation 262]
333. ~p(X0,X1) | X0 = X1 | pp(X0,X1) [cnf transformation 262]
364. ~path(X0,X1,X2) | node(X2) [cnf transformation 220]
366. X1 != X2 | ~path(X0,X1,X2) [cnf transformation 222]
367. ~arc(X0,X1,X2) | path(X0,X1,X2) [cnf transformation 223]
368. ~arc(X0,X1,X2) | ~p(X3,X0) | X0 = X3 [cnf transformation 225]
370. ~path(X0,X1,X2) | cont(X2,X0) [cnf transformation 226]
371. ~path(X1,X2,X3) | path(sK7(X0,X1),X0,sK8(X0,X1)) | ~cont(X0,X1) | path(sK7(X0,X1),sK8(X0,X1),X0) | ~node(X0) [cnf transformation 286]
372. ~path(X1,X2,X3) | ~cont(X0,X1) | p(sK7(X0,X1),X1) | ~node(X0) [cnf transformation 286]
373. ~path(X0,X3,X4) | X1 = X3 | ~path(X0,X1,X2) [cnf transformation 230]
377. path(X0,sK9(X0,X1),sK10(X0,X1)) | ~connects(X0,X1) [cnf transformation 290]
378. path(X1,sK10(X0,X1),sK11(X0,X1)) | ~connects(X0,X1) [cnf transformation 290]
394. ~path(X3,X4,X5) | ~pp(X0,X3) | X2 = X5 | X1 = X4 | ~path(X0,X1,X2) | connects(sK17(X0,X3),X0) [cnf transformation 300]
396. ~path(X3,X4,X5) | ~pp(X0,X3) | X2 = X5 | X1 = X4 | ~path(X0,X1,X2) | pp(sK17(X0,X3),X3) [cnf transformation 300]
404. ~path(X0,X1,X2) | arc(X0,X1,X2) | arc(sK22(X0,X1),X1,sK23(X0,X1)) [cnf transformation 304]
406. ~path(X0,X1,X2) | arc(X0,X1,X2) | pp(sK22(X0,X1),X0) [cnf transformation 304]
408. path(sK25,sK26,sK27) [cnf transformation 306]
409. ~p(X4,sK25) | ~path(X4,X6,sK27) | ~p(X3,sK25) | ~path(X3,sK26,X5) [cnf transformation 306]
411. ~path(X0,X2,X2) [equality resolution 366]
413. 0 <=> ! [X3,X5] : (~p(X3,sK25) | ~path(X3,sK26,X5)) [avatar definition]
414. ~path(X3,sK26,X5) | ~p(X3,sK25) <- {0} [avatar component clause 413]
416. 2 <=> ! [X4,X6] : (~p(X4,sK25) | ~path(X4,X6,sK27)) [avatar definition]
417. ~path(X4,X6,sK27) | ~p(X4,sK25) <- {2} [avatar component clause 416]
418. 0 | 2 [avatar split clause 409,416,413]
431. node(sK27) [resolution 364,408]
451. cont(sK27,sK25) [resolution 370,408]
609. ~p(X0,X1) | X0 = X1 | ~edge(X1) [resolution 368,312]
689. ~path(X11,X12,X13) | sK10(X10,X11) = X12 | ~connects(X10,X11) [resolution 378,373]
709. arc(sK25,sK26,sK27) | pp(sK22(sK25,sK26),sK25) [resolution 406,408]
716. 36 <=> pp(sK22(sK25,sK26),sK25) [avatar definition]
717. pp(sK22(sK25,sK26),sK25) <- {36} [avatar component clause 716]
722. 38 <=> arc(sK25,sK26,sK27) [avatar definition]
723. arc(sK25,sK26,sK27) <- {38} [avatar component clause 722]
724. 36 | 38 [avatar split clause 709,722,716]
743. p(sK7(X0,sK25),sK25) | ~cont(X0,sK25) | ~node(X0) [resolution 372,408]
765. arc(sK25,sK26,sK27) | arc(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) [resolution 404,408]
839. ~path(X0,X2,X1) | sK27 = X1 | sK26 = X2 | ~pp(X0,sK25) | connects(sK17(X0,sK25),X0) [resolution 394,408]
845. ~path(X0,X2,X1) | sK27 = X1 | sK26 = X2 | ~pp(X0,sK25) | pp(sK17(X0,sK25),sK25) [resolution 396,408]
852. path(sK7(X0,sK25),sK8(X0,sK25),X0) | ~cont(X0,sK25) | path(sK7(X0,sK25),X0,sK8(X0,sK25)) | ~node(X0) [resolution 371,408]
1867. 84 <=> path(sK7(sK27,sK25),sK27,sK8(sK27,sK25)) [avatar definition]
1868. path(sK7(sK27,sK25),sK27,sK8(sK27,sK25)) <- {84} [avatar component clause 1867]
1879. 88 <=> pp(sK7(sK27,sK25),sK25) [avatar definition]
1880. pp(sK7(sK27,sK25),sK25) <- {88} [avatar component clause 1879]
1897. 92 <=> p(sK7(sK27,sK25),sK25) [avatar definition]
1898. p(sK7(sK27,sK25),sK25) <- {92} [avatar component clause 1897]
1900. 93 <=> ~p(sK7(sK27,sK25),sK25) [avatar definition]
1901. ~p(sK7(sK27,sK25),sK25) <- {93} [avatar component clause 1900]
1928. ~cont(sK27,sK25) | ~node(sK27) <- {93} [resolution 1901,743]
1929. ~node(sK27) <- {93} [subsumption resolution 1928,451]
1930. $false <- {93} [subsumption resolution 1929,431]
1931. 93 [avatar contradiction clause 1930]
1959. 94 <=> sK7(sK27,sK25) = sK25 [avatar definition]
1960. sK7(sK27,sK25) = sK25 <- {94} [avatar component clause 1959]
1963. 96 <=> edge(sK25) [avatar definition]
1966. 97 <=> ~edge(sK25) [avatar definition]
1991. 104 <=> arc(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) [avatar definition]
1992. arc(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) <- {104} [avatar component clause 1991]
1993. 104 | 38 [avatar split clause 765,722,1991]
2008. 111 <=> ~p(sK25,sK25) [avatar definition]
2009. ~p(sK25,sK25) <- {111} [avatar component clause 2008]
2016. ~p(sK25,sK25) <- {0} [resolution 414,408]
2019. edge(sK25) <- {38} [resolution 723,308]
2033. 96 | ~38 [avatar split clause 2019,722,1963]
2036. p(sK25,sK25) <- {92, 94} [backward demodulation 1960,1898]
2042. $false <- {92, 94, 111} [subsumption resolution 2036,2009]
2043. ~92 | ~94 | 111 [avatar contradiction clause 2042]
2296. sK7(sK27,sK25) = sK25 | ~edge(sK25) <- {92} [resolution 1898,609]
2427. ~97 | 94 | ~92 [avatar split clause 2296,1897,1959,1966]
2434. ~111 | ~0 [avatar split clause 2016,413,2008]
2450. p(sK22(sK25,sK26),sK25) <- {36} [resolution 717,331]
2457. path(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) <- {104} [resolution 1992,367]
2472. sK7(sK27,sK25) = sK25 | pp(sK7(sK27,sK25),sK25) <- {92} [resolution 1898,333]
2560. sK8(sK27,sK25) = sK27 | sK26 = sK27 | ~pp(sK7(sK27,sK25),sK25) | connects(sK17(sK7(sK27,sK25),sK25),sK7(sK27,sK25)) <- {84} [resolution 1868,839]
2566. sK8(sK27,sK25) = sK27 | sK26 = sK27 | ~pp(sK7(sK27,sK25),sK25) | pp(sK17(sK7(sK27,sK25),sK25),sK25) <- {84} [resolution 1868,845]
2606. sK8(sK27,sK25) = sK27 | sK26 = sK27 | connects(sK17(sK7(sK27,sK25),sK25),sK7(sK27,sK25)) <- {84, 88} [subsumption resolution 2560,1880]
2611. 130 <=> connects(sK17(sK7(sK27,sK25),sK25),sK7(sK27,sK25)) [avatar definition]
2612. connects(sK17(sK7(sK27,sK25),sK25),sK7(sK27,sK25)) <- {130} [avatar component clause 2611]
2617. 132 <=> sK26 = sK27 [avatar definition]
2618. sK26 = sK27 <- {132} [avatar component clause 2617]
2623. 134 <=> sK8(sK27,sK25) = sK27 [avatar definition]
2624. sK8(sK27,sK25) = sK27 <- {134} [avatar component clause 2623]
2625. 130 | 132 | 134 | ~84 | ~88 [avatar split clause 2606,1879,1867,2623,2617,2611]
2634. sK8(sK27,sK25) = sK27 | sK26 = sK27 | pp(sK17(sK7(sK27,sK25),sK25),sK25) <- {84, 88} [subsumption resolution 2566,1880]
2639. 138 <=> pp(sK17(sK7(sK27,sK25),sK25),sK25) [avatar definition]
2640. pp(sK17(sK7(sK27,sK25),sK25),sK25) <- {138} [avatar component clause 2639]
2641. 138 | 132 | 134 | ~84 | ~88 [avatar split clause 2634,1879,1867,2623,2617,2639]
2657. ~p(sK22(sK25,sK26),sK25) <- {0, 104} [resolution 2457,414]
2709. $false <- {0, 36, 104} [subsumption resolution 2657,2450]
2710. ~0 | ~36 | ~104 [avatar contradiction clause 2709]
2717. path(sK25,sK26,sK26) <- {132} [backward demodulation 2618,408]
2862. $false <- {132} [subsumption resolution 2717,411]
2863. ~132 [avatar contradiction clause 2862]
2944. path(sK7(sK27,sK25),sK27,sK27) <- {84, 134} [backward demodulation 2624,1868]
2945. $false <- {84, 134} [subsumption resolution 2944,411]
2946. ~84 | ~134 [avatar contradiction clause 2945]
3001. ~p(sK25,sK25) <- {2} [resolution 417,408]
3210. ~connects(X53,sK7(sK27,sK25)) | sK10(X53,sK7(sK27,sK25)) = sK27 <- {84} [resolution 1868,689]
5255. p(sK17(sK7(sK27,sK25),sK25),sK25) <- {138} [resolution 2640,331]
25715. 88 | 94 | ~92 [avatar split clause 2472,1897,1959,1879]
25729. ~111 | ~2 [avatar split clause 3001,416,2008]
54208. sK10(sK17(sK7(sK27,sK25),sK25),sK7(sK27,sK25)) = sK27 <- {84, 130} [resolution 3210,2612]
137556. ~p(sK7(sK27,sK25),sK25) | ~cont(sK27,sK25) | path(sK7(sK27,sK25),sK27,sK8(sK27,sK25)) | ~node(sK27) <- {2} [resolution 417,852]
137559. ~cont(sK27,sK25) | path(sK7(sK27,sK25),sK27,sK8(sK27,sK25)) | ~node(sK27) <- {2, 92} [subsumption resolution 137556,1898]
137560. path(sK7(sK27,sK25),sK27,sK8(sK27,sK25)) | ~node(sK27) <- {2, 92} [subsumption resolution 137559,451]
137561. path(sK7(sK27,sK25),sK27,sK8(sK27,sK25)) <- {2, 92} [subsumption resolution 137560,431]
195159. 84 | ~2 | ~92 [avatar split clause 137561,1897,416,1867]
615820. 6640 <=> sK10(sK17(sK7(sK27,sK25),sK25),sK7(sK27,sK25)) = sK27 [avatar definition]
615821. sK10(sK17(sK7(sK27,sK25),sK25),sK7(sK27,sK25)) = sK27 <- {6640} [avatar component clause 615820]
617949. 6640 | ~84 | ~130 [avatar split clause 54208,2611,1867,615820]
619952. path(sK17(sK7(sK27,sK25),sK25),sK9(sK17(sK7(sK27,sK25),sK25),sK7(sK27,sK25)),sK27) | ~connects(sK17(sK7(sK27,sK25),sK25),sK7(sK27,sK25)) <- {6640} [superposition 377,615821]
620163. path(sK17(sK7(sK27,sK25),sK25),sK9(sK17(sK7(sK27,sK25),sK25),sK7(sK27,sK25)),sK27) <- {130, 6640} [subsumption resolution 619952,2612]
746852. ~p(sK17(sK7(sK27,sK25),sK25),sK25) <- {2, 130, 6640} [resolution 620163,417]
746913. $false <- {2, 130, 138, 6640} [subsumption resolution 746852,5255]
746914. ~2 | ~130 | ~138 | ~6640 [avatar contradiction clause 746913]
746927. $false [avatar sat refutation 418,724,1931,1993,2033,2043,2427,2434,2625,2641,2710,2863,2946,25715,25729,195159,617949,746914]
% SZS output end Proof for 24apr2018_full_T07
% ------------------------------
% Version: Vampire 4.2.2 (commit c6486af3 on 2018-01-23 13:53:44 +0000)
% Termination reason: Refutation

% Memory used [KB]: 471804
% Time elapsed: 510.323 s
% ------------------------------
----  Runtime statistics ----
binary resolutions skipped for weight limit before building clause: 17306
binary resolutions skipped for weight limit while building clause: 26247
clauses created: 729705
clauses deleted: 111393
clauses discarded by weight limit in forward simplification: 629
clauses discarded from active on weight limit update: 1699
clauses discarded from passive on weight limit update: 140395
splitter_tautology: 16
ssat_new_components: 5133
ssat_non_splittable_sat_clauses: 2081
ssat_reused_components: 9629
ssat_sat_clauses: 9275
ssat_self_dependent_component: 7
superpositions skipped for weight limit while constructing other literals: 676
superpositions weight skipped after rewrited literal weight retrieval: 374
superpositions weight skipped after rewriter weight retrieval with occurrence counting: 52
superpositions weight skipped early: 1204
total_frozen: 318432
total_unfrozen: 58631
unworthy child removed: 269234
-----------------------------
% ------------------------------
