% Refutation found. Thanks to Tanya!
% SZS status Theorem for T07
% SZS output start Proof for T07
2. ! [X2] : ! [X3] : ! [X4] : (arc(X2,X3,X4) => (node(X4) & node(X3) & edge(X2))) [input sos20]
4. ! [X6] : (edge(X6) => ? [X7] : ? [X8] : arc(X6,X7,X8)) [input sos40]
22. ! [X49] : ! [X50] : (pp(X49,X50) <=> (X49 != X50 & p(X49,X50))) [input sos220]
43. ! [X95] : ! [X96] : ! [X97] : (path(X95,X96,X97) => (node(X97) & node(X96) & curve(X95))) [input sos430]
46. ! [X102] : ! [X103] : ! [X104] : (arc(X102,X103,X104) => path(X102,X103,X104)) [input sos460]
47. ! [X105] : ! [X106] : ! [X107] : ! [X108] : ((p(X108,X105) & arc(X105,X106,X107)) => X105 = X108) [input sos470]
48. ! [X109] : ! [X110] : ! [X111] : (path(X109,X110,X111) => (cont(X111,X109) & cont(X110,X109))) [input sos480]
49. ! [X112] : ! [X113] : ! [X114] : ! [X115] : ((cont(X112,X113) & path(X113,X114,X115) & node(X112)) => ? [X116] : ? [X117] : (p(X116,X113) & (path(X116,X117,X112) | path(X116,X112,X117)))) [input sos490]
53. ! [X137] : ! [X138] : (connects(X137,X138) <=> ? [X139] : ? [X140] : ? [X141] : (path(X138,X140,X141) & path(X137,X139,X140))) [input sos530]
59. ! [X172] : ! [X173] : (connects(X172,X173) => (~p(X173,X172) & ~p(X172,X173))) [input sos595]
62. ! [X191] : ! [X192] : ! [X193] : ((~arc(X191,X192,X193) & path(X191,X192,X193)) => ? [X194] : ? [X195] : ? [X196] : (pp(X196,X191) & pp(X194,X191) & connects(X194,X196) & arc(X194,X192,X195))) [input sos630]
64. ! [X17] : ! [X203] : ! [X204] : (path(X17,X203,X204) => ? [X205] : ? [X18] : ? [X206] : ? [X207] : (p(X18,X17) & path(X18,X207,X204) & p(X205,X17) & path(X205,X203,X206))) [input t70]
65. ~! [X17] : ! [X203] : ! [X204] : (path(X17,X203,X204) => ? [X205] : ? [X18] : ? [X206] : ? [X207] : (p(X18,X17) & path(X18,X207,X204) & p(X205,X17) & path(X205,X203,X206))) [negated conjecture 64]
67. ! [X0] : ! [X1] : ! [X2] : (arc(X0,X1,X2) => (node(X2) & node(X1) & edge(X0))) [rectify 2]
68. ! [X0,X1,X2] : (arc(X0,X1,X2) => (node(X2) & node(X1) & edge(X0))) [flattening 67]
70. ! [X0] : (edge(X0) => ? [X1] : ? [X2] : arc(X0,X1,X2)) [rectify 4]
71. ! [X0] : (edge(X0) => ? [X1,X2] : arc(X0,X1,X2)) [flattening 70]
100. ! [X0] : ! [X1] : (pp(X0,X1) <=> (X0 != X1 & p(X0,X1))) [rectify 22]
101. ! [X0,X1] : (pp(X0,X1) <=> (X0 != X1 & p(X0,X1))) [flattening 100]
137. ! [X0] : ! [X1] : ! [X2] : (path(X0,X1,X2) => (node(X2) & node(X1) & curve(X0))) [rectify 43]
138. ! [X0,X1,X2] : (path(X0,X1,X2) => (node(X2) & node(X1) & curve(X0))) [flattening 137]
142. ! [X0] : ! [X1] : ! [X2] : (arc(X0,X1,X2) => path(X0,X1,X2)) [rectify 46]
143. ! [X0,X1,X2] : (arc(X0,X1,X2) => path(X0,X1,X2)) [flattening 142]
144. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((p(X3,X0) & arc(X0,X1,X2)) => X0 = X3) [rectify 47]
145. ! [X0,X1,X2,X3] : ((p(X3,X0) & arc(X0,X1,X2)) => X0 = X3) [flattening 144]
146. ! [X0] : ! [X1] : ! [X2] : (path(X0,X1,X2) => (cont(X2,X0) & cont(X1,X0))) [rectify 48]
147. ! [X0,X1,X2] : (path(X0,X1,X2) => (cont(X2,X0) & cont(X1,X0))) [flattening 146]
148. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((cont(X0,X1) & path(X1,X2,X3) & node(X0)) => ? [X4] : ? [X5] : (p(X4,X1) & (path(X4,X5,X0) | path(X4,X0,X5)))) [rectify 49]
149. ! [X0,X1,X2,X3] : ((cont(X0,X1) & path(X1,X2,X3) & node(X0)) => ? [X4,X5] : (p(X4,X1) & (path(X4,X5,X0) | path(X4,X0,X5)))) [flattening 148]
156. ! [X0] : ! [X1] : (connects(X0,X1) <=> ? [X2] : ? [X3] : ? [X4] : (path(X1,X3,X4) & path(X0,X2,X3))) [rectify 53]
157. ! [X0,X1] : (connects(X0,X1) <=> ? [X2,X3,X4] : (path(X1,X3,X4) & path(X0,X2,X3))) [flattening 156]
168. ! [X0] : ! [X1] : (connects(X0,X1) => (~p(X1,X0) & ~p(X0,X1))) [rectify 59]
169. ! [X0,X1] : (connects(X0,X1) => (~p(X1,X0) & ~p(X0,X1))) [flattening 168]
174. ! [X0] : ! [X1] : ! [X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3] : ? [X4] : ? [X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4))) [rectify 62]
175. ! [X0,X1,X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4))) [flattening 174]
178. ~! [X0] : ! [X1] : ! [X2] : (path(X0,X1,X2) => ? [X3] : ? [X4] : ? [X5] : ? [X6] : (p(X4,X0) & path(X4,X6,X2) & p(X3,X0) & path(X3,X1,X5))) [rectify 65]
179. ~! [X0,X1,X2] : (path(X0,X1,X2) => ? [X3,X4,X5,X6] : (p(X4,X0) & path(X4,X6,X2) & p(X3,X0) & path(X3,X1,X5))) [flattening 178]
183. ! [X0,X1,X2] : ((node(X2) & node(X1) & edge(X0)) | ~arc(X0,X1,X2)) [ennf transformation 68]
185. ! [X0] : (? [X1,X2] : arc(X0,X1,X2) | ~edge(X0)) [ennf transformation 71]
217. ! [X0,X1,X2] : ((node(X2) & node(X1) & curve(X0)) | ~path(X0,X1,X2)) [ennf transformation 138]
220. ! [X0,X1,X2] : (path(X0,X1,X2) | ~arc(X0,X1,X2)) [ennf transformation 143]
221. ! [X0,X1,X2,X3] : (X0 = X3 | (~p(X3,X0) | ~arc(X0,X1,X2))) [ennf transformation 145]
222. ! [X0,X1,X2,X3] : (X0 = X3 | ~p(X3,X0) | ~arc(X0,X1,X2)) [flattening 221]
223. ! [X0,X1,X2] : ((cont(X2,X0) & cont(X1,X0)) | ~path(X0,X1,X2)) [ennf transformation 147]
224. ! [X0,X1,X2,X3] : (? [X4,X5] : (p(X4,X1) & (path(X4,X5,X0) | path(X4,X0,X5))) | (~cont(X0,X1) | ~path(X1,X2,X3) | ~node(X0))) [ennf transformation 149]
225. ! [X0,X1,X2,X3] : (? [X4,X5] : (p(X4,X1) & (path(X4,X5,X0) | path(X4,X0,X5))) | ~cont(X0,X1) | ~path(X1,X2,X3) | ~node(X0)) [flattening 224]
241. ! [X0,X1] : ((~p(X1,X0) & ~p(X0,X1)) | ~connects(X0,X1)) [ennf transformation 169]
246. ! [X0,X1,X2] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) | (arc(X0,X1,X2) | ~path(X0,X1,X2))) [ennf transformation 175]
247. ! [X0,X1,X2] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) | arc(X0,X1,X2) | ~path(X0,X1,X2)) [flattening 246]
248. ? [X0,X1,X2] : (! [X3,X4,X5,X6] : (~p(X4,X0) | ~path(X4,X6,X2) | ~p(X3,X0) | ~path(X3,X1,X5)) & path(X0,X1,X2)) [ennf transformation 179]
249. ! [X0] : (? [X1,X2] : arc(X0,X1,X2) => arc(X0,sK0(X0),sK1(X0))) [choice axiom]
250. ! [X0] : (arc(X0,sK0(X0),sK1(X0)) | ~edge(X0)) [skolemisation 185,249]
255. ! [X0,X1] : ((pp(X0,X1) | (X0 = X1 | ~p(X0,X1))) & ((X0 != X1 & p(X0,X1)) | ~pp(X0,X1))) [nnf transformation 101]
256. ! [X0,X1] : ((pp(X0,X1) | X0 = X1 | ~p(X0,X1)) & ((X0 != X1 & p(X0,X1)) | ~pp(X0,X1))) [flattening 255]
279. ! [X1,X0] : (? [X4,X5] : (p(X4,X1) & (path(X4,X5,X0) | path(X4,X0,X5))) => (p(sK7(X0,X1),X1) & (path(sK7(X0,X1),sK8(X0,X1),X0) | path(sK7(X0,X1),X0,sK8(X0,X1))))) [choice axiom]
280. ! [X0,X1,X2,X3] : ((p(sK7(X0,X1),X1) & (path(sK7(X0,X1),sK8(X0,X1),X0) | path(sK7(X0,X1),X0,sK8(X0,X1)))) | ~cont(X0,X1) | ~path(X1,X2,X3) | ~node(X0)) [skolemisation 225,279]
281. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & (? [X2,X3,X4] : (path(X1,X3,X4) & path(X0,X2,X3)) | ~connects(X0,X1))) [nnf transformation 157]
282. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & (? [X5,X6,X7] : (path(X1,X6,X7) & path(X0,X5,X6)) | ~connects(X0,X1))) [rectify 281]
283. ! [X1,X0] : (? [X5,X6,X7] : (path(X1,X6,X7) & path(X0,X5,X6)) => (path(X1,sK10(X0,X1),sK11(X0,X1)) & path(X0,sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
284. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & ((path(X1,sK10(X0,X1),sK11(X0,X1)) & path(X0,sK9(X0,X1),sK10(X0,X1))) | ~connects(X0,X1))) [skolemisation 282,283]
297. ! [X1,X0] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) => (pp(sK24(X0,X1),X0) & pp(sK22(X0,X1),X0) & connects(sK22(X0,X1),sK24(X0,X1)) & arc(sK22(X0,X1),X1,sK23(X0,X1)))) [choice axiom]
298. ! [X0,X1,X2] : ((pp(sK24(X0,X1),X0) & pp(sK22(X0,X1),X0) & connects(sK22(X0,X1),sK24(X0,X1)) & arc(sK22(X0,X1),X1,sK23(X0,X1))) | arc(X0,X1,X2) | ~path(X0,X1,X2)) [skolemisation 247,297]
299. ? [X0,X1,X2] : (! [X3,X4,X5,X6] : (~p(X4,X0) | ~path(X4,X6,X2) | ~p(X3,X0) | ~path(X3,X1,X5)) & path(X0,X1,X2)) => (! [X6,X5,X4,X3] : (~p(X4,sK25) | ~path(X4,X6,sK27) | ~p(X3,sK25) | ~path(X3,sK26,X5)) & path(sK25,sK26,sK27)) [choice axiom]
300. ! [X3,X4,X5,X6] : (~p(X4,sK25) | ~path(X4,X6,sK27) | ~p(X3,sK25) | ~path(X3,sK26,X5)) & path(sK25,sK26,sK27) [skolemisation 248,299]
302. ~arc(X0,X1,X2) | edge(X0) [cnf transformation 183]
306. arc(X0,sK0(X0),sK1(X0)) | ~edge(X0) [cnf transformation 250]
325. ~pp(X0,X1) | p(X0,X1) [cnf transformation 256]
358. ~path(X0,X1,X2) | node(X2) [cnf transformation 217]
361. ~arc(X0,X1,X2) | path(X0,X1,X2) [cnf transformation 220]
362. ~arc(X0,X1,X2) | ~p(X3,X0) | X0 = X3 [cnf transformation 222]
364. ~path(X0,X1,X2) | cont(X2,X0) [cnf transformation 223]
365. ~path(X1,X2,X3) | path(sK7(X0,X1),X0,sK8(X0,X1)) | ~cont(X0,X1) | path(sK7(X0,X1),sK8(X0,X1),X0) | ~node(X0) [cnf transformation 280]
366. ~path(X1,X2,X3) | ~cont(X0,X1) | p(sK7(X0,X1),X1) | ~node(X0) [cnf transformation 280]
373. ~path(X1,X3,X4) | connects(X0,X1) | ~path(X0,X2,X3) [cnf transformation 284]
393. ~connects(X0,X1) | ~p(X1,X0) [cnf transformation 241]
398. ~path(X0,X1,X2) | arc(X0,X1,X2) | arc(sK22(X0,X1),X1,sK23(X0,X1)) [cnf transformation 298]
400. ~path(X0,X1,X2) | arc(X0,X1,X2) | pp(sK22(X0,X1),X0) [cnf transformation 298]
402. path(sK25,sK26,sK27) [cnf transformation 300]
403. ~p(X4,sK25) | ~path(X4,X6,sK27) | ~p(X3,sK25) | ~path(X3,sK26,X5) [cnf transformation 300]
407. 0 <=> ! [X3,X5] : (~p(X3,sK25) | ~path(X3,sK26,X5)) [avatar definition]
408. ~path(X3,sK26,X5) | ~p(X3,sK25) <- {0} [avatar component clause 407]
410. 2 <=> ! [X4,X6] : (~p(X4,sK25) | ~path(X4,X6,sK27)) [avatar definition]
411. ~path(X4,X6,sK27) | ~p(X4,sK25) <- {2} [avatar component clause 410]
412. 0 | 2 [avatar split clause 403,410,407]
425. node(sK27) [resolution 358,402]
445. cont(sK27,sK25) [resolution 364,402]
604. ~p(X0,X1) | X0 = X1 | ~edge(X1) [resolution 362,306]
704. arc(sK25,sK26,sK27) | pp(sK22(sK25,sK26),sK25) [resolution 400,402]
711. 36 <=> pp(sK22(sK25,sK26),sK25) [avatar definition]
712. pp(sK22(sK25,sK26),sK25) <- {36} [avatar component clause 711]
717. 38 <=> arc(sK25,sK26,sK27) [avatar definition]
718. arc(sK25,sK26,sK27) <- {38} [avatar component clause 717]
719. 36 | 38 [avatar split clause 704,717,711]
738. p(sK7(X0,sK25),sK25) | ~cont(X0,sK25) | ~node(X0) [resolution 366,402]
768. arc(sK25,sK26,sK27) | arc(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) [resolution 398,402]
840. path(sK7(X0,sK25),sK8(X0,sK25),X0) | ~cont(X0,sK25) | path(sK7(X0,sK25),X0,sK8(X0,sK25)) | ~node(X0) [resolution 365,402]
1776. ~cont(sK27,sK25) | path(sK7(sK27,sK25),sK27,sK8(sK27,sK25)) | ~node(sK27) | ~p(sK7(sK27,sK25),sK25) <- {2} [resolution 840,411]
1828. 84 <=> path(sK7(sK27,sK25),sK27,sK8(sK27,sK25)) [avatar definition]
1829. path(sK7(sK27,sK25),sK27,sK8(sK27,sK25)) <- {84} [avatar component clause 1828]
1840. path(sK7(sK27,sK25),sK27,sK8(sK27,sK25)) | ~node(sK27) | ~p(sK7(sK27,sK25),sK25) <- {2} [subsumption resolution 1776,445]
1841. path(sK7(sK27,sK25),sK27,sK8(sK27,sK25)) | ~p(sK7(sK27,sK25),sK25) <- {2} [subsumption resolution 1840,425]
1843. 88 <=> p(sK7(sK27,sK25),sK25) [avatar definition]
1844. p(sK7(sK27,sK25),sK25) <- {88} [avatar component clause 1843]
1846. 89 <=> ~p(sK7(sK27,sK25),sK25) [avatar definition]
1847. ~p(sK7(sK27,sK25),sK25) <- {89} [avatar component clause 1846]
1848. ~89 | 84 | ~2 [avatar split clause 1841,410,1828,1846]
1866. ~cont(sK27,sK25) | ~node(sK27) <- {89} [resolution 1847,738]
1867. ~node(sK27) <- {89} [subsumption resolution 1866,445]
1868. $false <- {89} [subsumption resolution 1867,425]
1869. 89 [avatar contradiction clause 1868]
1883. sK7(sK27,sK25) = sK25 | ~edge(sK25) <- {88} [resolution 1844,604]
1898. 90 <=> sK7(sK27,sK25) = sK25 [avatar definition]
1899. sK7(sK27,sK25) = sK25 <- {90} [avatar component clause 1898]
1902. 92 <=> edge(sK25) [avatar definition]
1905. 93 <=> ~edge(sK25) [avatar definition]
1907. ~93 | 90 | ~88 [avatar split clause 1883,1843,1898,1905]
1930. 100 <=> arc(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) [avatar definition]
1931. arc(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) <- {100} [avatar component clause 1930]
1932. 100 | 38 [avatar split clause 768,717,1930]
1941. ~p(sK25,sK25) <- {0} [resolution 408,402]
1944. edge(sK25) <- {38} [resolution 718,302]
1958. 92 | ~38 [avatar split clause 1944,717,1902]
1961. p(sK25,sK25) <- {88, 90} [backward demodulation 1899,1844]
1966. $false <- {0, 88, 90} [subsumption resolution 1961,1941]
1967. ~0 | ~88 | ~90 [avatar contradiction clause 1966]
1968. p(sK22(sK25,sK26),sK25) <- {36} [resolution 712,325]
2009. path(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) <- {100} [resolution 1931,361]
2634. ~p(sK22(sK25,sK26),sK25) <- {0, 100} [resolution 2009,408]
2683. $false <- {0, 36, 100} [subsumption resolution 2634,1968]
2684. ~0 | ~36 | ~100 [avatar contradiction clause 2683]
3139. ~path(X14,X15,sK27) | connects(X14,sK7(sK27,sK25)) <- {84} [resolution 1829,373]
25130. connects(sK25,sK7(sK27,sK25)) <- {84} [resolution 3139,402]
25147. ~p(sK7(sK27,sK25),sK25) <- {84} [resolution 25130,393]
25215. $false <- {84, 88} [subsumption resolution 25147,1844]
25216. ~84 | ~88 [avatar contradiction clause 25215]
25277. $false [avatar sat refutation 412,719,1848,1869,1907,1932,1958,1967,2684,25216]
% SZS output end Proof for T07
% ------------------------------
% Version: Vampire 4.2.2 (commit c6486af3 on 2018-01-23 13:53:44 +0000)
% Termination reason: Refutation

% Memory used [KB]: 31342
% Time elapsed: 12.106 s
% ------------------------------
----  Runtime statistics ----
clauses created: 23658
clauses deleted: 3580
ssat_new_components: 456
ssat_non_splittable_sat_clauses: 101
ssat_reused_components: 301
ssat_sat_clauses: 497
total_frozen: 4814
total_unfrozen: 1298
unworthy child removed: 4738
-----------------------------
% ------------------------------
