% Refutation found. Thanks to Tanya!
% SZS status Theorem for T08
% SZS output start Proof for T08
2. ! [X2] : ! [X3] : ! [X4] : (arc(X2,X3,X4) => (node(X4) & node(X3) & edge(X2))) [input sos20]
4. ! [X6] : (edge(X6) => ? [X7] : ? [X8] : arc(X6,X7,X8)) [input sos40]
5. ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ((arc(X9,X13,X14) & arc(X9,X10,X11)) => (X11 = X14 & X10 = X13)) [input sos50]
45. ! [X99] : ! [X100] : ! [X101] : (path(X99,X100,X101) => X100 != X101) [input sos450]
46. ! [X102] : ! [X103] : ! [X104] : (arc(X102,X103,X104) => path(X102,X103,X104)) [input sos460]
50. ! [X118] : ! [X119] : ! [X120] : ! [X121] : ! [X122] : ((path(X118,X121,X122) & path(X118,X119,X120)) => (X120 = X122 & X119 = X121)) [input sos500]
53. ! [X137] : ! [X138] : (connects(X137,X138) <=> ? [X139] : ? [X140] : ? [X141] : (path(X138,X140,X141) & path(X137,X139,X140))) [input sos530]
62. ! [X191] : ! [X192] : ! [X193] : ((~arc(X191,X192,X193) & path(X191,X192,X193)) => ? [X194] : ? [X195] : ? [X196] : (pp(X196,X191) & pp(X194,X191) & connects(X194,X196) & arc(X194,X192,X195))) [input sos630]
64. ! [X203] : ! [X204] : ! [X205] : ((~arc(X203,X204,X205) & path(X203,X204,X205)) => ? [X206] : ? [X20] : ? [X207] : ? [X19] : ? [X208] : ? [X209] : (pp(X20,X203) & pp(X206,X203) & X20 != X206 & path(X20,X208,X209) & path(X206,X207,X19))) [input t80]
65. ~! [X203] : ! [X204] : ! [X205] : ((~arc(X203,X204,X205) & path(X203,X204,X205)) => ? [X206] : ? [X20] : ? [X207] : ? [X19] : ? [X208] : ? [X209] : (pp(X20,X203) & pp(X206,X203) & X20 != X206 & path(X20,X208,X209) & path(X206,X207,X19))) [negated conjecture 64]
67. ! [X0] : ! [X1] : ! [X2] : (arc(X0,X1,X2) => (node(X2) & node(X1) & edge(X0))) [rectify 2]
68. ! [X0,X1,X2] : (arc(X0,X1,X2) => (node(X2) & node(X1) & edge(X0))) [flattening 67]
70. ! [X0] : (edge(X0) => ? [X1] : ? [X2] : arc(X0,X1,X2)) [rectify 4]
71. ! [X0] : (edge(X0) => ? [X1,X2] : arc(X0,X1,X2)) [flattening 70]
72. ! [X0] : ! [X1] : ! [X2] : ! [X4] : ! [X5] : ((arc(X0,X4,X5) & arc(X0,X1,X2)) => (X2 = X5 & X1 = X4)) [rectify 5]
73. ! [X0,X1,X2,X4,X5] : ((arc(X0,X4,X5) & arc(X0,X1,X2)) => (X2 = X5 & X1 = X4)) [flattening 72]
140. ! [X0] : ! [X1] : ! [X2] : (path(X0,X1,X2) => X1 != X2) [rectify 45]
141. ! [X0,X1,X2] : (path(X0,X1,X2) => X1 != X2) [flattening 140]
142. ! [X0] : ! [X1] : ! [X2] : (arc(X0,X1,X2) => path(X0,X1,X2)) [rectify 46]
143. ! [X0,X1,X2] : (arc(X0,X1,X2) => path(X0,X1,X2)) [flattening 142]
150. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ((path(X0,X3,X4) & path(X0,X1,X2)) => (X2 = X4 & X1 = X3)) [rectify 50]
151. ! [X0,X1,X2,X3,X4] : ((path(X0,X3,X4) & path(X0,X1,X2)) => (X2 = X4 & X1 = X3)) [flattening 150]
156. ! [X0] : ! [X1] : (connects(X0,X1) <=> ? [X2] : ? [X3] : ? [X4] : (path(X1,X3,X4) & path(X0,X2,X3))) [rectify 53]
157. ! [X0,X1] : (connects(X0,X1) <=> ? [X2,X3,X4] : (path(X1,X3,X4) & path(X0,X2,X3))) [flattening 156]
174. ! [X0] : ! [X1] : ! [X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3] : ? [X4] : ? [X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4))) [rectify 62]
175. ! [X0,X1,X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4))) [flattening 174]
178. ~! [X0] : ! [X1] : ! [X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3] : ? [X4] : ? [X5] : ? [X6] : ? [X7] : ? [X8] : (pp(X4,X0) & pp(X3,X0) & X3 != X4 & path(X4,X7,X8) & path(X3,X5,X6))) [rectify 65]
179. ~! [X0,X1,X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3,X4,X5,X6,X7,X8] : (pp(X4,X0) & pp(X3,X0) & X3 != X4 & path(X4,X7,X8) & path(X3,X5,X6))) [flattening 178]
183. ! [X0,X1,X2] : ((node(X2) & node(X1) & edge(X0)) | ~arc(X0,X1,X2)) [ennf transformation 68]
185. ! [X0] : (? [X1,X2] : arc(X0,X1,X2) | ~edge(X0)) [ennf transformation 71]
186. ! [X0,X1,X2,X4,X5] : ((X2 = X5 & X1 = X4) | (~arc(X0,X4,X5) | ~arc(X0,X1,X2))) [ennf transformation 73]
187. ! [X0,X1,X2,X4,X5] : ((X2 = X5 & X1 = X4) | ~arc(X0,X4,X5) | ~arc(X0,X1,X2)) [flattening 186]
219. ! [X0,X1,X2] : (X1 != X2 | ~path(X0,X1,X2)) [ennf transformation 141]
220. ! [X0,X1,X2] : (path(X0,X1,X2) | ~arc(X0,X1,X2)) [ennf transformation 143]
226. ! [X0,X1,X2,X3,X4] : ((X2 = X4 & X1 = X3) | (~path(X0,X3,X4) | ~path(X0,X1,X2))) [ennf transformation 151]
227. ! [X0,X1,X2,X3,X4] : ((X2 = X4 & X1 = X3) | ~path(X0,X3,X4) | ~path(X0,X1,X2)) [flattening 226]
246. ! [X0,X1,X2] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) | (arc(X0,X1,X2) | ~path(X0,X1,X2))) [ennf transformation 175]
247. ! [X0,X1,X2] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) | arc(X0,X1,X2) | ~path(X0,X1,X2)) [flattening 246]
248. ? [X0,X1,X2] : (! [X3,X4,X5,X6,X7,X8] : (~pp(X4,X0) | ~pp(X3,X0) | X3 = X4 | ~path(X4,X7,X8) | ~path(X3,X5,X6)) & (~arc(X0,X1,X2) & path(X0,X1,X2))) [ennf transformation 179]
249. ? [X0,X1,X2] : (! [X3,X4,X5,X6,X7,X8] : (~pp(X4,X0) | ~pp(X3,X0) | X3 = X4 | ~path(X4,X7,X8) | ~path(X3,X5,X6)) & ~arc(X0,X1,X2) & path(X0,X1,X2)) [flattening 248]
250. ! [X0] : (? [X1,X2] : arc(X0,X1,X2) => arc(X0,sK0(X0),sK1(X0))) [choice axiom]
251. ! [X0] : (arc(X0,sK0(X0),sK1(X0)) | ~edge(X0)) [skolemisation 185,250]
252. ! [X0,X1,X2,X3,X4] : ((X2 = X4 & X1 = X3) | ~arc(X0,X3,X4) | ~arc(X0,X1,X2)) [rectify 187]
282. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & (? [X2,X3,X4] : (path(X1,X3,X4) & path(X0,X2,X3)) | ~connects(X0,X1))) [nnf transformation 157]
283. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & (? [X5,X6,X7] : (path(X1,X6,X7) & path(X0,X5,X6)) | ~connects(X0,X1))) [rectify 282]
284. ! [X1,X0] : (? [X5,X6,X7] : (path(X1,X6,X7) & path(X0,X5,X6)) => (path(X1,sK10(X0,X1),sK11(X0,X1)) & path(X0,sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
285. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & ((path(X1,sK10(X0,X1),sK11(X0,X1)) & path(X0,sK9(X0,X1),sK10(X0,X1))) | ~connects(X0,X1))) [skolemisation 283,284]
298. ! [X1,X0] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) => (pp(sK24(X0,X1),X0) & pp(sK22(X0,X1),X0) & connects(sK22(X0,X1),sK24(X0,X1)) & arc(sK22(X0,X1),X1,sK23(X0,X1)))) [choice axiom]
299. ! [X0,X1,X2] : ((pp(sK24(X0,X1),X0) & pp(sK22(X0,X1),X0) & connects(sK22(X0,X1),sK24(X0,X1)) & arc(sK22(X0,X1),X1,sK23(X0,X1))) | arc(X0,X1,X2) | ~path(X0,X1,X2)) [skolemisation 247,298]
300. ? [X0,X1,X2] : (! [X3,X4,X5,X6,X7,X8] : (~pp(X4,X0) | ~pp(X3,X0) | X3 = X4 | ~path(X4,X7,X8) | ~path(X3,X5,X6)) & ~arc(X0,X1,X2) & path(X0,X1,X2)) => (! [X8,X7,X6,X5,X4,X3] : (~pp(X4,sK25) | ~pp(X3,sK25) | X3 = X4 | ~path(X4,X7,X8) | ~path(X3,X5,X6)) & ~arc(sK25,sK26,sK27) & path(sK25,sK26,sK27)) [choice axiom]
301. ! [X3,X4,X5,X6,X7,X8] : (~pp(X4,sK25) | ~pp(X3,sK25) | X3 = X4 | ~path(X4,X7,X8) | ~path(X3,X5,X6)) & ~arc(sK25,sK26,sK27) & path(sK25,sK26,sK27) [skolemisation 249,300]
303. ~arc(X0,X1,X2) | edge(X0) [cnf transformation 183]
307. arc(X0,sK0(X0),sK1(X0)) | ~edge(X0) [cnf transformation 251]
308. ~arc(X0,X3,X4) | X1 = X3 | ~arc(X0,X1,X2) [cnf transformation 252]
309. ~arc(X0,X3,X4) | X2 = X4 | ~arc(X0,X1,X2) [cnf transformation 252]
361. X1 != X2 | ~path(X0,X1,X2) [cnf transformation 219]
362. ~arc(X0,X1,X2) | path(X0,X1,X2) [cnf transformation 220]
368. ~path(X0,X3,X4) | X1 = X3 | ~path(X0,X1,X2) [cnf transformation 227]
369. ~path(X0,X3,X4) | X2 = X4 | ~path(X0,X1,X2) [cnf transformation 227]
372. path(X0,sK9(X0,X1),sK10(X0,X1)) | ~connects(X0,X1) [cnf transformation 285]
373. path(X1,sK10(X0,X1),sK11(X0,X1)) | ~connects(X0,X1) [cnf transformation 285]
399. ~path(X0,X1,X2) | arc(X0,X1,X2) | arc(sK22(X0,X1),X1,sK23(X0,X1)) [cnf transformation 299]
400. ~path(X0,X1,X2) | arc(X0,X1,X2) | connects(sK22(X0,X1),sK24(X0,X1)) [cnf transformation 299]
401. ~path(X0,X1,X2) | arc(X0,X1,X2) | pp(sK22(X0,X1),X0) [cnf transformation 299]
402. ~path(X0,X1,X2) | arc(X0,X1,X2) | pp(sK24(X0,X1),X0) [cnf transformation 299]
403. path(sK25,sK26,sK27) [cnf transformation 301]
404. ~arc(sK25,sK26,sK27) [cnf transformation 301]
405. ~path(X4,X7,X8) | ~pp(X3,sK25) | X3 = X4 | ~pp(X4,sK25) | ~path(X3,X5,X6) [cnf transformation 301]
407. ~path(X0,X2,X2) [equality resolution 361]
447. path(X0,sK0(X0),sK1(X0)) | ~edge(X0) [resolution 362,307]
583. sK0(X1) = X0 | ~arc(X1,X0,X2) | ~edge(X1) [resolution 308,307]
584. ~arc(X1,X0,X2) | sK0(X1) = X0 [subsumption resolution 583,303]
586. sK1(X1) = X0 | ~arc(X1,X2,X0) | ~edge(X1) [resolution 309,307]
587. ~arc(X1,X2,X0) | sK1(X1) = X0 [subsumption resolution 586,303]
609. ~path(X14,X17,X16) | sK10(X14,X15) = X16 | ~connects(X14,X15) [resolution 372,369]
615. ~path(X11,X12,X13) | sK10(X10,X11) = X12 | ~connects(X10,X11) [resolution 373,368]
633. arc(sK25,sK26,sK27) | pp(sK22(sK25,sK26),sK25) [resolution 401,403]
636. pp(sK22(sK25,sK26),sK25) [subsumption resolution 633,404]
638. arc(sK25,sK26,sK27) | pp(sK24(sK25,sK26),sK25) [resolution 402,403]
641. pp(sK24(sK25,sK26),sK25) [subsumption resolution 638,404]
649. arc(sK25,sK26,sK27) | connects(sK22(sK25,sK26),sK24(sK25,sK26)) [resolution 400,403]
652. connects(sK22(sK25,sK26),sK24(sK25,sK26)) [subsumption resolution 649,404]
686. arc(sK25,sK26,sK27) | arc(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) [resolution 399,403]
689. arc(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) [subsumption resolution 686,404]
697. ~path(X8,X10,X11) | X8 = X9 | ~pp(X9,sK25) | ~pp(X8,sK25) | ~connects(X12,X9) [resolution 405,373]
921. edge(sK22(sK25,sK26)) [resolution 689,303]
926. path(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) [resolution 689,362]
1023. sK0(sK22(sK25,sK26)) = sK26 [resolution 584,689]
1039. sK1(sK22(sK25,sK26)) = sK23(sK25,sK26) [resolution 587,689]
1314. ~connects(X7,X8) | sK1(X7) = sK10(X7,X8) | ~edge(X7) [resolution 609,447]
1318. ~connects(X7,X8) | sK0(X8) = sK10(X7,X8) | ~edge(X8) [resolution 615,447]
1575. ~connects(X10,X9) | ~pp(X9,sK25) | ~pp(X8,sK25) | X8 = X9 | ~edge(X8) [resolution 697,447]
2980. sK1(sK22(sK25,sK26)) = sK10(sK22(sK25,sK26),sK24(sK25,sK26)) | ~edge(sK22(sK25,sK26)) [resolution 1314,652]
2984. sK1(sK22(sK25,sK26)) = sK10(sK22(sK25,sK26),sK24(sK25,sK26)) [subsumption resolution 2980,921]
2985. sK10(sK22(sK25,sK26),sK24(sK25,sK26)) = sK23(sK25,sK26) [forward demodulation 2984,1039]
4043. ~pp(sK24(sK25,sK26),sK25) | ~pp(X0,sK25) | sK24(sK25,sK26) = X0 | ~edge(X0) [resolution 1575,652]
4047. ~pp(X0,sK25) | sK24(sK25,sK26) = X0 | ~edge(X0) [subsumption resolution 4043,641]
8312. sK0(sK24(sK25,sK26)) = sK10(sK22(sK25,sK26),sK24(sK25,sK26)) | ~edge(sK24(sK25,sK26)) [resolution 652,1318]
8788. 397 <=> sK23(sK25,sK26) != sK26 [avatar definition]
8789. sK23(sK25,sK26) != sK26 <- {397} [avatar component clause 8788]
8791. 396 <=> sK23(sK25,sK26) = sK26 [avatar definition]
8792. sK23(sK25,sK26) = sK26 <- {396} [avatar component clause 8791]
24876. path(sK22(sK25,sK26),sK26,sK26) <- {396} [backward demodulation 8792,926]
24965. $false <- {396} [subsumption resolution 24876,407]
24966. ~396 [avatar contradiction clause 24965]
31545. sK22(sK25,sK26) = sK24(sK25,sK26) | ~edge(sK22(sK25,sK26)) [resolution 4047,636]
31580. sK22(sK25,sK26) = sK24(sK25,sK26) [subsumption resolution 31545,921]
31612. sK10(sK22(sK25,sK26),sK22(sK25,sK26)) = sK23(sK25,sK26) [backward demodulation 31580,2985]
33116. sK0(sK22(sK25,sK26)) = sK10(sK22(sK25,sK26),sK22(sK25,sK26)) | ~edge(sK24(sK25,sK26)) [forward demodulation 8312,31580]
33153. sK10(sK22(sK25,sK26),sK22(sK25,sK26)) = sK26 | ~edge(sK24(sK25,sK26)) [forward demodulation 33116,1023]
33154. ~edge(sK22(sK25,sK26)) | sK10(sK22(sK25,sK26),sK22(sK25,sK26)) = sK26 [forward demodulation 33153,31580]
33155. sK10(sK22(sK25,sK26),sK22(sK25,sK26)) = sK26 [subsumption resolution 33154,921]
33157. sK23(sK25,sK26) = sK26 [forward demodulation 31612,33155]
33158. $false <- {397} [subsumption resolution 33157,8789]
33159. 397 [avatar contradiction clause 33158]
33235. $false [avatar sat refutation 24966,33159]
% SZS output end Proof for T08
% ------------------------------
% Version: Vampire 4.2.2 (commit c6486af3 on 2018-01-23 13:53:44 +0000)
% Termination reason: Refutation

% Memory used [KB]: 35820
% Time elapsed: 17.699 s
% ------------------------------
----  Runtime statistics ----
clauses created: 31414
clauses deleted: 7503
ssat_new_components: 540
ssat_non_splittable_sat_clauses: 101
ssat_reused_components: 293
ssat_sat_clauses: 550
total_frozen: 6118
total_unfrozen: 1662
unworthy child removed: 7028
-----------------------------
% ------------------------------
