% Refutation found. Thanks to Tanya!
% SZS status Theorem for 24apr2018_full_T08
% SZS output start Proof for 24apr2018_full_T08
2. ! [X2] : ! [X3] : ! [X4] : (arc(X2,X3,X4) => (node(X4) & node(X3) & edge(X2))) [input sos20]
4. ! [X6] : (edge(X6) => ? [X7] : ? [X8] : arc(X6,X7,X8)) [input sos40]
5. ! [X9] : ! [X10] : ! [X11] : ! [X12] : ! [X13] : ! [X14] : ((arc(X9,X13,X14) & arc(X9,X10,X11)) => (X11 = X14 & X10 = X13)) [input sos50]
45. ! [X99] : ! [X100] : ! [X101] : (path(X99,X100,X101) => X100 != X101) [input sos450]
46. ! [X102] : ! [X103] : ! [X104] : (arc(X102,X103,X104) => path(X102,X103,X104)) [input sos460]
50. ! [X118] : ! [X119] : ! [X120] : ! [X121] : ! [X122] : ((path(X118,X121,X122) & path(X118,X119,X120)) => (X120 = X122 & X119 = X121)) [input sos500]
53. ! [X137] : ! [X138] : (connects(X137,X138) <=> ? [X139] : ? [X140] : ? [X141] : (path(X138,X140,X141) & path(X137,X139,X140))) [input sos530]
63. ! [X201] : ! [X202] : ! [X203] : ((~arc(X201,X202,X203) & path(X201,X202,X203)) => ? [X204] : ? [X205] : ? [X206] : (pp(X206,X201) & pp(X204,X201) & connects(X204,X206) & arc(X204,X202,X205))) [input sos630]
65. ! [X213] : ! [X214] : ! [X215] : ((~arc(X213,X214,X215) & path(X213,X214,X215)) => ? [X216] : ? [X20] : ? [X217] : ? [X19] : ? [X218] : ? [X219] : (pp(X20,X213) & pp(X216,X213) & X20 != X216 & path(X20,X218,X219) & path(X216,X217,X19))) [input t80]
66. ~! [X213] : ! [X214] : ! [X215] : ((~arc(X213,X214,X215) & path(X213,X214,X215)) => ? [X216] : ? [X20] : ? [X217] : ? [X19] : ? [X218] : ? [X219] : (pp(X20,X213) & pp(X216,X213) & X20 != X216 & path(X20,X218,X219) & path(X216,X217,X19))) [negated conjecture 65]
68. ! [X0] : ! [X1] : ! [X2] : (arc(X0,X1,X2) => (node(X2) & node(X1) & edge(X0))) [rectify 2]
69. ! [X0,X1,X2] : (arc(X0,X1,X2) => (node(X2) & node(X1) & edge(X0))) [flattening 68]
71. ! [X0] : (edge(X0) => ? [X1] : ? [X2] : arc(X0,X1,X2)) [rectify 4]
72. ! [X0] : (edge(X0) => ? [X1,X2] : arc(X0,X1,X2)) [flattening 71]
73. ! [X0] : ! [X1] : ! [X2] : ! [X4] : ! [X5] : ((arc(X0,X4,X5) & arc(X0,X1,X2)) => (X2 = X5 & X1 = X4)) [rectify 5]
74. ! [X0,X1,X2,X4,X5] : ((arc(X0,X4,X5) & arc(X0,X1,X2)) => (X2 = X5 & X1 = X4)) [flattening 73]
141. ! [X0] : ! [X1] : ! [X2] : (path(X0,X1,X2) => X1 != X2) [rectify 45]
142. ! [X0,X1,X2] : (path(X0,X1,X2) => X1 != X2) [flattening 141]
143. ! [X0] : ! [X1] : ! [X2] : (arc(X0,X1,X2) => path(X0,X1,X2)) [rectify 46]
144. ! [X0,X1,X2] : (arc(X0,X1,X2) => path(X0,X1,X2)) [flattening 143]
151. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ((path(X0,X3,X4) & path(X0,X1,X2)) => (X2 = X4 & X1 = X3)) [rectify 50]
152. ! [X0,X1,X2,X3,X4] : ((path(X0,X3,X4) & path(X0,X1,X2)) => (X2 = X4 & X1 = X3)) [flattening 151]
157. ! [X0] : ! [X1] : (connects(X0,X1) <=> ? [X2] : ? [X3] : ? [X4] : (path(X1,X3,X4) & path(X0,X2,X3))) [rectify 53]
158. ! [X0,X1] : (connects(X0,X1) <=> ? [X2,X3,X4] : (path(X1,X3,X4) & path(X0,X2,X3))) [flattening 157]
177. ! [X0] : ! [X1] : ! [X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3] : ? [X4] : ? [X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4))) [rectify 63]
178. ! [X0,X1,X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4))) [flattening 177]
181. ~! [X0] : ! [X1] : ! [X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3] : ? [X4] : ? [X5] : ? [X6] : ? [X7] : ? [X8] : (pp(X4,X0) & pp(X3,X0) & X3 != X4 & path(X4,X7,X8) & path(X3,X5,X6))) [rectify 66]
182. ~! [X0,X1,X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3,X4,X5,X6,X7,X8] : (pp(X4,X0) & pp(X3,X0) & X3 != X4 & path(X4,X7,X8) & path(X3,X5,X6))) [flattening 181]
186. ! [X0,X1,X2] : ((node(X2) & node(X1) & edge(X0)) | ~arc(X0,X1,X2)) [ennf transformation 69]
188. ! [X0] : (? [X1,X2] : arc(X0,X1,X2) | ~edge(X0)) [ennf transformation 72]
189. ! [X0,X1,X2,X4,X5] : ((X2 = X5 & X1 = X4) | (~arc(X0,X4,X5) | ~arc(X0,X1,X2))) [ennf transformation 74]
190. ! [X0,X1,X2,X4,X5] : ((X2 = X5 & X1 = X4) | ~arc(X0,X4,X5) | ~arc(X0,X1,X2)) [flattening 189]
222. ! [X0,X1,X2] : (X1 != X2 | ~path(X0,X1,X2)) [ennf transformation 142]
223. ! [X0,X1,X2] : (path(X0,X1,X2) | ~arc(X0,X1,X2)) [ennf transformation 144]
229. ! [X0,X1,X2,X3,X4] : ((X2 = X4 & X1 = X3) | (~path(X0,X3,X4) | ~path(X0,X1,X2))) [ennf transformation 152]
230. ! [X0,X1,X2,X3,X4] : ((X2 = X4 & X1 = X3) | ~path(X0,X3,X4) | ~path(X0,X1,X2)) [flattening 229]
252. ! [X0,X1,X2] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) | (arc(X0,X1,X2) | ~path(X0,X1,X2))) [ennf transformation 178]
253. ! [X0,X1,X2] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) | arc(X0,X1,X2) | ~path(X0,X1,X2)) [flattening 252]
254. ? [X0,X1,X2] : (! [X3,X4,X5,X6,X7,X8] : (~pp(X4,X0) | ~pp(X3,X0) | X3 = X4 | ~path(X4,X7,X8) | ~path(X3,X5,X6)) & (~arc(X0,X1,X2) & path(X0,X1,X2))) [ennf transformation 182]
255. ? [X0,X1,X2] : (! [X3,X4,X5,X6,X7,X8] : (~pp(X4,X0) | ~pp(X3,X0) | X3 = X4 | ~path(X4,X7,X8) | ~path(X3,X5,X6)) & ~arc(X0,X1,X2) & path(X0,X1,X2)) [flattening 254]
256. ! [X0] : (? [X1,X2] : arc(X0,X1,X2) => arc(X0,sK0(X0),sK1(X0))) [choice axiom]
257. ! [X0] : (arc(X0,sK0(X0),sK1(X0)) | ~edge(X0)) [skolemisation 188,256]
258. ! [X0,X1,X2,X3,X4] : ((X2 = X4 & X1 = X3) | ~arc(X0,X3,X4) | ~arc(X0,X1,X2)) [rectify 190]
288. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & (? [X2,X3,X4] : (path(X1,X3,X4) & path(X0,X2,X3)) | ~connects(X0,X1))) [nnf transformation 158]
289. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & (? [X5,X6,X7] : (path(X1,X6,X7) & path(X0,X5,X6)) | ~connects(X0,X1))) [rectify 288]
290. ! [X1,X0] : (? [X5,X6,X7] : (path(X1,X6,X7) & path(X0,X5,X6)) => (path(X1,sK10(X0,X1),sK11(X0,X1)) & path(X0,sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
291. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & ((path(X1,sK10(X0,X1),sK11(X0,X1)) & path(X0,sK9(X0,X1),sK10(X0,X1))) | ~connects(X0,X1))) [skolemisation 289,290]
304. ! [X1,X0] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) => (pp(sK24(X0,X1),X0) & pp(sK22(X0,X1),X0) & connects(sK22(X0,X1),sK24(X0,X1)) & arc(sK22(X0,X1),X1,sK23(X0,X1)))) [choice axiom]
305. ! [X0,X1,X2] : ((pp(sK24(X0,X1),X0) & pp(sK22(X0,X1),X0) & connects(sK22(X0,X1),sK24(X0,X1)) & arc(sK22(X0,X1),X1,sK23(X0,X1))) | arc(X0,X1,X2) | ~path(X0,X1,X2)) [skolemisation 253,304]
306. ? [X0,X1,X2] : (! [X3,X4,X5,X6,X7,X8] : (~pp(X4,X0) | ~pp(X3,X0) | X3 = X4 | ~path(X4,X7,X8) | ~path(X3,X5,X6)) & ~arc(X0,X1,X2) & path(X0,X1,X2)) => (! [X8,X7,X6,X5,X4,X3] : (~pp(X4,sK25) | ~pp(X3,sK25) | X3 = X4 | ~path(X4,X7,X8) | ~path(X3,X5,X6)) & ~arc(sK25,sK26,sK27) & path(sK25,sK26,sK27)) [choice axiom]
307. ! [X3,X4,X5,X6,X7,X8] : (~pp(X4,sK25) | ~pp(X3,sK25) | X3 = X4 | ~path(X4,X7,X8) | ~path(X3,X5,X6)) & ~arc(sK25,sK26,sK27) & path(sK25,sK26,sK27) [skolemisation 255,306]
309. ~arc(X0,X1,X2) | edge(X0) [cnf transformation 186]
313. arc(X0,sK0(X0),sK1(X0)) | ~edge(X0) [cnf transformation 257]
314. ~arc(X0,X3,X4) | X1 = X3 | ~arc(X0,X1,X2) [cnf transformation 258]
315. ~arc(X0,X3,X4) | X2 = X4 | ~arc(X0,X1,X2) [cnf transformation 258]
367. X1 != X2 | ~path(X0,X1,X2) [cnf transformation 222]
368. ~arc(X0,X1,X2) | path(X0,X1,X2) [cnf transformation 223]
374. ~path(X0,X3,X4) | X1 = X3 | ~path(X0,X1,X2) [cnf transformation 230]
375. ~path(X0,X3,X4) | X2 = X4 | ~path(X0,X1,X2) [cnf transformation 230]
378. path(X0,sK9(X0,X1),sK10(X0,X1)) | ~connects(X0,X1) [cnf transformation 291]
379. path(X1,sK10(X0,X1),sK11(X0,X1)) | ~connects(X0,X1) [cnf transformation 291]
405. ~path(X0,X1,X2) | arc(X0,X1,X2) | arc(sK22(X0,X1),X1,sK23(X0,X1)) [cnf transformation 305]
406. ~path(X0,X1,X2) | arc(X0,X1,X2) | connects(sK22(X0,X1),sK24(X0,X1)) [cnf transformation 305]
407. ~path(X0,X1,X2) | arc(X0,X1,X2) | pp(sK22(X0,X1),X0) [cnf transformation 305]
408. ~path(X0,X1,X2) | arc(X0,X1,X2) | pp(sK24(X0,X1),X0) [cnf transformation 305]
409. path(sK25,sK26,sK27) [cnf transformation 307]
410. ~arc(sK25,sK26,sK27) [cnf transformation 307]
411. ~path(X4,X7,X8) | ~pp(X3,sK25) | X3 = X4 | ~pp(X4,sK25) | ~path(X3,X5,X6) [cnf transformation 307]
413. ~path(X0,X2,X2) [equality resolution 367]
453. path(X0,sK0(X0),sK1(X0)) | ~edge(X0) [resolution 368,313]
588. sK0(X1) = X0 | ~arc(X1,X0,X2) | ~edge(X1) [resolution 314,313]
589. ~arc(X1,X0,X2) | sK0(X1) = X0 [subsumption resolution 588,309]
591. sK1(X1) = X0 | ~arc(X1,X2,X0) | ~edge(X1) [resolution 315,313]
592. ~arc(X1,X2,X0) | sK1(X1) = X0 [subsumption resolution 591,309]
615. ~path(X14,X17,X16) | sK10(X14,X15) = X16 | ~connects(X14,X15) [resolution 378,375]
621. ~path(X11,X12,X13) | sK10(X10,X11) = X12 | ~connects(X10,X11) [resolution 379,374]
636. arc(sK25,sK26,sK27) | pp(sK22(sK25,sK26),sK25) [resolution 407,409]
639. pp(sK22(sK25,sK26),sK25) [subsumption resolution 636,410]
641. arc(sK25,sK26,sK27) | pp(sK24(sK25,sK26),sK25) [resolution 408,409]
644. pp(sK24(sK25,sK26),sK25) [subsumption resolution 641,410]
652. arc(sK25,sK26,sK27) | connects(sK22(sK25,sK26),sK24(sK25,sK26)) [resolution 406,409]
655. connects(sK22(sK25,sK26),sK24(sK25,sK26)) [subsumption resolution 652,410]
689. arc(sK25,sK26,sK27) | arc(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) [resolution 405,409]
692. arc(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) [subsumption resolution 689,410]
716. ~path(X8,X10,X11) | X8 = X9 | ~pp(X9,sK25) | ~pp(X8,sK25) | ~connects(X12,X9) [resolution 411,379]
933. edge(sK22(sK25,sK26)) [resolution 692,309]
938. path(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) [resolution 692,368]
1037. sK0(sK22(sK25,sK26)) = sK26 [resolution 589,692]
1053. sK1(sK22(sK25,sK26)) = sK23(sK25,sK26) [resolution 592,692]
1328. ~connects(X7,X8) | sK1(X7) = sK10(X7,X8) | ~edge(X7) [resolution 615,453]
1332. ~connects(X7,X8) | sK0(X8) = sK10(X7,X8) | ~edge(X8) [resolution 621,453]
1606. ~connects(X10,X9) | ~pp(X9,sK25) | ~pp(X8,sK25) | X8 = X9 | ~edge(X8) [resolution 716,453]
3220. sK1(sK22(sK25,sK26)) = sK10(sK22(sK25,sK26),sK24(sK25,sK26)) | ~edge(sK22(sK25,sK26)) [resolution 1328,655]
3225. sK1(sK22(sK25,sK26)) = sK10(sK22(sK25,sK26),sK24(sK25,sK26)) [subsumption resolution 3220,933]
3226. sK10(sK22(sK25,sK26),sK24(sK25,sK26)) = sK23(sK25,sK26) [forward demodulation 3225,1053]
4535. ~pp(sK24(sK25,sK26),sK25) | ~pp(X0,sK25) | sK24(sK25,sK26) = X0 | ~edge(X0) [resolution 1606,655]
4540. ~pp(X0,sK25) | sK24(sK25,sK26) = X0 | ~edge(X0) [subsumption resolution 4535,644]
9375. 457 <=> sK23(sK25,sK26) != sK26 [avatar definition]
9376. sK23(sK25,sK26) != sK26 <- {457} [avatar component clause 9375]
9378. 456 <=> sK23(sK25,sK26) = sK26 [avatar definition]
9379. sK23(sK25,sK26) = sK26 <- {456} [avatar component clause 9378]
16812. sK0(sK24(sK25,sK26)) = sK10(sK22(sK25,sK26),sK24(sK25,sK26)) | ~edge(sK24(sK25,sK26)) [resolution 655,1332]
28815. path(sK22(sK25,sK26),sK26,sK26) <- {456} [backward demodulation 9379,938]
28908. $false <- {456} [subsumption resolution 28815,413]
28909. ~456 [avatar contradiction clause 28908]
31075. sK22(sK25,sK26) = sK24(sK25,sK26) | ~edge(sK22(sK25,sK26)) [resolution 4540,639]
31110. sK22(sK25,sK26) = sK24(sK25,sK26) [subsumption resolution 31075,933]
31140. sK10(sK22(sK25,sK26),sK22(sK25,sK26)) = sK23(sK25,sK26) [backward demodulation 31110,3226]
32473. sK0(sK22(sK25,sK26)) = sK10(sK22(sK25,sK26),sK22(sK25,sK26)) | ~edge(sK24(sK25,sK26)) [forward demodulation 16812,31110]
32474. sK0(sK22(sK25,sK26)) = sK23(sK25,sK26) | ~edge(sK24(sK25,sK26)) [forward demodulation 32473,31140]
32475. sK23(sK25,sK26) = sK26 | ~edge(sK24(sK25,sK26)) [forward demodulation 32474,1037]
32476. ~edge(sK24(sK25,sK26)) <- {457} [subsumption resolution 32475,9376]
32477. ~edge(sK22(sK25,sK26)) <- {457} [forward demodulation 32476,31110]
32478. $false <- {457} [subsumption resolution 32477,933]
32479. 457 [avatar contradiction clause 32478]
32654. $false [avatar sat refutation 28909,32479]
% SZS output end Proof for 24apr2018_full_T08
% ------------------------------
% Version: Vampire 4.2.2 (commit c6486af3 on 2018-01-23 13:53:44 +0000)
% Termination reason: Refutation

% Memory used [KB]: 36332
% Time elapsed: 17.310 s
% ------------------------------
----  Runtime statistics ----
clauses created: 30518
clauses deleted: 7505
ssat_new_components: 652
ssat_non_splittable_sat_clauses: 107
ssat_reused_components: 557
ssat_sat_clauses: 680
total_frozen: 5705
total_unfrozen: 1441
unworthy child removed: 8866
-----------------------------
% ------------------------------
