% Refutation found. Thanks to Tanya!
% SZS status Theorem for 24apr2018_full_T11
% SZS output start Proof for 24apr2018_full_T11
54. ! [X142] : ! [X143] : (pi(X142,X143) <=> ? [X144] : ? [X145] : (X144 != X145 & cont(X145,X143) & cont(X145,X142) & cont(X144,X143) & cont(X144,X142) & node(X145) & node(X144))) [input sos540]
65. ! [X24] : ! [X25] : (pi(X24,X25) => pi(X25,X24)) [input t110]
66. ~! [X24] : ! [X25] : (pi(X24,X25) => pi(X25,X24)) [negated conjecture 65]
159. ! [X0] : ! [X1] : (pi(X0,X1) <=> ? [X2] : ? [X3] : (X2 != X3 & cont(X3,X1) & cont(X3,X0) & cont(X2,X1) & cont(X2,X0) & node(X3) & node(X2))) [rectify 54]
160. ! [X0,X1] : (pi(X0,X1) <=> ? [X2,X3] : (X2 != X3 & cont(X3,X1) & cont(X3,X0) & cont(X2,X1) & cont(X2,X0) & node(X3) & node(X2))) [flattening 159]
181. ~! [X0] : ! [X1] : (pi(X0,X1) => pi(X1,X0)) [rectify 66]
182. ~! [X0,X1] : (pi(X0,X1) => pi(X1,X0)) [flattening 181]
252. ? [X0,X1] : (~pi(X1,X0) & pi(X0,X1)) [ennf transformation 182]
289. ! [X0,X1] : ((pi(X0,X1) | ! [X2,X3] : (X2 = X3 | ~cont(X3,X1) | ~cont(X3,X0) | ~cont(X2,X1) | ~cont(X2,X0) | ~node(X3) | ~node(X2))) & (? [X2,X3] : (X2 != X3 & cont(X3,X1) & cont(X3,X0) & cont(X2,X1) & cont(X2,X0) & node(X3) & node(X2)) | ~pi(X0,X1))) [nnf transformation 160]
290. ! [X0,X1] : ((pi(X0,X1) | ! [X2,X3] : (X2 = X3 | ~cont(X3,X1) | ~cont(X3,X0) | ~cont(X2,X1) | ~cont(X2,X0) | ~node(X3) | ~node(X2))) & (? [X4,X5] : (X4 != X5 & cont(X5,X1) & cont(X5,X0) & cont(X4,X1) & cont(X4,X0) & node(X5) & node(X4)) | ~pi(X0,X1))) [rectify 289]
291. ! [X1,X0] : (? [X4,X5] : (X4 != X5 & cont(X5,X1) & cont(X5,X0) & cont(X4,X1) & cont(X4,X0) & node(X5) & node(X4)) => (sK12(X0,X1) != sK13(X0,X1) & cont(sK13(X0,X1),X1) & cont(sK13(X0,X1),X0) & cont(sK12(X0,X1),X1) & cont(sK12(X0,X1),X0) & node(sK13(X0,X1)) & node(sK12(X0,X1)))) [choice axiom]
292. ! [X0,X1] : ((pi(X0,X1) | ! [X2,X3] : (X2 = X3 | ~cont(X3,X1) | ~cont(X3,X0) | ~cont(X2,X1) | ~cont(X2,X0) | ~node(X3) | ~node(X2))) & ((sK12(X0,X1) != sK13(X0,X1) & cont(sK13(X0,X1),X1) & cont(sK13(X0,X1),X0) & cont(sK12(X0,X1),X1) & cont(sK12(X0,X1),X0) & node(sK13(X0,X1)) & node(sK12(X0,X1))) | ~pi(X0,X1))) [skolemisation 290,291]
305. ? [X0,X1] : (~pi(X1,X0) & pi(X0,X1)) => (~pi(sK26,sK25) & pi(sK25,sK26)) [choice axiom]
306. ~pi(sK26,sK25) & pi(sK25,sK26) [skolemisation 252,305]
380. node(sK12(X0,X1)) | ~pi(X0,X1) [cnf transformation 292]
381. node(sK13(X0,X1)) | ~pi(X0,X1) [cnf transformation 292]
382. cont(sK12(X0,X1),X0) | ~pi(X0,X1) [cnf transformation 292]
383. cont(sK12(X0,X1),X1) | ~pi(X0,X1) [cnf transformation 292]
384. cont(sK13(X0,X1),X0) | ~pi(X0,X1) [cnf transformation 292]
385. cont(sK13(X0,X1),X1) | ~pi(X0,X1) [cnf transformation 292]
386. sK12(X0,X1) != sK13(X0,X1) | ~pi(X0,X1) [cnf transformation 292]
387. ~cont(X3,X1) | X2 = X3 | pi(X0,X1) | ~cont(X3,X0) | ~cont(X2,X1) | ~cont(X2,X0) | ~node(X3) | ~node(X2) [cnf transformation 292]
409. pi(sK25,sK26) [cnf transformation 306]
410. ~pi(sK26,sK25) [cnf transformation 306]
609. sK13(X9,X10) = X8 | pi(X11,X9) | ~cont(sK13(X9,X10),X11) | ~cont(X8,X9) | ~cont(X8,X11) | ~node(sK13(X9,X10)) | ~node(X8) | ~pi(X9,X10) [resolution 387,384]
615. ~cont(sK13(X9,X10),X11) | pi(X11,X9) | sK13(X9,X10) = X8 | ~cont(X8,X9) | ~cont(X8,X11) | ~node(X8) | ~pi(X9,X10) [subsumption resolution 609,381]
1694. pi(X3,X4) | sK13(X4,X3) = X5 | ~cont(X5,X4) | ~cont(X5,X3) | ~node(X5) | ~pi(X4,X3) | ~pi(X4,X3) [resolution 615,385]
1695. ~pi(X4,X3) | sK13(X4,X3) = X5 | ~cont(X5,X4) | ~cont(X5,X3) | ~node(X5) | pi(X3,X4) [duplicate literal removal 1694]
4987. sK13(sK25,sK26) = X0 | ~cont(X0,sK25) | ~cont(X0,sK26) | ~node(X0) | pi(sK26,sK25) [resolution 1695,409]
4992. ~cont(X0,sK26) | ~cont(X0,sK25) | sK13(sK25,sK26) = X0 | ~node(X0) [subsumption resolution 4987,410]
39621. 415 <=> sK12(sK25,sK26) != sK13(sK25,sK26) [avatar definition]
39622. sK12(sK25,sK26) != sK13(sK25,sK26) <- {415} [avatar component clause 39621]
39624. 414 <=> sK12(sK25,sK26) = sK13(sK25,sK26) [avatar definition]
39625. sK12(sK25,sK26) = sK13(sK25,sK26) <- {414} [avatar component clause 39624]
46151. sK12(sK25,sK26) != sK12(sK25,sK26) | ~pi(sK25,sK26) <- {414} [superposition 386,39625]
46278. ~pi(sK25,sK26) <- {414} [trivial inequality removal 46151]
46279. $false <- {414} [subsumption resolution 46278,409]
46280. ~414 [avatar contradiction clause 46279]
46530. ~cont(sK12(X106,sK26),sK25) | sK12(X106,sK26) = sK13(sK25,sK26) | ~node(sK12(X106,sK26)) | ~pi(X106,sK26) [resolution 4992,383]
47384. ~cont(sK12(X106,sK26),sK25) | sK12(X106,sK26) = sK13(sK25,sK26) | ~pi(X106,sK26) [subsumption resolution 46530,380]
78923. sK12(sK25,sK26) = sK13(sK25,sK26) | ~pi(sK25,sK26) | ~pi(sK25,sK26) [resolution 47384,382]
78925. sK12(sK25,sK26) = sK13(sK25,sK26) | ~pi(sK25,sK26) [duplicate literal removal 78923]
78926. ~pi(sK25,sK26) <- {415} [subsumption resolution 78925,39622]
78927. $false <- {415} [subsumption resolution 78926,409]
78928. 415 [avatar contradiction clause 78927]
78932. $false [avatar sat refutation 46280,78928]
% SZS output end Proof for 24apr2018_full_T11
% ------------------------------
% Version: Vampire 4.2.2 (commit c6486af3 on 2018-01-23 13:53:44 +0000)
% Termination reason: Refutation

% Memory used [KB]: 42472
% Time elapsed: 85.754 s
% ------------------------------
----  Runtime statistics ----
clauses created: 77123
clauses deleted: 41806
ssat_new_components: 391
ssat_non_splittable_sat_clauses: 378
ssat_reused_components: 500
ssat_sat_clauses: 912
total_frozen: 24203
total_unfrozen: 7031
unworthy child removed: 23420
-----------------------------
% ------------------------------
