% Refutation found. Thanks to Tanya!
% SZS status Theorem for T11
% SZS output start Proof for T11
54. ! [X142] : ! [X143] : (pi(X142,X143) <=> ? [X144] : ? [X145] : (X144 != X145 & cont(X145,X143) & cont(X145,X142) & cont(X144,X143) & cont(X144,X142) & node(X145) & node(X144))) [input sos540]
64. ! [X24] : ! [X25] : (pi(X24,X25) => pi(X25,X24)) [input t110]
65. ~! [X24] : ! [X25] : (pi(X24,X25) => pi(X25,X24)) [negated conjecture 64]
158. ! [X0] : ! [X1] : (pi(X0,X1) <=> ? [X2] : ? [X3] : (X2 != X3 & cont(X3,X1) & cont(X3,X0) & cont(X2,X1) & cont(X2,X0) & node(X3) & node(X2))) [rectify 54]
159. ! [X0,X1] : (pi(X0,X1) <=> ? [X2,X3] : (X2 != X3 & cont(X3,X1) & cont(X3,X0) & cont(X2,X1) & cont(X2,X0) & node(X3) & node(X2))) [flattening 158]
178. ~! [X0] : ! [X1] : (pi(X0,X1) => pi(X1,X0)) [rectify 65]
179. ~! [X0,X1] : (pi(X0,X1) => pi(X1,X0)) [flattening 178]
246. ? [X0,X1] : (~pi(X1,X0) & pi(X0,X1)) [ennf transformation 179]
283. ! [X0,X1] : ((pi(X0,X1) | ! [X2,X3] : (X2 = X3 | ~cont(X3,X1) | ~cont(X3,X0) | ~cont(X2,X1) | ~cont(X2,X0) | ~node(X3) | ~node(X2))) & (? [X2,X3] : (X2 != X3 & cont(X3,X1) & cont(X3,X0) & cont(X2,X1) & cont(X2,X0) & node(X3) & node(X2)) | ~pi(X0,X1))) [nnf transformation 159]
284. ! [X0,X1] : ((pi(X0,X1) | ! [X2,X3] : (X2 = X3 | ~cont(X3,X1) | ~cont(X3,X0) | ~cont(X2,X1) | ~cont(X2,X0) | ~node(X3) | ~node(X2))) & (? [X4,X5] : (X4 != X5 & cont(X5,X1) & cont(X5,X0) & cont(X4,X1) & cont(X4,X0) & node(X5) & node(X4)) | ~pi(X0,X1))) [rectify 283]
285. ! [X1,X0] : (? [X4,X5] : (X4 != X5 & cont(X5,X1) & cont(X5,X0) & cont(X4,X1) & cont(X4,X0) & node(X5) & node(X4)) => (sK12(X0,X1) != sK13(X0,X1) & cont(sK13(X0,X1),X1) & cont(sK13(X0,X1),X0) & cont(sK12(X0,X1),X1) & cont(sK12(X0,X1),X0) & node(sK13(X0,X1)) & node(sK12(X0,X1)))) [choice axiom]
286. ! [X0,X1] : ((pi(X0,X1) | ! [X2,X3] : (X2 = X3 | ~cont(X3,X1) | ~cont(X3,X0) | ~cont(X2,X1) | ~cont(X2,X0) | ~node(X3) | ~node(X2))) & ((sK12(X0,X1) != sK13(X0,X1) & cont(sK13(X0,X1),X1) & cont(sK13(X0,X1),X0) & cont(sK12(X0,X1),X1) & cont(sK12(X0,X1),X0) & node(sK13(X0,X1)) & node(sK12(X0,X1))) | ~pi(X0,X1))) [skolemisation 284,285]
299. ? [X0,X1] : (~pi(X1,X0) & pi(X0,X1)) => (~pi(sK26,sK25) & pi(sK25,sK26)) [choice axiom]
300. ~pi(sK26,sK25) & pi(sK25,sK26) [skolemisation 246,299]
374. node(sK12(X0,X1)) | ~pi(X0,X1) [cnf transformation 286]
375. node(sK13(X0,X1)) | ~pi(X0,X1) [cnf transformation 286]
376. cont(sK12(X0,X1),X0) | ~pi(X0,X1) [cnf transformation 286]
377. cont(sK12(X0,X1),X1) | ~pi(X0,X1) [cnf transformation 286]
378. cont(sK13(X0,X1),X0) | ~pi(X0,X1) [cnf transformation 286]
379. cont(sK13(X0,X1),X1) | ~pi(X0,X1) [cnf transformation 286]
380. sK12(X0,X1) != sK13(X0,X1) | ~pi(X0,X1) [cnf transformation 286]
381. ~cont(X3,X1) | X2 = X3 | pi(X0,X1) | ~cont(X3,X0) | ~cont(X2,X1) | ~cont(X2,X0) | ~node(X3) | ~node(X2) [cnf transformation 286]
403. pi(sK25,sK26) [cnf transformation 300]
404. ~pi(sK26,sK25) [cnf transformation 300]
596. sK13(X9,X10) = X8 | pi(X11,X9) | ~cont(sK13(X9,X10),X11) | ~cont(X8,X9) | ~cont(X8,X11) | ~node(sK13(X9,X10)) | ~node(X8) | ~pi(X9,X10) [resolution 381,378]
602. ~cont(sK13(X9,X10),X11) | pi(X11,X9) | sK13(X9,X10) = X8 | ~cont(X8,X9) | ~cont(X8,X11) | ~node(X8) | ~pi(X9,X10) [subsumption resolution 596,375]
1624. pi(X3,X4) | sK13(X4,X3) = X5 | ~cont(X5,X4) | ~cont(X5,X3) | ~node(X5) | ~pi(X4,X3) | ~pi(X4,X3) [resolution 602,379]
1625. ~pi(X4,X3) | sK13(X4,X3) = X5 | ~cont(X5,X4) | ~cont(X5,X3) | ~node(X5) | pi(X3,X4) [duplicate literal removal 1624]
4832. sK13(sK25,sK26) = X0 | ~cont(X0,sK25) | ~cont(X0,sK26) | ~node(X0) | pi(sK26,sK25) [resolution 1625,403]
4837. ~cont(X0,sK26) | ~cont(X0,sK25) | sK13(sK25,sK26) = X0 | ~node(X0) [subsumption resolution 4832,404]
39278. 407 <=> sK12(sK25,sK26) != sK13(sK25,sK26) [avatar definition]
39279. sK12(sK25,sK26) != sK13(sK25,sK26) <- {407} [avatar component clause 39278]
39281. 406 <=> sK12(sK25,sK26) = sK13(sK25,sK26) [avatar definition]
39282. sK12(sK25,sK26) = sK13(sK25,sK26) <- {406} [avatar component clause 39281]
45800. sK12(sK25,sK26) != sK12(sK25,sK26) | ~pi(sK25,sK26) <- {406} [superposition 380,39282]
45927. ~pi(sK25,sK26) <- {406} [trivial inequality removal 45800]
45928. $false <- {406} [subsumption resolution 45927,403]
45929. ~406 [avatar contradiction clause 45928]
46545. ~cont(sK12(X106,sK26),sK25) | sK12(X106,sK26) = sK13(sK25,sK26) | ~node(sK12(X106,sK26)) | ~pi(X106,sK26) [resolution 4837,377]
46929. ~cont(sK12(X106,sK26),sK25) | sK12(X106,sK26) = sK13(sK25,sK26) | ~pi(X106,sK26) [subsumption resolution 46545,374]
80073. sK12(sK25,sK26) = sK13(sK25,sK26) | ~pi(sK25,sK26) | ~pi(sK25,sK26) [resolution 46929,376]
80075. sK12(sK25,sK26) = sK13(sK25,sK26) | ~pi(sK25,sK26) [duplicate literal removal 80073]
80076. ~pi(sK25,sK26) <- {407} [subsumption resolution 80075,39279]
80077. $false <- {407} [subsumption resolution 80076,403]
80078. 407 [avatar contradiction clause 80077]
80082. $false [avatar sat refutation 45929,80078]
% SZS output end Proof for T11
% ------------------------------
% Version: Vampire 4.2.2 (commit c6486af3 on 2018-01-23 13:53:44 +0000)
% Termination reason: Refutation

% Memory used [KB]: 42728
% Time elapsed: 90.529 s
% ------------------------------
----  Runtime statistics ----
clauses created: 78430
clauses deleted: 42245
ssat_new_components: 369
ssat_non_splittable_sat_clauses: 295
ssat_reused_components: 415
ssat_sat_clauses: 793
total_frozen: 20304
total_unfrozen: 2537
unworthy child removed: 25530
-----------------------------
% ------------------------------
