% Refutation found. Thanks to Tanya!
% SZS status Theorem for full_T12
% SZS output start Proof for full_T12
1. ! [X0] : ! [X1] : (edge(X0) => ~arc(X0,X1,X1)) [input sos10]
2. ! [X2] : ! [X3] : ! [X4] : (arc(X2,X3,X4) => (node(X4) & node(X3) & edge(X2))) [input sos20]
22. ! [X49] : ! [X50] : (pp(X49,X50) <=> (X49 != X50 & p(X49,X50))) [input sos220]
45. ! [X99] : ! [X100] : ! [X101] : (arc(X99,X100,X101) => path(X99,X100,X101)) [input sos460]
52. ! [X134] : ! [X135] : (connects(X134,X135) <=> ? [X136] : ? [X137] : ? [X138] : (path(X135,X137,X138) & path(X134,X136,X137))) [input sos530]
58. ! [X169] : ! [X170] : (connects(X169,X170) => (~p(X170,X169) & ~p(X169,X170))) [input sos595]
61. ! [X188] : ! [X189] : ! [X190] : ((~arc(X188,X189,X190) & path(X188,X189,X190)) => ? [X191] : ? [X192] : ? [X193] : (pp(X193,X188) & pp(X191,X188) & connects(X191,X193) & arc(X191,X189,X192))) [input sos630]
63. ! [X200] : ! [X27] : ! [X201] : (path(X200,X27,X201) => X27 != X201) [input t120]
64. ~! [X200] : ! [X27] : ! [X201] : (path(X200,X27,X201) => X27 != X201) [negated conjecture 63]
65. ! [X0,X1] : (edge(X0) => ~arc(X0,X1,X1)) [flattening 1]
66. ! [X0] : ! [X1] : ! [X2] : (arc(X0,X1,X2) => (node(X2) & node(X1) & edge(X0))) [rectify 2]
67. ! [X0,X1,X2] : (arc(X0,X1,X2) => (node(X2) & node(X1) & edge(X0))) [flattening 66]
99. ! [X0] : ! [X1] : (pp(X0,X1) <=> (X0 != X1 & p(X0,X1))) [rectify 22]
100. ! [X0,X1] : (pp(X0,X1) <=> (X0 != X1 & p(X0,X1))) [flattening 99]
139. ! [X0] : ! [X1] : ! [X2] : (arc(X0,X1,X2) => path(X0,X1,X2)) [rectify 45]
140. ! [X0,X1,X2] : (arc(X0,X1,X2) => path(X0,X1,X2)) [flattening 139]
153. ! [X0] : ! [X1] : (connects(X0,X1) <=> ? [X2] : ? [X3] : ? [X4] : (path(X1,X3,X4) & path(X0,X2,X3))) [rectify 52]
154. ! [X0,X1] : (connects(X0,X1) <=> ? [X2,X3,X4] : (path(X1,X3,X4) & path(X0,X2,X3))) [flattening 153]
165. ! [X0] : ! [X1] : (connects(X0,X1) => (~p(X1,X0) & ~p(X0,X1))) [rectify 58]
166. ! [X0,X1] : (connects(X0,X1) => (~p(X1,X0) & ~p(X0,X1))) [flattening 165]
171. ! [X0] : ! [X1] : ! [X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3] : ? [X4] : ? [X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4))) [rectify 61]
172. ! [X0,X1,X2] : ((~arc(X0,X1,X2) & path(X0,X1,X2)) => ? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4))) [flattening 171]
175. ~! [X0] : ! [X1] : ! [X2] : (path(X0,X1,X2) => X1 != X2) [rectify 64]
176. ~! [X0,X1,X2] : (path(X0,X1,X2) => X1 != X2) [flattening 175]
179. ! [X0,X1] : (~arc(X0,X1,X1) | ~edge(X0)) [ennf transformation 65]
180. ! [X0,X1,X2] : ((node(X2) & node(X1) & edge(X0)) | ~arc(X0,X1,X2)) [ennf transformation 67]
216. ! [X0,X1,X2] : (path(X0,X1,X2) | ~arc(X0,X1,X2)) [ennf transformation 140]
237. ! [X0,X1] : ((~p(X1,X0) & ~p(X0,X1)) | ~connects(X0,X1)) [ennf transformation 166]
242. ! [X0,X1,X2] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) | (arc(X0,X1,X2) | ~path(X0,X1,X2))) [ennf transformation 172]
243. ! [X0,X1,X2] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) | arc(X0,X1,X2) | ~path(X0,X1,X2)) [flattening 242]
244. ? [X0,X1,X2] : (X1 = X2 & path(X0,X1,X2)) [ennf transformation 176]
251. ! [X0,X1] : ((pp(X0,X1) | (X0 = X1 | ~p(X0,X1))) & ((X0 != X1 & p(X0,X1)) | ~pp(X0,X1))) [nnf transformation 100]
252. ! [X0,X1] : ((pp(X0,X1) | X0 = X1 | ~p(X0,X1)) & ((X0 != X1 & p(X0,X1)) | ~pp(X0,X1))) [flattening 251]
277. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & (? [X2,X3,X4] : (path(X1,X3,X4) & path(X0,X2,X3)) | ~connects(X0,X1))) [nnf transformation 154]
278. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & (? [X5,X6,X7] : (path(X1,X6,X7) & path(X0,X5,X6)) | ~connects(X0,X1))) [rectify 277]
279. ! [X1,X0] : (? [X5,X6,X7] : (path(X1,X6,X7) & path(X0,X5,X6)) => (path(X1,sK10(X0,X1),sK11(X0,X1)) & path(X0,sK9(X0,X1),sK10(X0,X1)))) [choice axiom]
280. ! [X0,X1] : ((connects(X0,X1) | ! [X2,X3,X4] : (~path(X1,X3,X4) | ~path(X0,X2,X3))) & ((path(X1,sK10(X0,X1),sK11(X0,X1)) & path(X0,sK9(X0,X1),sK10(X0,X1))) | ~connects(X0,X1))) [skolemisation 278,279]
293. ! [X1,X0] : (? [X3,X4,X5] : (pp(X5,X0) & pp(X3,X0) & connects(X3,X5) & arc(X3,X1,X4)) => (pp(sK24(X0,X1),X0) & pp(sK22(X0,X1),X0) & connects(sK22(X0,X1),sK24(X0,X1)) & arc(sK22(X0,X1),X1,sK23(X0,X1)))) [choice axiom]
294. ! [X0,X1,X2] : ((pp(sK24(X0,X1),X0) & pp(sK22(X0,X1),X0) & connects(sK22(X0,X1),sK24(X0,X1)) & arc(sK22(X0,X1),X1,sK23(X0,X1))) | arc(X0,X1,X2) | ~path(X0,X1,X2)) [skolemisation 243,293]
295. ? [X0,X1,X2] : (X1 = X2 & path(X0,X1,X2)) => (sK26 = sK27 & path(sK25,sK26,sK27)) [choice axiom]
296. sK26 = sK27 & path(sK25,sK26,sK27) [skolemisation 244,295]
297. ~arc(X0,X1,X1) | ~edge(X0) [cnf transformation 179]
298. ~arc(X0,X1,X2) | edge(X0) [cnf transformation 180]
321. ~pp(X0,X1) | p(X0,X1) [cnf transformation 252]
356. ~arc(X0,X1,X2) | path(X0,X1,X2) [cnf transformation 216]
368. ~path(X1,X3,X4) | connects(X0,X1) | ~path(X0,X2,X3) [cnf transformation 280]
388. ~connects(X0,X1) | ~p(X1,X0) [cnf transformation 237]
393. ~path(X0,X1,X2) | arc(X0,X1,X2) | arc(sK22(X0,X1),X1,sK23(X0,X1)) [cnf transformation 294]
395. ~path(X0,X1,X2) | arc(X0,X1,X2) | pp(sK22(X0,X1),X0) [cnf transformation 294]
397. path(sK25,sK26,sK27) [cnf transformation 296]
398. sK26 = sK27 [cnf transformation 296]
400. path(sK25,sK26,sK26) [forward demodulation 397,398]
622. arc(sK25,sK26,sK26) | pp(sK22(sK25,sK26),sK25) [resolution 395,400]
629. 20 <=> pp(sK22(sK25,sK26),sK25) [avatar definition]
630. pp(sK22(sK25,sK26),sK25) <- {20} [avatar component clause 629]
632. 23 <=> ~arc(sK25,sK26,sK26) [avatar definition]
633. ~arc(sK25,sK26,sK26) <- {23} [avatar component clause 632]
635. 22 <=> arc(sK25,sK26,sK26) [avatar definition]
636. arc(sK25,sK26,sK26) <- {22} [avatar component clause 635]
637. 20 | 22 [avatar split clause 622,635,629]
651. ~edge(sK25) <- {22} [resolution 636,297]
652. edge(sK25) <- {22} [resolution 636,298]
659. $false <- {22} [subsumption resolution 652,651]
660. ~22 [avatar contradiction clause 659]
690. arc(sK25,sK26,sK26) | arc(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) [resolution 393,400]
693. arc(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) <- {23} [subsumption resolution 690,633]
698. p(sK22(sK25,sK26),sK25) <- {20} [resolution 630,321]
836. path(sK22(sK25,sK26),sK26,sK23(sK25,sK26)) <- {23} [resolution 693,356]
2262. ~path(X14,X15,sK26) | connects(X14,sK22(sK25,sK26)) <- {23} [resolution 836,368]
21594. connects(sK25,sK22(sK25,sK26)) <- {23} [resolution 2262,400]
21650. ~p(sK22(sK25,sK26),sK25) <- {23} [resolution 21594,388]
21725. $false <- {20, 23} [subsumption resolution 21650,698]
21726. ~20 | 23 [avatar contradiction clause 21725]
21787. $false [avatar sat refutation 637,660,21726]
% SZS output end Proof for full_T12
% ------------------------------
% Version: Vampire 4.2.2 (commit c6486af3 on 2018-01-23 13:53:44 +0000)
% Termination reason: Refutation

% Memory used [KB]: 22643
% Time elapsed: 9.064 s
% ------------------------------
----  Runtime statistics ----
clauses created: 20916
clauses deleted: 3718
ssat_new_components: 202
ssat_non_splittable_sat_clauses: 36
ssat_reused_components: 98
ssat_sat_clauses: 222
total_frozen: 4757
total_unfrozen: 437
unworthy child removed: 1014
-----------------------------
% ------------------------------
